Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rich_list
% (18734)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1082
% Passive clauses: 8009
% Generated clauses: 34650
% Final active clauses: 1060
% Final passive clauses: 6671
% Input formulas: 261
% Initial clauses: 546
% 
% Duplicate literals: 191
% Trivial inequalities: 194
% Fw subsumption resolutions: 2221
% Fw demodulations: 3658
% Bw demodulations: 11
% 
% Simple tautologies: 348
% Equational tautologies: 28
% Forward subsumptions: 19538
% 
% Binary resolution: 6417
% Factoring: 479
% Forward superposition: 10792
% Backward superposition: 9929
% Self superposition: 37
% Equality factoring: 2
% Equality resolution: 25
% 
% Split clauses: 449
% Split components: 900
% SAT solver clauses: 604
% SAT solver unit clauses: 19
% SAT solver binary clauses: 74
% 
% Sat splits: 449
% 
% TWLsolver clauses: 603
% TWLsolver calls for satisfiability: 352
% 
% Memory used [KB]: 18549
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% (18773)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 575
% Passive clauses: 3453
% Generated clauses: 23855
% Final active clauses: 538
% Final passive clauses: 2583
% Input formulas: 261
% Initial clauses: 546
% 
% Duplicate literals: 94
% Trivial inequalities: 22
% Fw subsumption resolutions: 1041
% Fw demodulations: 3254
% Bw demodulations: 297
% Condensations: 64
% Global subsumptions: 25
% 
% Simple tautologies: 291
% Equational tautologies: 13
% Forward subsumptions: 11626
% Bw demodulations to eq. taut.: 13
% 
% Binary resolution: 2162
% Factoring: 45
% Forward superposition: 6990
% Backward superposition: 5423
% Self superposition: 48
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 30
% Split components: 60
% SAT solver clauses: 4238
% SAT solver unit clauses: 648
% SAT solver binary clauses: 175
% SAT solver learnt clauses: 51
% SAT solver learnt literals: 13
% 
% Sat splits: 30
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 4137
% TWLsolver calls for satisfiability: 3181
% 
% Memory used [KB]: 22771
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rich_list
% (18813)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1172
% Passive clauses: 14138
% Generated clauses: 167074
% Final active clauses: 882
% Final passive clauses: 4473
% Input formulas: 261
% Initial clauses: 546
% Discarded non-redundant clauses: 30229
% 
% Duplicate literals: 8
% Trivial inequalities: 39
% Fw subsumption resolutions: 1172
% Fw demodulations: 57669
% Bw demodulations: 3757
% Condensations: 52
% Global subsumptions: 8
% 
% Simple tautologies: 131
% Equational tautologies: 64
% Forward subsumptions: 39461
% Fw demodulations to eq. taut.: 32591
% Bw demodulations to eq. taut.: 17
% 
% Unit resulting resolution: 72808
% Factoring: 1
% Forward superposition: 16870
% Backward superposition: 9680
% Self superposition: 93
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 48
% Split components: 96
% SAT solver clauses: 14992
% SAT solver unit clauses: 4777
% SAT solver binary clauses: 7991
% 
% Sat splits: 48
% Sat splitting refutations: 748
% 
% Memory used [KB]: 180892
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f47518,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f991,f992,f994,f997,f999,f1000,f1001,f1004,f1095,f1116,f1121,f1187,f1188,f1357,f1190,f1359,f1361,f1193,f1195,f1196,f1197,f1200,f1291,f1312,f1317,f1485,f1606,f1607,f1608,f1609,f1617,f1622,f1633,f1638,f1657,f1662,f1673,f1678,f1780,f1779,f1778,f1785,f1804,f1805,f1809,f1811,f1813,f1815,f1816,f1911,f1931,f1936,f1995,f2000,f2118,f2117,f2281,f2280,f2279,f2274,f2262,f2254,f2253,f2252,f2397,f2398,f2399,f2400,f2401,f2402,f2417,f2419,f2421,f2423,f2424,f2519,f2539,f2544,f2601,f2606,f1649,f2617,f2622,f2638,f2643,f1689,f2654,f2659,f2693,f2698,f2712,f2720,f2725,f2736,f2741,f2775,f2780,f2794,f2802,f2807,f2818,f2823,f2842,f2847,f2858,f2863,f2882,f2887,f2898,f2903,f2922,f2927,f2938,f2943,f2962,f2967,f2978,f2983,f3002,f3007,f3018,f3023,f3043,f3048,f3059,f3064,f3084,f3089,f3100,f3105,f3129,f3131,f3133,f3135,f3136,f3231,f3251,f3256,f3551,f3766,f3767,f1161,f3769,f1163,f1165,f3772,f3774,f3775,f3776,f3779,f3870,f3891,f3896,f3942,f3947,f3957,f3994,f3999,f4030,f4035,f4051,f4052,f4056,f4058,f4060,f4062,f4063,f4158,f4178,f4183,f4243,f4245,f4247,f4249,f4250,f4345,f4365,f4370,f4429,f4434,f4445,f4450,f4467,f4472,f4483,f4488,f4505,f4510,f4521,f4526,f4543,f4548,f4559,f4564,f4581,f4586,f4597,f4602,f4619,f4624,f4635,f4640,f4657,f4662,f4673,f4678,f4696,f4701,f4712,f4717,f4735,f4740,f4751,f4756,f4778,f4780,f4782,f4784,f4785,f4880,f4900,f4905,f1174,f4966,f4968,f4970,f4972,f4973,f5068,f5137,f5179,f5181,f5183,f5185,f5186,f5281,f5540,f5543,f5545,f5372,f2009,f5556,f3965,f3960,f5473,f5382,f5379,f5378,f5377,f5375,f5559,f5803,f6025,f6030,f2010,f6050,f6052,f6053,f6054,f6057,f6148,f1166,f6174,f1176,f6264,f6300,f6316,f6317,f6321,f6323,f6325,f6327,f6328,f6423,f6443,f6448,f6508,f6510,f6512,f6514,f6515,f6610,f6630,f6635,f6694,f6699,f6715,f6732,f6737,f6753,f6770,f6775,f6791,f6808,f6813,f6829,f6846,f6851,f6867,f6884,f6889,f6905,f6922,f6927,f6943,f6961,f6966,f6982,f7000,f1175,f7005,f7021,f7043,f7045,f7047,f7049,f7050,f7145,f3296,f7230,f7232,f7234,f7236,f7237,f7332,f5350,f7420,f7425,f7441,f7459,f7464,f7480,f7498,f7503,f7519,f7537,f7542,f7558,f7576,f7581,f7597,f7615,f7620,f7636,f7654,f7659,f7675,f7694,f7699,f7715,f7734,f7739,f7755,f7778,f7780,f7782,f7784,f7785,f7880,f8198,f8454,f8456,f8457,f8458,f8461,f8552,f1362,f8578,f1366,f8650,f8686,f8702,f1976,f8703,f1978,f1979,f8707,f8709,f8711,f8713,f8714,f8809,f1980,f8893,f8895,f8897,f8899,f8900,f8995,f2586,f9020,f9079,f9084,f2751,f9100,f9117,f9122,f2833,f9138,f9155,f9160,f2873,f9176,f9193,f9198,f2914,f9214,f9231,f9236,f2954,f9252,f9269,f9274,f2994,f9290,f9307,f9312,f3035,f9328,f9346,f9351,f3076,f9367,f9385,f1365,f9390,f3117,f9406,f9428,f9430,f9432,f9434,f9435,f9530,f7945,f9615,f9617,f9619,f9621,f9622,f9717,f9808,f9982,f9980,f9811,f9813,f9814,f9815,f9818,f9909,f9983,f10022,f10024,f10026,f10028,f10029,f10124,f9786,f10217,f10233,f10244,f10257,f10273,f10284,f10297,f10313,f10325,f10338,f10354,f10365,f10378,f10394,f10405,f10417,f10433,f10445,f10457,f10473,f10486,f10498,f10514,f10527,f10539,f10555,f10568,f10569,f10582,f10598,f10609,f10622,f10638,f10649,f10662,f10678,f10690,f10703,f10719,f10730,f10743,f10759,f10770,f10782,f10798,f10810,f10822,f10838,f10851,f10863,f10879,f10892,f10904,f10920,f10933,f10964,f10977,f11008,f11021,f11033,f11049,f11061,f11092,f11105,f11136,f11150,f11180,f11194,f11224,f11238,f11268,f11283,f11295,f11311,f11323,f11353,f11368,f11398,f11413,f11426,f10608,f11442,f11452,f11465,f10648,f11481,f11491,f11504,f10689,f11520,f11531,f11545,f11561,f11573,f11586,f10729,f11602,f11612,f11625,f10769,f11641,f11651,f11663,f10809,f11679,f11690,f11702,f10850,f11718,f11730,f11742,f10891,f11758,f11770,f11784,f11800,f11812,f11824,f10932,f11840,f11852,f11865,f11909,f11953,f11998,f12042,f12085,f12129,f12174,f12219,f12265,f12308,f12351,f12395,f12488,f12490,f12449,f12492,f12543,f12585,f12628,f12672,f12716,f12774,f13167,f13392,f13426,f10243,f13464,f10283,f13502,f10324,f13541,f10364,f13579,f10404,f13617,f11060,f13633,f13644,f13656,f10444,f13694,f10485,f13733,f10526,f13772,f10567,f13812,f11322,f13828,f13839,f13854,f11572,f13870,f13883,f13896,f13939,f13984,f11811,f14000,f14013,f14026,f14070,f14113,f14155,f14198,f14242,f14286,f14342,f14383,f14413,f14806,f15031,f13412,f15064,f15080,f15096,f15110,f15126,f15142,f15158,f15174,f15186,f15202,f15218,f15230,f15255,f12491,f15926,f13166,f16000,f16016,f16034,f16051,f16067,f16086,f16101,f16117,f16133,f16148,f16164,f16180,f16197,f16213,f16227,f16244,f16260,f16274,f16299,f14382,f16972,f14805,f17046,f17062,f17080,f17097,f17113,f17132,f17157,f16132,f17173,f17188,f17204,f16179,f17220,f17235,f17253,f16226,f17269,f17284,f17302,f16273,f17318,f17333,f17348,f17364,f17379,f17404,f17799,f18068,f18084,f18101,f18125,f18165,f18201,f17079,f18217,f18234,f18249,f18265,f18277,f18294,f17131,f18310,f18328,f18345,f18398,f18449,f18465,f18481,f18500,f18557,f19380,f19396,f19413,f19430,f19487,f19543,f19559,f19576,f19592,f19608,f19621,f19637,f19653,f19671,f19687,f19703,f19716,f19734,f19787,f19842,f19897,f19913,f19931,f19951,f20777,f20833,f20892,f15095,f20939,f20955,f20977,f20996,f15141,f21049,f15185,f21107,f15229,f21159,f21175,f21197,f21972,f16033,f22023,f16085,f22078,f18378,f22133,f18431,f22190,f18535,f22255,f18592,f22315,f17378,f22331,f22349,f23047,f18100,f23063,f23084,f23189,f19467,f23243,f18276,f23259,f23272,f23289,f19525,f23345,f18480,f23361,f23381,f23399,f19410,f23415,f23435,f23442,f23443,f23474,f19573,f23490,f23510,f23656,f19619,f23672,f23685,f23760,f19669,f23776,f23798,f23815,f19714,f23831,f23844,f23923,f19929,f23939,f23961,f23976,f20974,f23992,f24013,f24028,f21194,f24044,f24065,f24163,f24179,f24208,f24229,f24245,f24271,f24288,f24304,f24329,f24348,f24364,f24379,f24398,f24414,f24441,f24460,f24476,f24503,f24518,f24534,f24550,f24565,f24581,f24597,f24613,f24629,f24657,f24677,f24693,f24713,f24733,f24749,f24769,f24785,f24801,f24829,f24847,f24863,f24895,f24916,f24932,f24958,f24976,f24992,f25018,f25039,f25055,f25073,f25092,f25108,f25135,f25154,f25170,f25197,f25215,f25231,f25251,f25269,f25285,f25305,f25323,f25339,f25372,f25393,f25409,f25429,f25450,f25466,f25486,f25504,f25520,f25553,f25561,f25626,f25702,f24892,f25718,f25752,f25773,f24955,f25789,f25814,f25833,f25016,f25849,f25875,f25898,f25071,f25914,f25934,f25953,f25133,f25969,f25995,f26014,f25195,f26030,f26056,f26077,f25249,f26093,f26116,f26137,f25303,f26153,f26176,f26196,f25369,f26212,f26249,f26271,f25426,f26287,f26306,f26328,f25483,f26344,f26363,f26383,f25550,f26399,f26436,f26456,f26533,f26597,f26669,f26731,f26796,f26866,f26937,f27006,f27068,f27070,f27071,f27072,f27075,f27166,f27251,f27176,f27202,f27261,f27287,f27348,f27408,f27470,f27472,f27473,f27474,f27477,f27568,f27653,f27578,f27257,f27604,f27661,f27677,f27679,f27680,f27681,f27684,f27775,f27872,f27814,f27877,f27901,f27917,f27932,f27953,f28032,f28096,f28170,f28235,f28251,f28283,f28302,f28366,f28439,f28516,f28590,f28679,f28741,f28802,f28892,f24205,f28969,f24268,f29029,f24327,f29101,f24377,f29164,f24439,f29223,f24501,f29294,f24548,f29373,f24595,f29448,f24654,f29526,f29528,f29529,f29530,f29533,f29624,f29721,f29663,f29726,f29758,f24710,f29815,f24766,f29873,f24826,f29955,f27930,f29971,f29985,f30008,f26509,f30093,f26574,f30159,f26641,f30239,f26709,f30310,f28281,f30326,f30357,f30376,f26774,f30440,f26839,f30519,f26910,f30610,f26981,f30696,f27064,f30797,f27322,f30861,f27383,f30926,f27466,f31023,f31039,f31061,f31081,f31097,f31119,f31143,f31159,f31181,f31201,f31217,f31239,f31293,f31335,f31384,f31419,f31467,f31502,f31591,f31628,f31647,f31625,f31683,f31703,f31719,f31746,f31766,f31744,f31782,f31808,f31817,f25683,f31881,f31948,f32014,f25618,f32079,f31940,f32154,f32180,f32204,f32221,f32247,f32271,f32289,f32202,f32315,f32338,f32355,f32269,f32381,f32404,f32418,f32444,f32467,f32480,f32506,f32529,f32542,f32568,f32591,f32604,f32630,f32654,f32666,f32692,f32716,f32728,f32754,f32778,f32791,f32817,f32842,f32854,f32880,f32905,f32910,f32967,f32918,f32944,f32971,f32984,f32465,f33010,f33032,f33045,f32527,f33071,f33093,f33106,f32589,f33132,f33154,f33168,f32652,f33194,f33217,f33229,f32714,f33255,f33278,f33290,f32776,f33316,f33339,f33352,f32840,f33378,f33402,f33414,f32903,f33440,f33464,f33470,f33527,f33478,f32969,f33504,f33530,f33547,f33576,f33608,f33588,f33594,f1652,f33662,f33664,f33666,f33668,f33670,f33672,f33674,f33695,f33724,f33756,f33736,f33742,f1690,f33774,f33606,f33803,f33834,f33815,f33821,f2633,f33851,f33754,f33880,f33911,f33892,f33898,f2670,f34083,f31179,f34099,f34120,f34140,f31237,f34156,f34177,f34186,f34460,f34197,f34462,f34202,f34464,f34466,f34685,f35208,f35254,f35272,f35333,f35400,f35446,f35481,f35545,f35571,f35617,f35638,f35706,f35751,f35752,f35772,f35838,f35876,f35896,f31059,f35951,f31117,f36009,f36049,f36050,f36079,f36125,f36144,f35310,f36204,f35371,f36269,f36315,f36345,f36398,f36423,f36476,f36506,f36553,f36554,f36569,f36598,f36629,f36610,f36616,f2752,f36643,f36672,f36704,f36684,f36690,f2834,f36718,f36747,f36778,f36759,f36765,f2874,f36793,f36822,f36854,f36834,f36840,f2915,f36866,f36895,f36927,f36907,f36913,f2955,f36966,f37048,f37049,f37061,f37090,f37122,f37102,f37108,f2995,f37135,f37164,f37197,f37176,f37182,f3036,f37209,f37238,f37271,f37250,f37256,f3077,f37276,f37343,f37284,f37313,f37347,f37325,f33609,f37331,f3118,f37359,f37361,f37363,f37365,f37366,f37461,f37561,f37483,f37568,f37614,f37696,f37697,f37718,f37748,f37765,f37811,f36627,f37840,f37870,f37858,f7451,f37884,f36702,f37913,f37944,f37931,f7490,f37964,f37980,f38013,f38027,f36776,f38056,f38086,f38074,f7529,f38101,f36852,f38130,f38161,f38148,f7569,f38173,f36925,f38202,f38233,f38220,f7608,f38245,f37120,f38274,f38305,f38292,f7647,f38318,f37195,f38347,f38379,f38365,f7687,f38402,f38418,f38454,f38466,f37269,f38495,f38527,f38513,f7727,f38532,f38599,f38540,f37345,f38569,f38602,f38587,f7767,f38619,f37796,f38662,f38678,f38690,f38717,f38763,f38775,f38791,f38803,f38816,f38832,f38844,f38856,f38872,f38884,f38912,f38961,f38990,f39043,f39056,f38842,f39072,f39083,f39113,f39167,f39181,f38882,f39197,f39210,f39223,f39268,f39314,f39358,f39419,f39435,f39479,f39492,f38688,f39531,f38801,f39571,f39251,f39615,f39298,f39674,f39721,f39692,f11893,f39750,f39766,f39809,f39836,f39883,f39854,f11937,f39912,f39960,f39930,f11982,f39990,f40038,f40008,f12026,f40068,f40115,f40086,f12070,f40123,f40197,f40124,f40198,f40150,f40200,f40168,f12114,f40232,f40248,f40282,f40312,f40361,f40330,f12159,f40390,f40439,f40408,f12204,f40468,f40517,f40486,f12249,f33610,f40541,f40572,f40593,f40569,f40623,f2123,f40644,f40660,f40685,f40706,f40683,f40722,f40749,f40770,f40801,f40823,f40798,f40853,f40802,f1178,f40875,f40907,f40891,f40910,f40921,f41211,f40926,f41213,f41215,f41980,f42012,f41996,f42015,f42053,f42071,f42151,f42196,f42470,f42207,f42472,f42212,f42474,f42476,f42695,f34961,f43263,f42185,f43326,f35251,f43397,f35443,f43472,f35541,f43559,f35614,f43632,f35746,f43704,f35871,f43763,f36044,f43827,f36122,f43898,f36312,f43968,f36395,f44045,f36473,f44128,f36548,f44221,f37045,f44344,f37693,f44448,f38011,f44523,f38452,f44601,f38760,f44674,f38958,f44756,f39040,f44858,f39164,f44958,f39475,f45047,f39805,f45136,f40280,f45195,f45197,f45199,f45201,f45202,f45297,f45417,f45306,f45424,f45377,f45430,f45461,f45502,f45477,f45504,f45543,f45619,f45559,f45621,f45648,f45707,f45664,f45709,f45739,f45795,f45755,f45797,f45824,f45867,f45840,f45869,f45896,f45939,f45912,f45941,f45973,f46044,f45989,f46048,f46068,f46104,f46084,f46108,f46140,f46211,f46156,f46215,f46240,f46290,f46256,f46292,f46317,f46367,f46333,f46369,f46397,f46438,f46413,f46440,f46475,f46534,f46491,f46536,f46565,f46609,f46581,f46611,f46641,f46688,f46657,f46690,f46722,f46763,f46738,f46765,f46799,f46853,f46815,f46857,f46886,f46921,f46902,f46925,f46959,f47013,f46975,f47017,f47051,f47108,f47067,f47110,f47139,f47189,f47155,f47191,f47218,f47261,f47234,f47263,f1989,f47274,f47276,f47278,f1167,f47280,f1169,f47281,f1168,f47376,f1170,f47498,f47401,f47508,f1651,f47510,f47512])).
fof(f47512,plain,(
  $false | ($spl44 | $spl760)),
  inference(subsumption_resolution,[],[f47511,f948])).
fof(f948,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rich_list.prefixes_is_prefix_total',aHL_TRUTH)).
fof(f47511,plain,(
  ~p(s(bool,'T_0')) | ($spl44 | $spl760)),
  inference(forward_demodulation,[],[f47446,f1288])).
fof(f1288,plain,(
  s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))),
  inference(resolution,[],[f653,f928])).
fof(f928,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f635])).
fof(f635,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X19] : (s(bool,'T_0') = s(bool,X19) <=> p(s(bool,X19)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.EQ_CLAUSES|split|1')).
fof(f653,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))))),
  inference(cnf_transformation,[],[f498])).
fof(f498,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) & ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) & ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f420])).
fof(f420,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1)))) & ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) & ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f419])).
fof(f419,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))) & (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) & ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))) => (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))) => (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),X7)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X7))))) => (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),X9)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),X7)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X7))))) => (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),X9)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X8))))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.prefixes_is_prefix_total')).
fof(f47446,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) | ($spl44 | $spl760)),
  inference(superposition,[],[f1976,f34464])).
fof(f47510,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f47509,f948])).
fof(f47509,plain,(
  ~p(s(bool,'T_0')) | $spl44),
  inference(forward_demodulation,[],[f47445,f1288])).
fof(f47445,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) | $spl44),
  inference(superposition,[],[f1976,f995])).
fof(f995,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK35(sK14,sK13,sK12))))),
  inference(resolution,[],[f652,f958])).
fof(f958,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),sK35(X2,X1,X0)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))) )),
  inference(cnf_transformation,[],[f645])).
fof(f645,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),sK35(X2,X1,X0))))) & (! [X3] : s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f644])).
fof(f644,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) | ? [X4] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4)))) & (! [X3] : s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f643])).
fof(f643,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) | ? [X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) & (! [X3] : s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(nnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) <=> ? [X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f389])).
fof(f389,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) <=> ? [X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X8,X9] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X8)))) <=> ? [X7] : s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X7))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_APPEND')).
fof(f652,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))))),
  inference(cnf_transformation,[],[f498])).
fof(f1651,plain,(
  s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) | $spl34),
  inference(cnf_transformation,[],[f1651_D])).
fof(f1651_D,plain,(
  s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f47508,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))))))) ) | $spl1082),
  inference(cnf_transformation,[],[f47508_D])).
fof(f47508_D,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))))))) ) <=> ~$spl1082),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1082])])).
fof(f47401,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19)))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl44),
  inference(superposition,[],[f1976,f977])).
fof(f977,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK38(X1,X0)),s('type.list.list'(X0),sK39(X1,X0)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f651])).
fof(f651,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK38(X1,X0)),s('type.list.list'(X0),sK39(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK38,sK39])],[f418])).
fof(f418,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f417])).
fof(f417,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6] : ! [X7] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X7) | ? [X20,X19] : s('type.list.list'(X6),X7) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X20),s('type.list.list'(X6),X19))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.list.list_CASES')).
fof(f47498,plain,(
  ( ! [X23,X25] : (~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X25)))))) ) | $spl1080),
  inference(cnf_transformation,[],[f47498_D])).
fof(f47498_D,plain,(
  ( ! [X23,X25] : (~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X25)))))) ) <=> ~$spl1080),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1080])])).
fof(f1170,plain,(
  ( ! [X21] : (p(s(bool,X21))) ) | $spl14),
  inference(cnf_transformation,[],[f1170_D])).
fof(f1170_D,plain,(
  ( ! [X21] : (p(s(bool,X21))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f47376,plain,(
  ( ! [X282,X281,X283] : (p(s(bool,X281)) | ~p(s(bool,X282)) | s(bool,X282) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X283)))))) ) | $spl44),
  inference(resolution,[],[f1976,f978])).
fof(f978,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f676])).
fof(f676,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f500])).
fof(f500,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f499])).
fof(f499,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f425])).
fof(f425,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X11,X31,X18,X32] : ((s(bool,X11) = s(bool,X31) & (p(s(bool,X31)) => s(bool,X18) = s(bool,X32))) => ((p(s(bool,X11)) => p(s(bool,X18))) <=> (p(s(bool,X31)) => p(s(bool,X32)))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.IMP_CONG')).
fof(f1168,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X16))) ) | $spl10),
  inference(cnf_transformation,[],[f1168_D])).
fof(f1168_D,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X16))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f47281,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X51))))) | s(bool,X49) != s(bool,X52)) ) | $spl44),
  inference(resolution,[],[f1976,f679])).
fof(f679,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f500])).
fof(f1169,plain,(
  ( ! [X22] : (~p(s(bool,X22))) ) | $spl12),
  inference(cnf_transformation,[],[f1169_D])).
fof(f1169_D,plain,(
  ( ! [X22] : (~p(s(bool,X22))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f47280,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X48)))))) ) | $spl44),
  inference(resolution,[],[f1976,f678])).
fof(f678,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f500])).
fof(f1167,plain,(
  ( ! [X11,X9] : (s(bool,X9) != s(bool,X11) | p(s(bool,X9))) ) | $spl8),
  inference(cnf_transformation,[],[f1167_D])).
fof(f1167_D,plain,(
  ( ! [X11,X9] : (s(bool,X9) != s(bool,X11) | p(s(bool,X9))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f47278,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X41)))))) ) | $spl44),
  inference(resolution,[],[f1976,f677])).
fof(f677,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f500])).
fof(f47276,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X32))))) | s(bool,X30) != s(bool,X33)) ) | $spl44),
  inference(resolution,[],[f1976,f675])).
fof(f675,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f500])).
fof(f47274,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X25)))))) ) | $spl44),
  inference(resolution,[],[f1976,f673])).
fof(f673,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f500])).
fof(f1989,plain,(
  s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12)))) | $spl52),
  inference(cnf_transformation,[],[f1989_D])).
fof(f1989_D,plain,(
  s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f47263,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl1078),
  inference(cnf_transformation,[],[f47263_D])).
fof(f47263_D,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl1078),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1078])])).
fof(f47234,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1948,f977])).
fof(f1948,plain,(
  ( ! [X6,X7] : (~sP9(X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f654,f752])).
fof(f752,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f493])).
fof(f493,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f275,f492,f491,f490,f489])).
fof(f489,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f490,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f491,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f492,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X35,X33,X34] : ((p(s(bool,X34)) <=> s(bool,X33) = s(bool,X35)) <=> ((p(s(bool,X34)) | p(s(bool,X33)) | p(s(bool,X35))) & (p(s(bool,X34)) | ~p(s(bool,X35)) | ~p(s(bool,X33))) & (p(s(bool,X33)) | ~p(s(bool,X35)) | ~p(s(bool,X34))) & (p(s(bool,X35)) | ~p(s(bool,X33)) | ~p(s(bool,X34)))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.sat.sat_2|49|_prove_ltb')).
fof(f654,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))))),
  inference(cnf_transformation,[],[f498])).
fof(f47261,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl1076),
  inference(cnf_transformation,[],[f47261_D])).
fof(f47261_D,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl1076),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1076])])).
fof(f47218,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1948,f977])).
fof(f47191,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1074),
  inference(cnf_transformation,[],[f47191_D])).
fof(f47191_D,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1074),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1074])])).
fof(f47155,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1946,f977])).
fof(f1946,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X2) | ~p(s(bool,X2)) | p(s(bool,X3))) )),
  inference(superposition,[],[f654,f750])).
fof(f750,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f535])).
fof(f47189,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1072),
  inference(cnf_transformation,[],[f47189_D])).
fof(f47189_D,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1072),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1072])])).
fof(f47139,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1946,f977])).
fof(f47110,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1070),
  inference(cnf_transformation,[],[f47110_D])).
fof(f47110_D,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1070),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1070])])).
fof(f47067,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1945,f977])).
fof(f1945,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(superposition,[],[f654,f750])).
fof(f47108,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1068),
  inference(cnf_transformation,[],[f47108_D])).
fof(f47108_D,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1068),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1068])])).
fof(f47051,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1945,f977])).
fof(f47017,plain,(
  ( ! [X26,X27] : (~sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1066),
  inference(cnf_transformation,[],[f47017_D])).
fof(f47017_D,plain,(
  ( ! [X26,X27] : (~sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1066),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1066])])).
fof(f46975,plain,(
  ( ! [X26,X27] : (~sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1873,f977])).
fof(f1873,plain,(
  ( ! [X149,X148] : (~sP6(X148,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X149) | p(s(bool,X149)) | p(s(bool,X148))) )),
  inference(resolution,[],[f654,f746])).
fof(f746,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f534])).
fof(f534,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f533])).
fof(f533,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f489])).
fof(f47013,plain,(
  ( ! [X26,X27] : (~sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1064),
  inference(cnf_transformation,[],[f47013_D])).
fof(f47013_D,plain,(
  ( ! [X26,X27] : (~sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1064),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1064])])).
fof(f46959,plain,(
  ( ! [X26,X27] : (~sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1873,f977])).
fof(f46925,plain,(
  ( ! [X26,X27] : (~sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1062),
  inference(cnf_transformation,[],[f46925_D])).
fof(f46925_D,plain,(
  ( ! [X26,X27] : (~sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1062),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1062])])).
fof(f46902,plain,(
  ( ! [X26,X27] : (~sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1872,f977])).
fof(f1872,plain,(
  ( ! [X146,X147] : (~sP6(X146,X147,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | p(s(bool,X147)) | p(s(bool,X146))) )),
  inference(resolution,[],[f654,f746])).
fof(f46921,plain,(
  ( ! [X26,X27] : (~sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1060),
  inference(cnf_transformation,[],[f46921_D])).
fof(f46921_D,plain,(
  ( ! [X26,X27] : (~sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1060),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1060])])).
fof(f46886,plain,(
  ( ! [X26,X27] : (~sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1872,f977])).
fof(f46857,plain,(
  ( ! [X26,X27] : (~sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1058),
  inference(cnf_transformation,[],[f46857_D])).
fof(f46857_D,plain,(
  ( ! [X26,X27] : (~sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1058),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1058])])).
fof(f46815,plain,(
  ( ! [X26,X27] : (~sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1871,f977])).
fof(f1871,plain,(
  ( ! [X144,X145] : (~sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X144,X145) | p(s(bool,X145)) | p(s(bool,X144))) )),
  inference(resolution,[],[f654,f746])).
fof(f46853,plain,(
  ( ! [X26,X27] : (~sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1056),
  inference(cnf_transformation,[],[f46853_D])).
fof(f46853_D,plain,(
  ( ! [X26,X27] : (~sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1056),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1056])])).
fof(f46799,plain,(
  ( ! [X26,X27] : (~sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1871,f977])).
fof(f46765,plain,(
  ( ! [X26,X27] : (~sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1054),
  inference(cnf_transformation,[],[f46765_D])).
fof(f46765_D,plain,(
  ( ! [X26,X27] : (~sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1054),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1054])])).
fof(f46738,plain,(
  ( ! [X26,X27] : (~sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1868,f977])).
fof(f1868,plain,(
  ( ! [X138,X139] : (~sP7(X138,X139,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X139)) | ~p(s(bool,X138))) )),
  inference(resolution,[],[f654,f742])).
fof(f742,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f532])).
fof(f532,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f531])).
fof(f531,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f530])).
fof(f530,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f490])).
fof(f46763,plain,(
  ( ! [X26,X27] : (~sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1052),
  inference(cnf_transformation,[],[f46763_D])).
fof(f46763_D,plain,(
  ( ! [X26,X27] : (~sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1052),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1052])])).
fof(f46722,plain,(
  ( ! [X26,X27] : (~sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1868,f977])).
fof(f46690,plain,(
  ( ! [X26,X27] : (~sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1050),
  inference(cnf_transformation,[],[f46690_D])).
fof(f46690_D,plain,(
  ( ! [X26,X27] : (~sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1050),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1050])])).
fof(f46657,plain,(
  ( ! [X26,X27] : (~sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1865,f977])).
fof(f1865,plain,(
  ( ! [X132,X133] : (~sP8(X132,X133,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X133)) | ~p(s(bool,X132))) )),
  inference(resolution,[],[f654,f738])).
fof(f738,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f529])).
fof(f529,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f528])).
fof(f528,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f527])).
fof(f527,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f491])).
fof(f46688,plain,(
  ( ! [X26,X27] : (~sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1048),
  inference(cnf_transformation,[],[f46688_D])).
fof(f46688_D,plain,(
  ( ! [X26,X27] : (~sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1048),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1048])])).
fof(f46641,plain,(
  ( ! [X26,X27] : (~sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1865,f977])).
fof(f46611,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1046),
  inference(cnf_transformation,[],[f46611_D])).
fof(f46611_D,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1046),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1046])])).
fof(f46581,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1862,f977])).
fof(f1862,plain,(
  ( ! [X127,X126] : (~sP9(X126,X127,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X127)) | ~p(s(bool,X126))) )),
  inference(resolution,[],[f654,f734])).
fof(f734,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f525])).
fof(f525,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f524])).
fof(f524,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f492])).
fof(f46609,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1044),
  inference(cnf_transformation,[],[f46609_D])).
fof(f46609_D,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1044),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1044])])).
fof(f46565,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1862,f977])).
fof(f46536,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1042),
  inference(cnf_transformation,[],[f46536_D])).
fof(f46536_D,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1042),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1042])])).
fof(f46491,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1861,f977])).
fof(f1861,plain,(
  ( ! [X125,X124] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X124,X125) | ~p(s(bool,X125)) | ~p(s(bool,X124))) )),
  inference(resolution,[],[f654,f730])).
fof(f730,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f523])).
fof(f523,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f522])).
fof(f522,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f488])).
fof(f488,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f273,f487,f486])).
fof(f486,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f487,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X35,X33,X34] : ((p(s(bool,X34)) <=> (p(s(bool,X33)) & p(s(bool,X35)))) <=> ((p(s(bool,X34)) | ~p(s(bool,X33)) | ~p(s(bool,X35))) & (p(s(bool,X33)) | ~p(s(bool,X34))) & (p(s(bool,X35)) | ~p(s(bool,X34)))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.sat.sat_23_prove_ltb')).
fof(f46534,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1040),
  inference(cnf_transformation,[],[f46534_D])).
fof(f46534_D,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1040),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1040])])).
fof(f46475,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1861,f977])).
fof(f46440,plain,(
  ( ! [X26,X27] : (~sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1038),
  inference(cnf_transformation,[],[f46440_D])).
fof(f46440_D,plain,(
  ( ! [X26,X27] : (~sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1038),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1038])])).
fof(f46413,plain,(
  ( ! [X26,X27] : (~sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1852,f977])).
fof(f1852,plain,(
  ( ! [X107,X106] : (~sP4(X106,X107,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X107)) | ~p(s(bool,X106))) )),
  inference(resolution,[],[f654,f721])).
fof(f721,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f521])).
fof(f521,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f520])).
fof(f520,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f486])).
fof(f46438,plain,(
  ( ! [X26,X27] : (~sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl1036),
  inference(cnf_transformation,[],[f46438_D])).
fof(f46438_D,plain,(
  ( ! [X26,X27] : (~sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl1036),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1036])])).
fof(f46397,plain,(
  ( ! [X26,X27] : (~sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1852,f977])).
fof(f46369,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl1034),
  inference(cnf_transformation,[],[f46369_D])).
fof(f46369_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl1034),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1034])])).
fof(f46333,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1845,f977])).
fof(f1845,plain,(
  ( ! [X92,X93] : (~sP3(X93,X92,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X93)) | p(s(bool,X92))) )),
  inference(resolution,[],[f654,f711])).
fof(f711,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP3(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f516])).
fof(f516,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f515])).
fof(f515,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f485])).
fof(f485,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f271,f484,f483])).
fof(f483,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f484,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X35,X33,X34] : ((p(s(bool,X34)) <=> (p(s(bool,X33)) | p(s(bool,X35)))) <=> ((p(s(bool,X34)) | ~p(s(bool,X33))) & (p(s(bool,X34)) | ~p(s(bool,X35))) & (p(s(bool,X33)) | p(s(bool,X35)) | ~p(s(bool,X34)))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.sat.sat_25_prove_ltb')).
fof(f46367,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl1032),
  inference(cnf_transformation,[],[f46367_D])).
fof(f46367_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl1032),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1032])])).
fof(f46317,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1845,f977])).
fof(f46292,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl1030),
  inference(cnf_transformation,[],[f46292_D])).
fof(f46292_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl1030),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1030])])).
fof(f46256,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1844,f977])).
fof(f1844,plain,(
  ( ! [X90,X91] : (~sP3(X91,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X90) | ~p(s(bool,X91)) | p(s(bool,X90))) )),
  inference(resolution,[],[f654,f711])).
fof(f46290,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl1028),
  inference(cnf_transformation,[],[f46290_D])).
fof(f46290_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl1028),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1028])])).
fof(f46240,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1844,f977])).
fof(f46215,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1026),
  inference(cnf_transformation,[],[f46215_D])).
fof(f46215_D,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1026),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1026])])).
fof(f46156,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1843,f977])).
fof(f1843,plain,(
  ( ! [X88,X89] : (sP3(X88,X89,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | p(s(bool,X89)) | p(s(bool,X88))) )),
  inference(resolution,[],[f654,f708])).
fof(f708,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f516])).
fof(f46211,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1024),
  inference(cnf_transformation,[],[f46211_D])).
fof(f46211_D,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1024),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1024])])).
fof(f46140,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1843,f977])).
fof(f46108,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | p(s(bool,X26))) ) | $spl1022),
  inference(cnf_transformation,[],[f46108_D])).
fof(f46108_D,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | p(s(bool,X26))) ) <=> ~$spl1022),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1022])])).
fof(f46084,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1842,f977])).
fof(f1842,plain,(
  ( ! [X87,X86] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X86,X87) | p(s(bool,X86)) | p(s(bool,X87))) )),
  inference(resolution,[],[f654,f708])).
fof(f46104,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27)) | p(s(bool,X26))) ) | $spl1020),
  inference(cnf_transformation,[],[f46104_D])).
fof(f46104_D,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27)) | p(s(bool,X26))) ) <=> ~$spl1020),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1020])])).
fof(f46068,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1842,f977])).
fof(f46048,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1018),
  inference(cnf_transformation,[],[f46048_D])).
fof(f46048_D,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1018),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1018])])).
fof(f45989,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1841,f977])).
fof(f1841,plain,(
  ( ! [X85,X84] : (sP3(X84,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X85) | p(s(bool,X85)) | p(s(bool,X84))) )),
  inference(resolution,[],[f654,f708])).
fof(f46044,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1016),
  inference(cnf_transformation,[],[f46044_D])).
fof(f46044_D,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1016),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1016])])).
fof(f45973,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1841,f977])).
fof(f45941,plain,(
  ( ! [X26,X27] : (~sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1014),
  inference(cnf_transformation,[],[f45941_D])).
fof(f45941_D,plain,(
  ( ! [X26,X27] : (~sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1014),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1014])])).
fof(f45912,plain,(
  ( ! [X26,X27] : (~sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1839,f977])).
fof(f1839,plain,(
  ( ! [X80,X81] : (~sP2(X80,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X81) | p(s(bool,X81)) | ~p(s(bool,X80))) )),
  inference(resolution,[],[f654,f704])).
fof(f704,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f514])).
fof(f514,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f513])).
fof(f513,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f512])).
fof(f512,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f483])).
fof(f45939,plain,(
  ( ! [X26,X27] : (~sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1012),
  inference(cnf_transformation,[],[f45939_D])).
fof(f45939_D,plain,(
  ( ! [X26,X27] : (~sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1012),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1012])])).
fof(f45896,plain,(
  ( ! [X26,X27] : (~sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1839,f977])).
fof(f45869,plain,(
  ( ! [X26,X27] : (~sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1010),
  inference(cnf_transformation,[],[f45869_D])).
fof(f45869_D,plain,(
  ( ! [X26,X27] : (~sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1010),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1010])])).
fof(f45840,plain,(
  ( ! [X26,X27] : (~sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1838,f977])).
fof(f1838,plain,(
  ( ! [X78,X79] : (~sP2(X78,X79,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | p(s(bool,X79)) | ~p(s(bool,X78))) )),
  inference(resolution,[],[f654,f704])).
fof(f45867,plain,(
  ( ! [X26,X27] : (~sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl1008),
  inference(cnf_transformation,[],[f45867_D])).
fof(f45867_D,plain,(
  ( ! [X26,X27] : (~sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl1008),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1008])])).
fof(f45824,plain,(
  ( ! [X26,X27] : (~sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1838,f977])).
fof(f45797,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl1006),
  inference(cnf_transformation,[],[f45797_D])).
fof(f45797_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl1006),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1006])])).
fof(f45755,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1828,f977])).
fof(f1828,plain,(
  ( ! [X59,X58] : (~sP1(X59,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X58) | ~p(s(bool,X59)) | ~p(s(bool,X58))) )),
  inference(resolution,[],[f654,f694])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f508])).
fof(f508,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f507])).
fof(f507,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f427,f481,f480])).
fof(f480,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f481,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f427,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X35,X33,X34] : ((p(s(bool,X34)) <=> (p(s(bool,X33)) => p(s(bool,X35)))) <=> ((p(s(bool,X34)) | p(s(bool,X33))) & (p(s(bool,X34)) | ~p(s(bool,X35))) & (~p(s(bool,X33)) | p(s(bool,X35)) | ~p(s(bool,X34)))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.sat.sat_27_prove_ltb')).
fof(f45795,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl1004),
  inference(cnf_transformation,[],[f45795_D])).
fof(f45795_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl1004),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1004])])).
fof(f45739,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1828,f977])).
fof(f45709,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26))) ) | $spl1002),
  inference(cnf_transformation,[],[f45709_D])).
fof(f45709_D,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26))) ) <=> ~$spl1002),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1002])])).
fof(f45664,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1826,f977])).
fof(f1826,plain,(
  ( ! [X54,X55] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X54,X55) | p(s(bool,X54)) | ~p(s(bool,X55))) )),
  inference(resolution,[],[f654,f691])).
fof(f691,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f508])).
fof(f45707,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26))) ) | $spl1000),
  inference(cnf_transformation,[],[f45707_D])).
fof(f45707_D,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26))) ) <=> ~$spl1000),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1000])])).
fof(f45648,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1826,f977])).
fof(f45621,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl998),
  inference(cnf_transformation,[],[f45621_D])).
fof(f45621_D,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl998),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl998])])).
fof(f45559,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1825,f977])).
fof(f1825,plain,(
  ( ! [X52,X53] : (sP1(X52,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X53) | ~p(s(bool,X53)) | p(s(bool,X52))) )),
  inference(resolution,[],[f654,f691])).
fof(f45619,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl996),
  inference(cnf_transformation,[],[f45619_D])).
fof(f45619_D,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl996),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl996])])).
fof(f45543,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1825,f977])).
fof(f45504,plain,(
  ( ! [X26,X27] : (~sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl994),
  inference(cnf_transformation,[],[f45504_D])).
fof(f45504_D,plain,(
  ( ! [X26,X27] : (~sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl994),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl994])])).
fof(f45477,plain,(
  ( ! [X26,X27] : (~sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1822,f977])).
fof(f1822,plain,(
  ( ! [X47,X46] : (~sP0(X46,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X47) | ~p(s(bool,X47)) | ~p(s(bool,X46))) )),
  inference(resolution,[],[f654,f687])).
fof(f687,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f506])).
fof(f506,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f505])).
fof(f505,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f504])).
fof(f504,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f480])).
fof(f45502,plain,(
  ( ! [X26,X27] : (~sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl992),
  inference(cnf_transformation,[],[f45502_D])).
fof(f45502_D,plain,(
  ( ! [X26,X27] : (~sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl992),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl992])])).
fof(f45461,plain,(
  ( ! [X26,X27] : (~sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1822,f977])).
fof(f45430,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X19))),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) | $spl990),
  inference(cnf_transformation,[],[f45430_D])).
fof(f45430_D,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X19))),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) <=> ~$spl990),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl990])])).
fof(f45377,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X19))),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1806,f977])).
fof(f1806,plain,(
  ( ! [X10] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X10))),s('type.list.list'(sK12),sK15))))) )),
  inference(resolution,[],[f654,f953])).
fof(f953,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f478])).
fof(f478,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3)))))),
  inference(ennf_transformation,[],[f384])).
fof(f384,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) => p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3)))))),
  inference(flattening,[],[f383])).
fof(f383,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) => p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6] : ! [X24,X25,X26] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X24),s('type.list.list'(X6),X25))),s('type.list.list'(X6),X26)))) => p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X24),s('type.list.list'(X6),X26)))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_APPEND|49|')).
fof(f45424,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))),s('type.list.list'(sK12),sK15))))) ) | $spl988),
  inference(cnf_transformation,[],[f45424_D])).
fof(f45424_D,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))),s('type.list.list'(sK12),sK15))))) ) <=> ~$spl988),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl988])])).
fof(f45306,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))),s('type.list.list'(sK12),sK15)))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1806,f977])).
fof(f45417,plain,(
  ( ! [X23,X25] : (~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X25))),s('type.list.list'(sK12),sK15)))) ) | $spl986),
  inference(cnf_transformation,[],[f45417_D])).
fof(f45417_D,plain,(
  ( ! [X23,X25] : (~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X25))),s('type.list.list'(sK12),sK15)))) ) <=> ~$spl986),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl986])])).
fof(f45297,plain,(
  ( ! [X282,X281,X283] : (p(s(bool,X281)) | ~p(s(bool,X282)) | s(bool,X282) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X283))),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f1806,f978])).
fof(f45202,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X51))),s('type.list.list'(sK12),sK15))) | s(bool,X49) != s(bool,X52)) )),
  inference(resolution,[],[f1806,f679])).
fof(f45201,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X48))),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f1806,f678])).
fof(f45199,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X41))),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f1806,f677])).
fof(f45197,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X32))),s('type.list.list'(sK12),sK15))) | s(bool,X30) != s(bool,X33)) )),
  inference(resolution,[],[f1806,f675])).
fof(f45195,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X25))),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f1806,f673])).
fof(f40280,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl938),
  inference(cnf_transformation,[],[f40280_D])).
fof(f40280_D,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl938),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl938])])).
fof(f45136,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1328,f977])).
fof(f1328,plain,(
  ( ! [X4,X5] : (~sP9(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(superposition,[],[f653,f752])).
fof(f39805,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl920),
  inference(cnf_transformation,[],[f39805_D])).
fof(f39805_D,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl920),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl920])])).
fof(f45047,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1327,f977])).
fof(f1327,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X2) | p(s(bool,X2)) | p(s(bool,X3))) )),
  inference(superposition,[],[f653,f750])).
fof(f39475,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl914),
  inference(cnf_transformation,[],[f39475_D])).
fof(f39475_D,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl914),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl914])])).
fof(f44958,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1326,f977])).
fof(f1326,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(superposition,[],[f653,f750])).
fof(f39164,plain,(
  ( ! [X26,X27] : (~sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl906),
  inference(cnf_transformation,[],[f39164_D])).
fof(f39164_D,plain,(
  ( ! [X26,X27] : (~sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl906),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl906])])).
fof(f44858,plain,(
  ( ! [X26,X27] : (~sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1253,f977])).
fof(f1253,plain,(
  ( ! [X138,X137] : (~sP7(X137,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X138) | p(s(bool,X138)) | ~p(s(bool,X137))) )),
  inference(resolution,[],[f653,f742])).
fof(f39040,plain,(
  ( ! [X26,X27] : (~sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl902),
  inference(cnf_transformation,[],[f39040_D])).
fof(f39040_D,plain,(
  ( ! [X26,X27] : (~sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl902),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl902])])).
fof(f44756,plain,(
  ( ! [X26,X27] : (~sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1252,f977])).
fof(f1252,plain,(
  ( ! [X136,X135] : (~sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X135,X136) | p(s(bool,X136)) | ~p(s(bool,X135))) )),
  inference(resolution,[],[f653,f742])).
fof(f38958,plain,(
  ( ! [X26,X27] : (~sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl900),
  inference(cnf_transformation,[],[f38958_D])).
fof(f38958_D,plain,(
  ( ! [X26,X27] : (~sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl900),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl900])])).
fof(f44674,plain,(
  ( ! [X26,X27] : (~sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1250,f977])).
fof(f1250,plain,(
  ( ! [X132,X131] : (~sP8(X131,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X132) | p(s(bool,X132)) | ~p(s(bool,X131))) )),
  inference(resolution,[],[f653,f738])).
fof(f38760,plain,(
  ( ! [X26,X27] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl886),
  inference(cnf_transformation,[],[f38760_D])).
fof(f38760_D,plain,(
  ( ! [X26,X27] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl886),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl886])])).
fof(f44601,plain,(
  ( ! [X26,X27] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1249,f977])).
fof(f1249,plain,(
  ( ! [X130,X129] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X129,X130) | p(s(bool,X130)) | ~p(s(bool,X129))) )),
  inference(resolution,[],[f653,f738])).
fof(f38452,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl872),
  inference(cnf_transformation,[],[f38452_D])).
fof(f38452_D,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl872),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl872])])).
fof(f44523,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1247,f977])).
fof(f1247,plain,(
  ( ! [X125,X126] : (~sP9(X125,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X126) | p(s(bool,X126)) | ~p(s(bool,X125))) )),
  inference(resolution,[],[f653,f734])).
fof(f38011,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl858),
  inference(cnf_transformation,[],[f38011_D])).
fof(f38011_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl858),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl858])])).
fof(f44448,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1246,f977])).
fof(f1246,plain,(
  ( ! [X123,X124] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X123,X124) | p(s(bool,X124)) | ~p(s(bool,X123))) )),
  inference(resolution,[],[f653,f734])).
fof(f37693,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl846),
  inference(cnf_transformation,[],[f37693_D])).
fof(f37693_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl846),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl846])])).
fof(f44344,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1245,f977])).
fof(f1245,plain,(
  ( ! [X121,X122] : (~sP5(X121,X122,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X122)) | p(s(bool,X121))) )),
  inference(resolution,[],[f653,f730])).
fof(f37045,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl820),
  inference(cnf_transformation,[],[f37045_D])).
fof(f37045_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl820),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl820])])).
fof(f44221,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1244,f977])).
fof(f1244,plain,(
  ( ! [X120,X119] : (~sP5(X119,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X120) | ~p(s(bool,X120)) | p(s(bool,X119))) )),
  inference(resolution,[],[f653,f730])).
fof(f36548,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl796),
  inference(cnf_transformation,[],[f36548_D])).
fof(f36548_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl796),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl796])])).
fof(f44128,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1241,f977])).
fof(f1241,plain,(
  ( ! [X114,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X113,X114) | ~p(s(bool,X113)) | ~p(s(bool,X114))) )),
  inference(resolution,[],[f653,f727])).
fof(f727,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f523])).
fof(f36473,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl794),
  inference(cnf_transformation,[],[f36473_D])).
fof(f36473_D,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl794),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl794])])).
fof(f44045,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1240,f977])).
fof(f1240,plain,(
  ( ! [X111,X112] : (sP5(X111,X112,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X112)) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f653,f727])).
fof(f36395,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl792),
  inference(cnf_transformation,[],[f36395_D])).
fof(f36395_D,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl792),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl792])])).
fof(f43968,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1239,f977])).
fof(f1239,plain,(
  ( ! [X109,X110] : (sP5(X109,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X110) | ~p(s(bool,X110)) | ~p(s(bool,X109))) )),
  inference(resolution,[],[f653,f727])).
fof(f36312,plain,(
  ( ! [X26,X27] : (~sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl790),
  inference(cnf_transformation,[],[f36312_D])).
fof(f36312_D,plain,(
  ( ! [X26,X27] : (~sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl790),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl790])])).
fof(f43898,plain,(
  ( ! [X26,X27] : (~sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1237,f977])).
fof(f1237,plain,(
  ( ! [X105,X106] : (~sP4(X105,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X106) | p(s(bool,X106)) | ~p(s(bool,X105))) )),
  inference(resolution,[],[f653,f721])).
fof(f36122,plain,(
  ( ! [X26,X27] : (~sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl788),
  inference(cnf_transformation,[],[f36122_D])).
fof(f36122_D,plain,(
  ( ! [X26,X27] : (~sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl788),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl788])])).
fof(f43827,plain,(
  ( ! [X26,X27] : (~sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1236,f977])).
fof(f1236,plain,(
  ( ! [X103,X104] : (~sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X103,X104) | p(s(bool,X104)) | ~p(s(bool,X103))) )),
  inference(resolution,[],[f653,f721])).
fof(f36044,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl784),
  inference(cnf_transformation,[],[f36044_D])).
fof(f36044_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl784),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl784])])).
fof(f43763,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1227,f977])).
fof(f1227,plain,(
  ( ! [X85,X86] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X85,X86) | p(s(bool,X86)) | p(s(bool,X85))) )),
  inference(resolution,[],[f653,f711])).
fof(f35871,plain,(
  ( ! [X26,X27] : (~sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl782),
  inference(cnf_transformation,[],[f35871_D])).
fof(f35871_D,plain,(
  ( ! [X26,X27] : (~sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl782),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl782])])).
fof(f43704,plain,(
  ( ! [X26,X27] : (~sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1220,f977])).
fof(f1220,plain,(
  ( ! [X72,X71] : (~sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X71,X72) | p(s(bool,X72)) | p(s(bool,X71))) )),
  inference(resolution,[],[f653,f704])).
fof(f35746,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl778),
  inference(cnf_transformation,[],[f35746_D])).
fof(f35746_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl778),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl778])])).
fof(f43632,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1214,f977])).
fof(f1214,plain,(
  ( ! [X59,X60] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X59,X60) | ~p(s(bool,X60)) | p(s(bool,X59))) )),
  inference(resolution,[],[f653,f694])).
fof(f35614,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl776),
  inference(cnf_transformation,[],[f35614_D])).
fof(f35614_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl776),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl776])])).
fof(f43559,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1213,f977])).
fof(f1213,plain,(
  ( ! [X57,X58] : (~sP1(X58,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X58)) | p(s(bool,X57))) )),
  inference(resolution,[],[f653,f694])).
fof(f35541,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl774),
  inference(cnf_transformation,[],[f35541_D])).
fof(f35541_D,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl774),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl774])])).
fof(f43472,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1209,f977])).
fof(f1209,plain,(
  ( ! [X50,X49] : (sP1(X49,X50,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X50)) | p(s(bool,X49))) )),
  inference(resolution,[],[f653,f691])).
fof(f35443,plain,(
  ( ! [X26,X27] : (~sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl772),
  inference(cnf_transformation,[],[f35443_D])).
fof(f35443_D,plain,(
  ( ! [X26,X27] : (~sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl772),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl772])])).
fof(f43397,plain,(
  ( ! [X26,X27] : (~sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1207,f977])).
fof(f1207,plain,(
  ( ! [X45,X46] : (~sP0(X45,X46,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X46)) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f653,f687])).
fof(f35251,plain,(
  ( ! [X26,X27] : (~sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl766),
  inference(cnf_transformation,[],[f35251_D])).
fof(f35251_D,plain,(
  ( ! [X26,X27] : (~sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | ~p(s(bool,X27))) ) <=> ~$spl766),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl766])])).
fof(f43326,plain,(
  ( ! [X26,X27] : (~sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1206,f977])).
fof(f1206,plain,(
  ( ! [X43,X44] : (~sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X43,X44) | ~p(s(bool,X44)) | p(s(bool,X43))) )),
  inference(resolution,[],[f653,f687])).
fof(f42185,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19)) ) | $spl976),
  inference(cnf_transformation,[],[f42185_D])).
fof(f42185_D,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19)) ) <=> ~$spl976),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl976])])).
fof(f43263,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f36240,f977])).
fof(f36240,plain,(
  ( ! [X15] : (~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X15)) ) | $spl31),
  inference(subsumption_resolution,[],[f36194,f9673])).
fof(f9673,plain,(
  ( ! [X136,X137] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X136,X137)) ) | $spl31),
  inference(resolution,[],[f1365,f741])).
fof(f741,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f529])).
fof(f36194,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,X15) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X15,'F_0') | ~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f10088,f23709])).
fof(f23709,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f23707,f942])).
fof(f942,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f376])).
fof(f376,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rich_list.prefixes_is_prefix_total',aHL_FALSITY)).
fof(f23707,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(resolution,[],[f23704,f736])).
fof(f736,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f526])).
fof(f23704,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f23688,f743])).
fof(f743,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f532])).
fof(f23688,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f23639,f741])).
fof(f23639,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f23635,f744])).
fof(f744,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f532])).
fof(f23635,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f23629,f732])).
fof(f732,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f526])).
fof(f23629,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f23628,f928])).
fof(f23628,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f23625,f23615])).
fof(f23615,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f23610,f747])).
fof(f747,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f534])).
fof(f23610,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) )),
  inference(resolution,[],[f12504,f731])).
fof(f731,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f526])).
fof(f12504,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f12497,f750])).
fof(f12497,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f12474,f652])).
fof(f12474,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f1090,f928])).
fof(f1090,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))),
  inference(resolution,[],[f652,f923])).
fof(f923,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f633])).
fof(f633,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f357])).
fof(f357,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X19] : (s(bool,'F_0') = s(bool,X19) <=> ~p(s(bool,X19)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.EQ_CLAUSES|split|3')).
fof(f23625,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) )),
  inference(resolution,[],[f23616,f735])).
fof(f735,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f526])).
fof(f23616,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f23611,f743])).
fof(f23611,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f12504,f732])).
fof(f10088,plain,(
  ( ! [X152,X153] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X153,X152) | s(bool,X152) != s(bool,X153)) ) | $spl31),
  inference(resolution,[],[f1365,f753])).
fof(f753,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f535])).
fof(f34961,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl764),
  inference(cnf_transformation,[],[f34961_D])).
fof(f34961_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl764),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl764])])).
fof(f42695,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12)))) | $spl982),
  inference(superposition,[],[f42474,f977])).
fof(f42476,plain,(
  s('type.list.list'(sK12),sK35(sK15,sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK35(sK15,sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK35(sK15,sK13,sK12),sK12)))) | $spl984),
  inference(cnf_transformation,[],[f42476_D])).
fof(f42476_D,plain,(
  s('type.list.list'(sK12),sK35(sK15,sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK35(sK15,sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK35(sK15,sK13,sK12),sK12)))) <=> ~$spl984),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl984])])).
fof(f42474,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl982),
  inference(cnf_transformation,[],[f42474_D])).
fof(f42474_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl982),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl982])])).
fof(f42212,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK35(sK15,sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK35(sK15,sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK35(sK15,sK13,sK12),sK12))))),
  inference(superposition,[],[f1191,f977])).
fof(f1191,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK35(sK15,sK13,sK12))))),
  inference(resolution,[],[f653,f958])).
fof(f42472,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK35(sK15,sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl981),
  inference(cnf_transformation,[],[f42472_D])).
fof(f42472_D,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK35(sK15,sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')))) <=> ~$spl981),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl981])])).
fof(f42207,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK35(sK15,sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f1191,f863])).
fof(f863,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f594])).
fof(f594,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.NIL_0'))))))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.NIL_0')))) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.NIL_0')))) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X11] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X11),s('type.list.list'(X6),'const.list.NIL_0')))) <=> s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X11))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_NIL|split|1')).
fof(f42470,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK35(sK15,sK13,sK12)))) | $spl978),
  inference(cnf_transformation,[],[f42470_D])).
fof(f42470_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK35(sK15,sK13,sK12)))) <=> ~$spl978),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl978])])).
fof(f42196,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK35(sK15,sK13,sK12)))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f1191,f977])).
fof(f42151,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f35808,f977])).
fof(f35808,plain,(
  ( ! [X15] : (~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X15)) ) | $spl19),
  inference(subsumption_resolution,[],[f35762,f5237])).
fof(f5237,plain,(
  ( ! [X136,X137] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X136,X137)) ) | $spl19),
  inference(resolution,[],[f1175,f741])).
fof(f35762,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,X15) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X15,'F_0') | ~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f7296,f23709])).
fof(f7296,plain,(
  ( ! [X152,X153] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X153,X152) | s(bool,X152) != s(bool,X153)) ) | $spl19),
  inference(resolution,[],[f1175,f753])).
fof(f42071,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19)) ) | $spl974),
  inference(cnf_transformation,[],[f42071_D])).
fof(f42071_D,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19)) ) <=> ~$spl974),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl974])])).
fof(f42053,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f34174,f977])).
fof(f34174,plain,(
  ( ! [X15] : (~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,'F_0') != s(bool,X15)) ) | $spl1),
  inference(subsumption_resolution,[],[f34130,f3187])).
fof(f3187,plain,(
  ( ! [X136,X137] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X136,X137)) ) | $spl1),
  inference(resolution,[],[f1161,f741])).
fof(f34130,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,X15) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X15,'F_0') | ~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f4844,f23709])).
fof(f4844,plain,(
  ( ! [X152,X153] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X153,X152) | s(bool,X152) != s(bool,X153)) ) | $spl1),
  inference(resolution,[],[f1161,f753])).
fof(f42015,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19)) ) | $spl972),
  inference(cnf_transformation,[],[f42015_D])).
fof(f42015_D,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19)) ) <=> ~$spl972),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl972])])).
fof(f41996,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f31235,f977])).
fof(f31235,plain,(
  ( ! [X15] : (~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,'F_0') != s(bool,X15)) )),
  inference(subsumption_resolution,[],[f31191,f2475])).
fof(f2475,plain,(
  ( ! [X136,X137] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X136,X137)) )),
  inference(resolution,[],[f655,f741])).
fof(f655,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))))),
  inference(cnf_transformation,[],[f498])).
fof(f31191,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,X15) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X15,'F_0') | ~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f2483,f23709])).
fof(f2483,plain,(
  ( ! [X152,X153] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X153,X152) | s(bool,X152) != s(bool,X153)) )),
  inference(resolution,[],[f655,f753])).
fof(f42012,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,'F_0') != s(bool,X19)) ) | $spl970),
  inference(cnf_transformation,[],[f42012_D])).
fof(f42012_D,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,'F_0') != s(bool,X19)) ) <=> ~$spl970),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl970])])).
fof(f41980,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,'F_0') != s(bool,X19) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f31235,f977])).
fof(f41215,plain,(
  s('type.list.list'(sK12),sK39(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK39(sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK39(sK13,sK12),sK12)))) | $spl968),
  inference(cnf_transformation,[],[f41215_D])).
fof(f41215_D,plain,(
  s('type.list.list'(sK12),sK39(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK39(sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK39(sK13,sK12),sK12)))) <=> ~$spl968),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl968])])).
fof(f41213,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl966),
  inference(cnf_transformation,[],[f41213_D])).
fof(f41213_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl966),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl966])])).
fof(f40926,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK39(sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK39(sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK39(sK13,sK12),sK12)))) | $spl20),
  inference(superposition,[],[f1178,f977])).
fof(f41211,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK39(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl965),
  inference(cnf_transformation,[],[f41211_D])).
fof(f41211_D,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK39(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')))) <=> ~$spl965),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl965])])).
fof(f40921,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK39(sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl20),
  inference(superposition,[],[f1178,f863])).
fof(f40910,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19)) ) | $spl962),
  inference(cnf_transformation,[],[f40910_D])).
fof(f40910_D,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19)) ) <=> ~$spl962),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl962])])).
fof(f40891,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X19) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f31115,f977])).
fof(f31115,plain,(
  ( ! [X15] : (~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,'F_0') != s(bool,X15)) )),
  inference(subsumption_resolution,[],[f31071,f1867])).
fof(f1867,plain,(
  ( ! [X136,X137] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X136,X137)) )),
  inference(resolution,[],[f654,f741])).
fof(f31071,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,X15) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X15,'F_0') | ~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f1875,f23709])).
fof(f1875,plain,(
  ( ! [X152,X153] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X153,X152) | s(bool,X152) != s(bool,X153)) )),
  inference(resolution,[],[f654,f753])).
fof(f40907,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,'F_0') != s(bool,X19)) ) | $spl960),
  inference(cnf_transformation,[],[f40907_D])).
fof(f40907_D,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,'F_0') != s(bool,X19)) ) <=> ~$spl960),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl960])])).
fof(f40875,plain,(
  ( ! [X19] : (~sP6(X19,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,'F_0') != s(bool,X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f31115,f977])).
fof(f1178,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12)))) | $spl20),
  inference(cnf_transformation,[],[f1178_D])).
fof(f1178_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f40802,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,X19)) ) | $spl958),
  inference(cnf_transformation,[],[f40802_D])).
fof(f40802_D,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,X19)) ) <=> ~$spl958),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl958])])).
fof(f40853,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl31),
  inference(subsumption_resolution,[],[f40839,f9673])).
fof(f40839,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19,'F_0') | s(bool,'F_0') = s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f30351,f977])).
fof(f30351,plain,(
  ( ! [X15] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X15,'F_0') | s(bool,'F_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f30300,f1255])).
fof(f1255,plain,(
  ( ! [X142,X141] : (sP6(X141,X142,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f747])).
fof(f30300,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X15,'F_0') | ~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f1259,f23709])).
fof(f1259,plain,(
  ( ! [X149,X150] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X150,X149) | s(bool,X149) = s(bool,X150)) )),
  inference(resolution,[],[f653,f752])).
fof(f40798,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | s(bool,'F_0') = s(bool,X19)) ) | $spl956),
  inference(cnf_transformation,[],[f40798_D])).
fof(f40798_D,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | s(bool,'F_0') = s(bool,X19)) ) <=> ~$spl956),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl956])])).
fof(f40823,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | s(bool,'F_0') = s(bool,X19) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f30351,f977])).
fof(f40801,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f40786,f5237])).
fof(f40786,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19,'F_0') | s(bool,'F_0') = s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f28276,f977])).
fof(f28276,plain,(
  ( ! [X15] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X15,'F_0') | s(bool,'F_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f28225,f1059])).
fof(f1059,plain,(
  ( ! [X142,X141] : (sP6(X141,X142,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f747])).
fof(f28225,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X15,'F_0') | ~sP6(X15,'F_0','const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f1063,f23709])).
fof(f1063,plain,(
  ( ! [X149,X150] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X150,X149) | s(bool,X149) = s(bool,X150)) )),
  inference(resolution,[],[f652,f752])).
fof(f40770,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | s(bool,'F_0') = s(bool,X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f28276,f977])).
fof(f40749,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | s(bool,'T_0') != s(bool,X19)) ) | $spl954),
  inference(cnf_transformation,[],[f40749_D])).
fof(f40749_D,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | s(bool,'T_0') != s(bool,X19)) ) <=> ~$spl954),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl954])])).
fof(f40722,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl440),
  inference(superposition,[],[f40577,f977])).
fof(f40577,plain,(
  ( ! [X4] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl440),
  inference(resolution,[],[f30347,f733])).
fof(f733,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X2,X1) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f526])).
fof(f30347,plain,(
  ( ! [X13] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X13,'F_0') | s(bool,'T_0') != s(bool,X13)) ) | $spl440),
  inference(subsumption_resolution,[],[f30298,f23443])).
fof(f30298,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,'T_0') != s(bool,X13) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X13,'F_0')) )),
  inference(resolution,[],[f1259,f23623])).
fof(f23623,plain,(
  ( ! [X4,X5] : (sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4) | ~sP8(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f23622,f942])).
fof(f23622,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f23619,f745])).
fof(f745,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f532])).
fof(f23619,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) )),
  inference(resolution,[],[f23615,f736])).
fof(f40683,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),'F_0',X19) | s(bool,'T_0') != s(bool,X19)) ) | $spl950),
  inference(cnf_transformation,[],[f40683_D])).
fof(f40683_D,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),'F_0',X19) | s(bool,'T_0') != s(bool,X19)) ) <=> ~$spl950),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl950])])).
fof(f40706,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),'F_0',X19) | s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl440),
  inference(superposition,[],[f40577,f977])).
fof(f40685,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | s(bool,'T_0') != s(bool,X19)) ) | $spl952),
  inference(cnf_transformation,[],[f40685_D])).
fof(f40685_D,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | s(bool,'T_0') != s(bool,X19)) ) <=> ~$spl952),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl952])])).
fof(f40660,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl440),
  inference(superposition,[],[f40525,f977])).
fof(f40525,plain,(
  ( ! [X4] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl440),
  inference(resolution,[],[f28272,f733])).
fof(f28272,plain,(
  ( ! [X13] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X13,'F_0') | s(bool,'T_0') != s(bool,X13)) ) | $spl440),
  inference(subsumption_resolution,[],[f28223,f23443])).
fof(f28223,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,'T_0') != s(bool,X13) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X13,'F_0')) )),
  inference(resolution,[],[f1063,f23623])).
fof(f40644,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),'F_0',X19) | s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl440),
  inference(superposition,[],[f40525,f977])).
fof(f2123,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,X49)) ) | $spl56),
  inference(cnf_transformation,[],[f2123_D])).
fof(f2123_D,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,X49)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f40623,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | ($spl31 | $spl440)),
  inference(subsumption_resolution,[],[f40609,f9673])).
fof(f40609,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19,'F_0') | s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl440),
  inference(superposition,[],[f30347,f977])).
fof(f40569,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | s(bool,'T_0') != s(bool,X19)) ) | $spl948),
  inference(cnf_transformation,[],[f40569_D])).
fof(f40569_D,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | s(bool,'T_0') != s(bool,X19)) ) <=> ~$spl948),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl948])])).
fof(f40593,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl440),
  inference(superposition,[],[f30347,f977])).
fof(f40572,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | ($spl19 | $spl440)),
  inference(subsumption_resolution,[],[f40557,f5237])).
fof(f40557,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19,'F_0') | s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl440),
  inference(superposition,[],[f28272,f977])).
fof(f40541,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | s(bool,'T_0') != s(bool,X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl440),
  inference(superposition,[],[f28272,f977])).
fof(f33610,plain,(
  ( ! [X31] : (s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl742),
  inference(cnf_transformation,[],[f33610_D])).
fof(f33610_D,plain,(
  ( ! [X31] : (s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12))))) ) <=> ~$spl742),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl742])])).
fof(f12249,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl284),
  inference(cnf_transformation,[],[f12249_D])).
fof(f12249_D,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f40486,plain,(
  ( ! [X41,X42,X40] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X41,X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37420,f977])).
fof(f37420,plain,(
  ( ! [X215,X213,X214] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X213),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')),X214,X215)) ) | $spl19),
  inference(resolution,[],[f27879,f745])).
fof(f27879,plain,(
  ( ! [X6] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X6),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) | $spl19),
  inference(subsumption_resolution,[],[f27878,f938])).
fof(f938,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f472])).
fof(f472,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f369])).
fof(f369,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X19] : ((p(s(bool,X19)) => p(s(bool,'F_0'))) => ~p(s(bool,X19)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.IMP_F')).
fof(f27878,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X6),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) | $spl19),
  inference(forward_demodulation,[],[f27826,f7317])).
fof(f7317,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl19),
  inference(resolution,[],[f1175,f878])).
fof(f878,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f604])).
fof(f604,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f603])).
fof(f603,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f236])).
fof(f236,axiom,(
  ! [X19] : ((p(s(bool,X19)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X19))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.IMP_F_EQ_F')).
fof(f27826,plain,(
  ( ! [X6] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X6),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f993,f863])).
fof(f993,plain,(
  ( ! [X8] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X8),s('type.list.list'(sK12),sK13))))))) )),
  inference(resolution,[],[f652,f950])).
fof(f950,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f640])).
fof(f640,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))) & ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))) & s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2))))))))),
  inference(flattening,[],[f639])).
fof(f639,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) | (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2))))) & ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))) & s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2))))))))),
  inference(nnf_transformation,[],[f380])).
fof(f380,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) <=> (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f379])).
fof(f379,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) <=> (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.SNOC_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6] : ! [X11,X18,X29] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X29),s('type.list.list'(X6),'const.list.SNOC_2'(s(X6,X11),s('type.list.list'(X6),X18)))))) <=> (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X29),s('type.list.list'(X6),X18)))) | s('type.list.list'(X6),X29) = s('type.list.list'(X6),'const.list.SNOC_2'(s(X6,X11),s('type.list.list'(X6),X18)))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_SNOC')).
fof(f40517,plain,(
  ( ! [X35,X33,X34] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) | $spl946),
  inference(cnf_transformation,[],[f40517_D])).
fof(f40517_D,plain,(
  ( ! [X35,X33,X34] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) <=> ~$spl946),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl946])])).
fof(f40468,plain,(
  ( ! [X35,X33,X34] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37420,f977])).
fof(f12204,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl282),
  inference(cnf_transformation,[],[f12204_D])).
fof(f12204_D,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f40408,plain,(
  ( ! [X41,X42,X40] : (sP7(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37419,f977])).
fof(f37419,plain,(
  ( ! [X212,X210,X211] : (sP7(X210,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X211),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')),X212)) ) | $spl19),
  inference(resolution,[],[f27879,f744])).
fof(f40439,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35)) ) | $spl944),
  inference(cnf_transformation,[],[f40439_D])).
fof(f40439_D,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35)) ) <=> ~$spl944),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl944])])).
fof(f40390,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37419,f977])).
fof(f12159,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl280),
  inference(cnf_transformation,[],[f12159_D])).
fof(f12159_D,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f40330,plain,(
  ( ! [X41,X42,X40] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X41,X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37417,f977])).
fof(f37417,plain,(
  ( ! [X206,X204,X205] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X204),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')),X205,X206)) ) | $spl19),
  inference(resolution,[],[f27879,f741])).
fof(f40361,plain,(
  ( ! [X35,X33,X34] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) | $spl942),
  inference(cnf_transformation,[],[f40361_D])).
fof(f40361_D,plain,(
  ( ! [X35,X33,X34] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) <=> ~$spl942),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl942])])).
fof(f40312,plain,(
  ( ! [X35,X33,X34] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37417,f977])).
fof(f40282,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl940),
  inference(cnf_transformation,[],[f40282_D])).
fof(f40282_D,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | ~p(s(bool,X26))) ) <=> ~$spl940),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl940])])).
fof(f40248,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1132,f977])).
fof(f1132,plain,(
  ( ! [X4,X5] : (~sP9(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(superposition,[],[f652,f752])).
fof(f40232,plain,(
  ( ! [X26,X27] : (~sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1132,f977])).
fof(f12114,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl278),
  inference(cnf_transformation,[],[f12114_D])).
fof(f12114_D,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f40168,plain,(
  ( ! [X41,X42,X40] : (sP8(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37416,f977])).
fof(f37416,plain,(
  ( ! [X202,X203,X201] : (sP8(X201,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X202),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')),X203)) ) | $spl19),
  inference(resolution,[],[f27879,f740])).
fof(f740,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f529])).
fof(f40200,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35)) ) | $spl936),
  inference(cnf_transformation,[],[f40200_D])).
fof(f40200_D,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35)) ) <=> ~$spl936),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl936])])).
fof(f40150,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37416,f977])).
fof(f40198,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X14),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl934),
  inference(cnf_transformation,[],[f40198_D])).
fof(f40198_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X14),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl934),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl934])])).
fof(f40124,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X14),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X15))) ) | $spl19),
  inference(resolution,[],[f37416,f23905])).
fof(f23905,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f23904,f747])).
fof(f23904,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f23903,f876])).
fof(f876,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f602])).
fof(f602,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f601])).
fof(f601,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f340])).
fof(f340,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X19] : ((p(s(bool,X19)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X19)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.IMP_CLAUSES|split|4')).
fof(f23903,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f23896,f745])).
fof(f23896,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) )),
  inference(resolution,[],[f12506,f736])).
fof(f12506,plain,(
  ( ! [X10,X11] : (~sP9(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) )),
  inference(superposition,[],[f12497,f752])).
fof(f40197,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl932),
  inference(cnf_transformation,[],[f40197_D])).
fof(f40197_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl932),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl932])])).
fof(f40123,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X13))) ) | $spl19),
  inference(resolution,[],[f37416,f23863])).
fof(f23863,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f23857,f747])).
fof(f23857,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f23709,f752])).
fof(f12070,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl276),
  inference(cnf_transformation,[],[f12070_D])).
fof(f12070_D,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f40086,plain,(
  ( ! [X41,X42,X40] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X41,X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37404,f977])).
fof(f37404,plain,(
  ( ! [X167,X165,X166] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X165),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')),X166,X167)) ) | $spl19),
  inference(resolution,[],[f27879,f724])).
fof(f724,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f521])).
fof(f40115,plain,(
  ( ! [X35,X33,X34] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) | $spl930),
  inference(cnf_transformation,[],[f40115_D])).
fof(f40115_D,plain,(
  ( ! [X35,X33,X34] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) <=> ~$spl930),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl930])])).
fof(f40068,plain,(
  ( ! [X35,X33,X34] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37404,f977])).
fof(f12026,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl274),
  inference(cnf_transformation,[],[f12026_D])).
fof(f12026_D,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f40008,plain,(
  ( ! [X41,X42,X40] : (sP4(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37403,f977])).
fof(f37403,plain,(
  ( ! [X163,X164,X162] : (sP4(X162,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X163),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')),X164)) ) | $spl19),
  inference(resolution,[],[f27879,f723])).
fof(f723,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f521])).
fof(f40038,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35)) ) | $spl928),
  inference(cnf_transformation,[],[f40038_D])).
fof(f40038_D,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35)) ) <=> ~$spl928),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl928])])).
fof(f39990,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37403,f977])).
fof(f11982,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl272),
  inference(cnf_transformation,[],[f11982_D])).
fof(f11982_D,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f39930,plain,(
  ( ! [X41,X42,X40] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X41,X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37390,f977])).
fof(f37390,plain,(
  ( ! [X125,X123,X124] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X123),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')),X124,X125)) ) | $spl19),
  inference(resolution,[],[f27879,f707])).
fof(f707,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f514])).
fof(f39960,plain,(
  ( ! [X35,X33,X34] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) | $spl926),
  inference(cnf_transformation,[],[f39960_D])).
fof(f39960_D,plain,(
  ( ! [X35,X33,X34] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) <=> ~$spl926),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl926])])).
fof(f39912,plain,(
  ( ! [X35,X33,X34] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37390,f977])).
fof(f11937,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl270),
  inference(cnf_transformation,[],[f11937_D])).
fof(f11937_D,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f39854,plain,(
  ( ! [X41,X42,X40] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X41,X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37374,f977])).
fof(f37374,plain,(
  ( ! [X76,X77,X75] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X75),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')),X76,X77)) ) | $spl19),
  inference(resolution,[],[f27879,f690])).
fof(f690,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f506])).
fof(f39883,plain,(
  ( ! [X35,X33,X34] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) | $spl924),
  inference(cnf_transformation,[],[f39883_D])).
fof(f39883_D,plain,(
  ( ! [X35,X33,X34] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35)) ) <=> ~$spl924),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl924])])).
fof(f39836,plain,(
  ( ! [X35,X33,X34] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')),X34,X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37374,f977])).
fof(f39809,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl922),
  inference(cnf_transformation,[],[f39809_D])).
fof(f39809_D,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl922),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl922])])).
fof(f39766,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1131,f977])).
fof(f1131,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X2) | p(s(bool,X2)) | p(s(bool,X3))) )),
  inference(superposition,[],[f652,f750])).
fof(f39750,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1131,f977])).
fof(f11893,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl268),
  inference(cnf_transformation,[],[f11893_D])).
fof(f11893_D,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f39692,plain,(
  ( ! [X41,X42,X40] : (sP0(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) ) | $spl19),
  inference(superposition,[],[f37373,f977])).
fof(f37373,plain,(
  ( ! [X74,X72,X73] : (sP0(X72,X73,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X74),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f27879,f688])).
fof(f688,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f506])).
fof(f39721,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl918),
  inference(cnf_transformation,[],[f39721_D])).
fof(f39721_D,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl918),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl918])])).
fof(f39674,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37373,f977])).
fof(f39298,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) | $spl912),
  inference(cnf_transformation,[],[f39298_D])).
fof(f39298_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl912),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl912])])).
fof(f39615,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f36131,f977])).
fof(f36131,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X6) = s(bool,X7)) ) | $spl31),
  inference(resolution,[],[f10087,f732])).
fof(f10087,plain,(
  ( ! [X151,X150] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X150,X151) | s(bool,X150) = s(bool,X151)) ) | $spl31),
  inference(resolution,[],[f1365,f750])).
fof(f39251,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) | $spl910),
  inference(cnf_transformation,[],[f39251_D])).
fof(f39251_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl910),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl910])])).
fof(f39571,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f36130,f977])).
fof(f36130,plain,(
  ( ! [X4,X5] : (sP6(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X4) = s(bool,X5)) ) | $spl31),
  inference(resolution,[],[f10087,f731])).
fof(f38801,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,X26) = s(bool,X27)) ) | $spl888),
  inference(cnf_transformation,[],[f38801_D])).
fof(f38801_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl888),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl888])])).
fof(f39531,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f35883,f977])).
fof(f35883,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | s(bool,X6) = s(bool,X7)) ) | $spl23),
  inference(resolution,[],[f9493,f732])).
fof(f9493,plain,(
  ( ! [X151,X150] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X150,X151) | s(bool,X150) = s(bool,X151)) ) | $spl23),
  inference(resolution,[],[f1357,f750])).
fof(f38688,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,X26) = s(bool,X27)) ) | $spl882),
  inference(cnf_transformation,[],[f38688_D])).
fof(f38688_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl882),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl882])])).
fof(f39492,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f35882,f977])).
fof(f35882,plain,(
  ( ! [X4,X5] : (sP6(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | s(bool,X4) = s(bool,X5)) ) | $spl23),
  inference(resolution,[],[f9493,f731])).
fof(f39479,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | p(s(bool,X27))) ) | $spl916),
  inference(cnf_transformation,[],[f39479_D])).
fof(f39479_D,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl916),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl916])])).
fof(f39435,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1130,f977])).
fof(f1130,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(superposition,[],[f652,f750])).
fof(f39419,plain,(
  ( ! [X26,X27] : (sP9(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1130,f977])).
fof(f39358,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f35625,f977])).
fof(f35625,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X6) = s(bool,X7)) ) | $spl19),
  inference(resolution,[],[f7295,f732])).
fof(f7295,plain,(
  ( ! [X151,X150] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X150,X151) | s(bool,X150) = s(bool,X151)) ) | $spl19),
  inference(resolution,[],[f1175,f750])).
fof(f39314,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f35624,f977])).
fof(f35624,plain,(
  ( ! [X4,X5] : (sP6(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X4) = s(bool,X5)) ) | $spl19),
  inference(resolution,[],[f7295,f731])).
fof(f39268,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f35259,f977])).
fof(f35259,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X6) = s(bool,X7)) ) | $spl17),
  inference(resolution,[],[f5031,f732])).
fof(f5031,plain,(
  ( ! [X151,X150] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X150,X151) | s(bool,X150) = s(bool,X151)) ) | $spl17),
  inference(resolution,[],[f1174,f750])).
fof(f39223,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f35258,f977])).
fof(f35258,plain,(
  ( ! [X4,X5] : (sP6(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X4) = s(bool,X5)) ) | $spl17),
  inference(resolution,[],[f5031,f731])).
fof(f39210,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) | $spl908),
  inference(cnf_transformation,[],[f39210_D])).
fof(f39210_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl908),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl908])])).
fof(f39197,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f34070,f977])).
fof(f34070,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X6) = s(bool,X7)) ) | $spl1),
  inference(resolution,[],[f4843,f732])).
fof(f4843,plain,(
  ( ! [X151,X150] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X150,X151) | s(bool,X150) = s(bool,X151)) ) | $spl1),
  inference(resolution,[],[f1161,f750])).
fof(f38882,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,X26) = s(bool,X27)) ) | $spl896),
  inference(cnf_transformation,[],[f38882_D])).
fof(f38882_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl896),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl896])])).
fof(f39181,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f34070,f977])).
fof(f39167,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f39129,f5239])).
fof(f5239,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X141)) ) | $spl19),
  inference(resolution,[],[f1175,f744])).
fof(f39129,plain,(
  ( ! [X26,X27] : (~sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1057,f977])).
fof(f1057,plain,(
  ( ! [X138,X137] : (~sP7(X137,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X138) | p(s(bool,X138)) | ~p(s(bool,X137))) )),
  inference(resolution,[],[f652,f742])).
fof(f39113,plain,(
  ( ! [X26,X27] : (~sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1057,f977])).
fof(f39083,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) | $spl904),
  inference(cnf_transformation,[],[f39083_D])).
fof(f39083_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl904),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl904])])).
fof(f39072,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f34069,f977])).
fof(f34069,plain,(
  ( ! [X4,X5] : (sP6(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X4) = s(bool,X5)) ) | $spl1),
  inference(resolution,[],[f4843,f731])).
fof(f38842,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,X26) = s(bool,X27)) ) | $spl892),
  inference(cnf_transformation,[],[f38842_D])).
fof(f38842_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl892),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl892])])).
fof(f39056,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f34069,f977])).
fof(f39043,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f39006,f5240])).
fof(f5240,plain,(
  ( ! [X142,X143] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X142,X143)) ) | $spl19),
  inference(resolution,[],[f1175,f745])).
fof(f39006,plain,(
  ( ! [X26,X27] : (~sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1056,f977])).
fof(f1056,plain,(
  ( ! [X136,X135] : (~sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X135,X136) | p(s(bool,X136)) | ~p(s(bool,X135))) )),
  inference(resolution,[],[f652,f742])).
fof(f38990,plain,(
  ( ! [X26,X27] : (~sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1056,f977])).
fof(f38961,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f38928,f5236])).
fof(f5236,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X135)) ) | $spl19),
  inference(resolution,[],[f1175,f740])).
fof(f38928,plain,(
  ( ! [X26,X27] : (~sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1054,f977])).
fof(f1054,plain,(
  ( ! [X132,X131] : (~sP8(X131,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X132) | p(s(bool,X132)) | ~p(s(bool,X131))) )),
  inference(resolution,[],[f652,f738])).
fof(f38912,plain,(
  ( ! [X26,X27] : (~sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1054,f977])).
fof(f38884,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) | $spl898),
  inference(cnf_transformation,[],[f38884_D])).
fof(f38884_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl898),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl898])])).
fof(f38872,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f31130,f977])).
fof(f31130,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X6) = s(bool,X7)) )),
  inference(resolution,[],[f2482,f732])).
fof(f2482,plain,(
  ( ! [X151,X150] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X150,X151) | s(bool,X150) = s(bool,X151)) )),
  inference(resolution,[],[f655,f750])).
fof(f38856,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f31130,f977])).
fof(f38844,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) | $spl894),
  inference(cnf_transformation,[],[f38844_D])).
fof(f38844_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl894),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl894])])).
fof(f38832,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f31129,f977])).
fof(f31129,plain,(
  ( ! [X4,X5] : (sP6(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X4) = s(bool,X5)) )),
  inference(resolution,[],[f2482,f731])).
fof(f38816,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f31129,f977])).
fof(f38803,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) | $spl890),
  inference(cnf_transformation,[],[f38803_D])).
fof(f38803_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl890),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl890])])).
fof(f38791,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f31010,f977])).
fof(f31010,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X6) = s(bool,X7)) )),
  inference(resolution,[],[f1874,f732])).
fof(f1874,plain,(
  ( ! [X151,X150] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X150,X151) | s(bool,X150) = s(bool,X151)) )),
  inference(resolution,[],[f654,f750])).
fof(f38775,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f31010,f977])).
fof(f38763,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f38733,f5237])).
fof(f38733,plain,(
  ( ! [X26,X27] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1053,f977])).
fof(f1053,plain,(
  ( ! [X130,X129] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X129,X130) | p(s(bool,X130)) | ~p(s(bool,X129))) )),
  inference(resolution,[],[f652,f738])).
fof(f38717,plain,(
  ( ! [X26,X27] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1053,f977])).
fof(f38690,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) | $spl884),
  inference(cnf_transformation,[],[f38690_D])).
fof(f38690_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl884),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl884])])).
fof(f38678,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f31009,f977])).
fof(f31009,plain,(
  ( ! [X4,X5] : (sP6(X5,X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X4) = s(bool,X5)) )),
  inference(resolution,[],[f1874,f731])).
fof(f38662,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f31009,f977])).
fof(f37796,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl852),
  inference(cnf_transformation,[],[f37796_D])).
fof(f37796_D,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl852),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl852])])).
fof(f38619,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29943,f977])).
fof(f29943,plain,(
  ( ! [X8,X9] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X9,X8) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f1258,f733])).
fof(f1258,plain,(
  ( ! [X147,X148] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X147,X148) | s(bool,X147) != s(bool,X148)) )),
  inference(resolution,[],[f653,f751])).
fof(f751,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f535])).
fof(f7767,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl132),
  inference(cnf_transformation,[],[f7767_D])).
fof(f7767_D,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f38587,plain,(
  ( ! [X41,X42,X40] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X41,X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29590,f977])).
fof(f29590,plain,(
  ( ! [X222,X220,X221] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X220),s('type.list.list'(sK12),sK13)))),X221,X222)) )),
  inference(resolution,[],[f1189,f749])).
fof(f749,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f534])).
fof(f1189,plain,(
  ( ! [X8] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X8),s('type.list.list'(sK12),sK13))))))) )),
  inference(resolution,[],[f653,f950])).
fof(f38602,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0')))),X34,X35)) ) | $spl880),
  inference(cnf_transformation,[],[f38602_D])).
fof(f38602_D,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0')))),X34,X35)) ) <=> ~$spl880),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl880])])).
fof(f38569,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0')))),X34,X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29590,f977])).
fof(f37345,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),sK13)))),X34,X35)) ) | $spl838),
  inference(cnf_transformation,[],[f37345_D])).
fof(f37345_D,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),sK13)))),X34,X35)) ) <=> ~$spl838),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl838])])).
fof(f38540,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),sK13)))),X34,X35) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29590,f977])).
fof(f38599,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X13),s('type.list.list'(sK12),sK13)))))) ) | $spl878),
  inference(cnf_transformation,[],[f38599_D])).
fof(f38599_D,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X13),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl878),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl878])])).
fof(f38532,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X13),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f29590,f23864])).
fof(f23864,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f23858,f741])).
fof(f23858,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f23709,f753])).
fof(f7727,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl130),
  inference(cnf_transformation,[],[f7727_D])).
fof(f7727_D,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f38513,plain,(
  ( ! [X41,X42,X40] : (sP6(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29589,f977])).
fof(f29589,plain,(
  ( ! [X218,X219,X217] : (sP6(X217,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X218),s('type.list.list'(sK12),sK13)))),X219)) )),
  inference(resolution,[],[f1189,f748])).
fof(f748,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f534])).
fof(f38527,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) | $spl876),
  inference(cnf_transformation,[],[f38527_D])).
fof(f38527_D,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) <=> ~$spl876),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl876])])).
fof(f38495,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29589,f977])).
fof(f37269,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35)) ) | $spl832),
  inference(cnf_transformation,[],[f37269_D])).
fof(f37269_D,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35)) ) <=> ~$spl832),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl832])])).
fof(f38466,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29589,f977])).
fof(f38454,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl874),
  inference(cnf_transformation,[],[f38454_D])).
fof(f38454_D,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl874),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl874])])).
fof(f38418,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1051,f977])).
fof(f1051,plain,(
  ( ! [X125,X126] : (~sP9(X125,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X126) | p(s(bool,X126)) | ~p(s(bool,X125))) )),
  inference(resolution,[],[f652,f734])).
fof(f38402,plain,(
  ( ! [X26,X27] : (~sP9(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1051,f977])).
fof(f7687,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl128),
  inference(cnf_transformation,[],[f7687_D])).
fof(f7687_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f38365,plain,(
  ( ! [X41,X42,X40] : (sP6(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29588,f977])).
fof(f29588,plain,(
  ( ! [X216,X215,X214] : (sP6(X214,X215,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X216),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f747])).
fof(f38379,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl870),
  inference(cnf_transformation,[],[f38379_D])).
fof(f38379_D,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl870),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl870])])).
fof(f38347,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29588,f977])).
fof(f37195,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) | $spl828),
  inference(cnf_transformation,[],[f37195_D])).
fof(f37195_D,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl828),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl828])])).
fof(f38318,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29588,f977])).
fof(f7647,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl126),
  inference(cnf_transformation,[],[f7647_D])).
fof(f7647_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f38292,plain,(
  ( ! [X41,X42,X40] : (sP7(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29587,f977])).
fof(f29587,plain,(
  ( ! [X212,X213,X211] : (sP7(X211,X212,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X213),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f743])).
fof(f38305,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl868),
  inference(cnf_transformation,[],[f38305_D])).
fof(f38305_D,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl868),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl868])])).
fof(f38274,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29587,f977])).
fof(f37120,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) | $spl824),
  inference(cnf_transformation,[],[f37120_D])).
fof(f37120_D,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl824),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl824])])).
fof(f38245,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29587,f977])).
fof(f7608,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl124),
  inference(cnf_transformation,[],[f7608_D])).
fof(f7608_D,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f38220,plain,(
  ( ! [X41,X42,X40] : (sP8(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29584,f977])).
fof(f29584,plain,(
  ( ! [X204,X202,X203] : (sP8(X202,X203,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X204),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f739])).
fof(f739,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f529])).
fof(f38233,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl866),
  inference(cnf_transformation,[],[f38233_D])).
fof(f38233_D,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl866),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl866])])).
fof(f38202,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29584,f977])).
fof(f36925,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) | $spl816),
  inference(cnf_transformation,[],[f36925_D])).
fof(f36925_D,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl816),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl816])])).
fof(f38173,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29584,f977])).
fof(f7569,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl122),
  inference(cnf_transformation,[],[f7569_D])).
fof(f7569_D,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f38148,plain,(
  ( ! [X41,X42,X40] : (sP4(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29571,f977])).
fof(f29571,plain,(
  ( ! [X165,X163,X164] : (sP4(X163,X164,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X165),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f722])).
fof(f722,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f521])).
fof(f38161,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl864),
  inference(cnf_transformation,[],[f38161_D])).
fof(f38161_D,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl864),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl864])])).
fof(f38130,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29571,f977])).
fof(f36852,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) | $spl812),
  inference(cnf_transformation,[],[f36852_D])).
fof(f36852_D,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl812),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl812])])).
fof(f38101,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29571,f977])).
fof(f7529,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl120),
  inference(cnf_transformation,[],[f7529_D])).
fof(f7529_D,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f38074,plain,(
  ( ! [X41,X42,X40] : (sP2(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29555,f977])).
fof(f29555,plain,(
  ( ! [X116,X117,X115] : (sP2(X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X116),s('type.list.list'(sK12),sK13)))),X117)) )),
  inference(resolution,[],[f1189,f706])).
fof(f706,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f514])).
fof(f38086,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) | $spl862),
  inference(cnf_transformation,[],[f38086_D])).
fof(f38086_D,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) <=> ~$spl862),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl862])])).
fof(f38056,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29555,f977])).
fof(f36776,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35)) ) | $spl808),
  inference(cnf_transformation,[],[f36776_D])).
fof(f36776_D,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35)) ) <=> ~$spl808),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl808])])).
fof(f38027,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29555,f977])).
fof(f38013,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl860),
  inference(cnf_transformation,[],[f38013_D])).
fof(f38013_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) <=> ~$spl860),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl860])])).
fof(f37980,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1050,f977])).
fof(f1050,plain,(
  ( ! [X123,X124] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X123,X124) | p(s(bool,X124)) | ~p(s(bool,X123))) )),
  inference(resolution,[],[f652,f734])).
fof(f37964,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1050,f977])).
fof(f7490,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl118),
  inference(cnf_transformation,[],[f7490_D])).
fof(f7490_D,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f37931,plain,(
  ( ! [X41,X42,X40] : (sP2(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29554,f977])).
fof(f29554,plain,(
  ( ! [X114,X112,X113] : (sP2(X112,X113,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X114),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f705])).
fof(f705,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f514])).
fof(f37944,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl856),
  inference(cnf_transformation,[],[f37944_D])).
fof(f37944_D,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl856),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl856])])).
fof(f37913,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29554,f977])).
fof(f36702,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) | $spl804),
  inference(cnf_transformation,[],[f36702_D])).
fof(f36702_D,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl804),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl804])])).
fof(f37884,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29554,f977])).
fof(f7451,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl116),
  inference(cnf_transformation,[],[f7451_D])).
fof(f7451_D,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f37858,plain,(
  ( ! [X41,X42,X40] : (sP0(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29541,f977])).
fof(f29541,plain,(
  ( ! [X74,X75,X73] : (sP0(X73,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X74),s('type.list.list'(sK12),sK13)))),X75)) )),
  inference(resolution,[],[f1189,f689])).
fof(f689,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f506])).
fof(f37870,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) | $spl854),
  inference(cnf_transformation,[],[f37870_D])).
fof(f37870_D,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) <=> ~$spl854),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl854])])).
fof(f37840,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29541,f977])).
fof(f36627,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35)) ) | $spl800),
  inference(cnf_transformation,[],[f36627_D])).
fof(f36627_D,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35)) ) <=> ~$spl800),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl800])])).
fof(f37811,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29541,f977])).
fof(f37765,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27889,f977])).
fof(f27889,plain,(
  ( ! [X8,X9] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X9,X8) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f1062,f733])).
fof(f1062,plain,(
  ( ! [X147,X148] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X147,X148) | s(bool,X147) != s(bool,X148)) )),
  inference(resolution,[],[f652,f751])).
fof(f37748,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl850),
  inference(cnf_transformation,[],[f37748_D])).
fof(f37748_D,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl850),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl850])])).
fof(f37718,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f37348,f977])).
fof(f37348,plain,(
  ( ! [X0] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X0),s('type.list.list'(sK12),sK13)))) ) | $spl19),
  inference(resolution,[],[f27879,f864])).
fof(f864,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f594])).
fof(f37697,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27))) ) | $spl848),
  inference(cnf_transformation,[],[f37697_D])).
fof(f37697_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27))) ) <=> ~$spl848),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl848])])).
fof(f37696,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f37630,f7269])).
fof(f7269,plain,(
  ( ! [X99,X98] : (~sP5(X98,X99,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X98))) ) | $spl19),
  inference(resolution,[],[f1175,f715])).
fof(f715,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f519])).
fof(f519,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f518])).
fof(f518,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f487])).
fof(f37630,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1049,f977])).
fof(f1049,plain,(
  ( ! [X121,X122] : (~sP5(X121,X122,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X122)) | p(s(bool,X121))) )),
  inference(resolution,[],[f652,f730])).
fof(f37614,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1049,f977])).
fof(f37568,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) | $spl844),
  inference(cnf_transformation,[],[f37568_D])).
fof(f37568_D,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) <=> ~$spl844),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl844])])).
fof(f37483,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(superposition,[],[f27879,f977])).
fof(f37561,plain,(
  ( ! [X26,X24] : (~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl842),
  inference(cnf_transformation,[],[f37561_D])).
fof(f37561_D,plain,(
  ( ! [X26,X24] : (~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl842),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl842])])).
fof(f37461,plain,(
  ( ! [X282,X283,X284] : (p(s(bool,X282)) | ~p(s(bool,X283)) | s(bool,X283) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X284),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f27879,f978])).
fof(f37366,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X51) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X52),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X50) != s(bool,X53)) ) | $spl19),
  inference(resolution,[],[f27879,f679])).
fof(f37365,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X49),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f27879,f678])).
fof(f37363,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X42),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f27879,f677])).
fof(f37361,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X31) != s(bool,X34)) ) | $spl19),
  inference(resolution,[],[f27879,f675])).
fof(f37359,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f27879,f673])).
fof(f3118,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X6,X7)) ) | $spl102),
  inference(cnf_transformation,[],[f3118_D])).
fof(f3118_D,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X6,X7)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f37331,plain,(
  ( ! [X41,X42,X40] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X41,X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27741,f977])).
fof(f27741,plain,(
  ( ! [X222,X220,X221] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X220),s('type.list.list'(sK12),sK13)))),X221,X222)) )),
  inference(resolution,[],[f993,f749])).
fof(f33609,plain,(
  ( ! [X9] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X9),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) | $spl740),
  inference(cnf_transformation,[],[f33609_D])).
fof(f33609_D,plain,(
  ( ! [X9] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X9),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) <=> ~$spl740),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl740])])).
fof(f37325,plain,(
  ( ! [X14,X12,X13] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X13,X14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27741,f863])).
fof(f37347,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0')))),X34,X35)) ) | $spl840),
  inference(cnf_transformation,[],[f37347_D])).
fof(f37347_D,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0')))),X34,X35)) ) <=> ~$spl840),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl840])])).
fof(f37313,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),'const.list.NIL_0')))),X34,X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27741,f977])).
fof(f37284,plain,(
  ( ! [X35,X33,X34] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X33),s('type.list.list'(sK12),sK13)))),X34,X35) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27741,f977])).
fof(f37343,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X13),s('type.list.list'(sK12),sK13)))))) ) | $spl836),
  inference(cnf_transformation,[],[f37343_D])).
fof(f37343_D,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X13),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl836),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl836])])).
fof(f37276,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X13),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f27741,f23864])).
fof(f3077,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7)) ) | $spl98),
  inference(cnf_transformation,[],[f3077_D])).
fof(f3077_D,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f37256,plain,(
  ( ! [X41,X42,X40] : (sP6(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27740,f977])).
fof(f27740,plain,(
  ( ! [X218,X219,X217] : (sP6(X217,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X218),s('type.list.list'(sK12),sK13)))),X219)) )),
  inference(resolution,[],[f993,f748])).
fof(f37250,plain,(
  ( ! [X14,X12,X13] : (sP6(X13,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27740,f863])).
fof(f37271,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) | $spl834),
  inference(cnf_transformation,[],[f37271_D])).
fof(f37271_D,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) <=> ~$spl834),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl834])])).
fof(f37238,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27740,f977])).
fof(f37209,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27740,f977])).
fof(f3036,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f3036_D])).
fof(f3036_D,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f37182,plain,(
  ( ! [X41,X42,X40] : (sP6(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27739,f977])).
fof(f27739,plain,(
  ( ! [X216,X215,X214] : (sP6(X214,X215,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X216),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f747])).
fof(f37176,plain,(
  ( ! [X14,X12,X13] : (sP6(X13,X14,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27739,f863])).
fof(f37197,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl830),
  inference(cnf_transformation,[],[f37197_D])).
fof(f37197_D,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl830),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl830])])).
fof(f37164,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27739,f977])).
fof(f37135,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27739,f977])).
fof(f2995,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl90),
  inference(cnf_transformation,[],[f2995_D])).
fof(f2995_D,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f37108,plain,(
  ( ! [X41,X42,X40] : (sP7(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27738,f977])).
fof(f27738,plain,(
  ( ! [X212,X213,X211] : (sP7(X211,X212,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X213),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f743])).
fof(f37102,plain,(
  ( ! [X14,X12,X13] : (sP7(X13,X14,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27738,f863])).
fof(f37122,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl826),
  inference(cnf_transformation,[],[f37122_D])).
fof(f37122_D,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl826),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl826])])).
fof(f37090,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27738,f977])).
fof(f37061,plain,(
  ( ! [X35,X33,X34] : (sP7(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27738,f977])).
fof(f37049,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27))) ) | $spl822),
  inference(cnf_transformation,[],[f37049_D])).
fof(f37049_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27))) ) <=> ~$spl822),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl822])])).
fof(f37048,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36982,f7270])).
fof(f7270,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X101) | ~p(s(bool,X100))) ) | $spl19),
  inference(resolution,[],[f1175,f716])).
fof(f716,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f519])).
fof(f36982,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1048,f977])).
fof(f1048,plain,(
  ( ! [X120,X119] : (~sP5(X119,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X120) | ~p(s(bool,X120)) | p(s(bool,X119))) )),
  inference(resolution,[],[f652,f730])).
fof(f36966,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1048,f977])).
fof(f2955,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl86),
  inference(cnf_transformation,[],[f2955_D])).
fof(f2955_D,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f36913,plain,(
  ( ! [X41,X42,X40] : (sP8(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27735,f977])).
fof(f27735,plain,(
  ( ! [X204,X202,X203] : (sP8(X202,X203,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X204),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f739])).
fof(f36907,plain,(
  ( ! [X14,X12,X13] : (sP8(X13,X14,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27735,f863])).
fof(f36927,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl818),
  inference(cnf_transformation,[],[f36927_D])).
fof(f36927_D,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl818),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl818])])).
fof(f36895,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27735,f977])).
fof(f36866,plain,(
  ( ! [X35,X33,X34] : (sP8(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27735,f977])).
fof(f2915,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl82),
  inference(cnf_transformation,[],[f2915_D])).
fof(f2915_D,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f36840,plain,(
  ( ! [X41,X42,X40] : (sP4(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27722,f977])).
fof(f27722,plain,(
  ( ! [X165,X163,X164] : (sP4(X163,X164,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X165),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f722])).
fof(f36834,plain,(
  ( ! [X14,X12,X13] : (sP4(X13,X14,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27722,f863])).
fof(f36854,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl814),
  inference(cnf_transformation,[],[f36854_D])).
fof(f36854_D,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl814),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl814])])).
fof(f36822,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27722,f977])).
fof(f36793,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27722,f977])).
fof(f2874,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7)) ) | $spl78),
  inference(cnf_transformation,[],[f2874_D])).
fof(f2874_D,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f36765,plain,(
  ( ! [X41,X42,X40] : (sP2(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27706,f977])).
fof(f27706,plain,(
  ( ! [X116,X117,X115] : (sP2(X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X116),s('type.list.list'(sK12),sK13)))),X117)) )),
  inference(resolution,[],[f993,f706])).
fof(f36759,plain,(
  ( ! [X14,X12,X13] : (sP2(X13,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27706,f863])).
fof(f36778,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) | $spl810),
  inference(cnf_transformation,[],[f36778_D])).
fof(f36778_D,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) <=> ~$spl810),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl810])])).
fof(f36747,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27706,f977])).
fof(f36718,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27706,f977])).
fof(f2834,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f2834_D])).
fof(f2834_D,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f36690,plain,(
  ( ! [X41,X42,X40] : (sP2(X41,X42,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27705,f977])).
fof(f27705,plain,(
  ( ! [X114,X112,X113] : (sP2(X112,X113,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X114),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f705])).
fof(f36684,plain,(
  ( ! [X14,X12,X13] : (sP2(X13,X14,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27705,f863])).
fof(f36704,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl806),
  inference(cnf_transformation,[],[f36704_D])).
fof(f36704_D,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl806),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl806])])).
fof(f36672,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27705,f977])).
fof(f36643,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X35),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27705,f977])).
fof(f2752,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7)) ) | $spl68),
  inference(cnf_transformation,[],[f2752_D])).
fof(f2752_D,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f36616,plain,(
  ( ! [X41,X42,X40] : (sP0(X41,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X42) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27692,f977])).
fof(f27692,plain,(
  ( ! [X74,X75,X73] : (sP0(X73,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X74),s('type.list.list'(sK12),sK13)))),X75)) )),
  inference(resolution,[],[f993,f689])).
fof(f36610,plain,(
  ( ! [X14,X12,X13] : (sP0(X13,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X12),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27692,f863])).
fof(f36629,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) | $spl802),
  inference(cnf_transformation,[],[f36629_D])).
fof(f36629_D,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35)) ) <=> ~$spl802),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl802])])).
fof(f36598,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),'const.list.NIL_0')))),X35) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27692,f977])).
fof(f36569,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)))),X35) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27692,f977])).
fof(f36554,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl798),
  inference(cnf_transformation,[],[f36554_D])).
fof(f36554_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl798),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl798])])).
fof(f36553,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36552,f7279])).
fof(f7279,plain,(
  ( ! [X118,X119] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X118,X119) | p(s(bool,X118))) ) | $spl19),
  inference(resolution,[],[f1175,f726])).
fof(f726,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f523])).
fof(f36552,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36522,f7276])).
fof(f7276,plain,(
  ( ! [X112,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X112,X113) | p(s(bool,X113))) ) | $spl19),
  inference(resolution,[],[f1175,f725])).
fof(f725,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f523])).
fof(f36522,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1045,f977])).
fof(f1045,plain,(
  ( ! [X114,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X113,X114) | ~p(s(bool,X113)) | ~p(s(bool,X114))) )),
  inference(resolution,[],[f652,f727])).
fof(f36506,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1045,f977])).
fof(f36476,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36439,f7269])).
fof(f36439,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1044,f977])).
fof(f1044,plain,(
  ( ! [X111,X112] : (sP5(X111,X112,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X112)) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f652,f727])).
fof(f36423,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1044,f977])).
fof(f36398,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36361,f7270])).
fof(f36361,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1043,f977])).
fof(f1043,plain,(
  ( ! [X109,X110] : (sP5(X109,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X110) | ~p(s(bool,X110)) | ~p(s(bool,X109))) )),
  inference(resolution,[],[f652,f727])).
fof(f36345,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1043,f977])).
fof(f36315,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36285,f5223])).
fof(f5223,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X109)) ) | $spl19),
  inference(resolution,[],[f1175,f723])).
fof(f36285,plain,(
  ( ! [X26,X27] : (~sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1041,f977])).
fof(f1041,plain,(
  ( ! [X105,X106] : (~sP4(X105,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X106) | p(s(bool,X106)) | ~p(s(bool,X105))) )),
  inference(resolution,[],[f652,f721])).
fof(f36269,plain,(
  ( ! [X26,X27] : (~sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1041,f977])).
fof(f35371,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl770),
  inference(cnf_transformation,[],[f35371_D])).
fof(f35371_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl770),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl770])])).
fof(f36204,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10088,f977])).
fof(f35310,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl768),
  inference(cnf_transformation,[],[f35310_D])).
fof(f35310_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl768),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl768])])).
fof(f36144,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10087,f977])).
fof(f36125,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36095,f5224])).
fof(f5224,plain,(
  ( ! [X111,X110] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X110,X111)) ) | $spl19),
  inference(resolution,[],[f1175,f724])).
fof(f36095,plain,(
  ( ! [X26,X27] : (~sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1040,f977])).
fof(f1040,plain,(
  ( ! [X103,X104] : (~sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X103,X104) | p(s(bool,X104)) | ~p(s(bool,X103))) )),
  inference(resolution,[],[f652,f721])).
fof(f36079,plain,(
  ( ! [X26,X27] : (~sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1040,f977])).
fof(f36050,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl786),
  inference(cnf_transformation,[],[f36050_D])).
fof(f36050_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl786),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl786])])).
fof(f36049,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36048,f7254])).
fof(f7254,plain,(
  ( ! [X68,X69] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X68,X69) | ~p(s(bool,X68))) ) | $spl19),
  inference(resolution,[],[f1175,f698])).
fof(f698,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f511])).
fof(f511,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f510])).
fof(f510,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f509])).
fof(f509,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f484])).
fof(f36048,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f36025,f7253])).
fof(f7253,plain,(
  ( ! [X66,X67] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X67,X66) | ~p(s(bool,X66))) ) | $spl19),
  inference(resolution,[],[f1175,f697])).
fof(f697,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f511])).
fof(f36025,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1031,f977])).
fof(f1031,plain,(
  ( ! [X85,X86] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X85,X86) | p(s(bool,X86)) | p(s(bool,X85))) )),
  inference(resolution,[],[f652,f711])).
fof(f36009,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1031,f977])).
fof(f31117,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl638),
  inference(cnf_transformation,[],[f31117_D])).
fof(f31117_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl638),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl638])])).
fof(f35951,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9494,f977])).
fof(f9494,plain,(
  ( ! [X152,X153] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X153,X152) | s(bool,X152) != s(bool,X153)) ) | $spl23),
  inference(resolution,[],[f1357,f753])).
fof(f31059,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl634),
  inference(cnf_transformation,[],[f31059_D])).
fof(f31059_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl634),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl634])])).
fof(f35896,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9493,f977])).
fof(f35876,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f35854,f5210])).
fof(f5210,plain,(
  ( ! [X83,X82] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X82,X83)) ) | $spl19),
  inference(resolution,[],[f1175,f707])).
fof(f35854,plain,(
  ( ! [X26,X27] : (~sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1024,f977])).
fof(f1024,plain,(
  ( ! [X72,X71] : (~sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X71,X72) | p(s(bool,X72)) | p(s(bool,X71))) )),
  inference(resolution,[],[f652,f704])).
fof(f35838,plain,(
  ( ! [X26,X27] : (~sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1024,f977])).
fof(f35772,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7296,f977])).
fof(f35752,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl780),
  inference(cnf_transformation,[],[f35752_D])).
fof(f35752_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl780),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl780])])).
fof(f35751,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f35750,f7238])).
fof(f7238,plain,(
  ( ! [X37,X36] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X37,X36) | p(s(bool,X36))) ) | $spl19),
  inference(resolution,[],[f1175,f680])).
fof(f680,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f503])).
fof(f503,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f502])).
fof(f502,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f501])).
fof(f501,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f481])).
fof(f35750,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f35722,f7240])).
fof(f7240,plain,(
  ( ! [X41,X40] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X40,X41) | ~p(s(bool,X40))) ) | $spl19),
  inference(resolution,[],[f1175,f681])).
fof(f681,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f503])).
fof(f35722,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1018,f977])).
fof(f1018,plain,(
  ( ! [X59,X60] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X59,X60) | ~p(s(bool,X60)) | p(s(bool,X59))) )),
  inference(resolution,[],[f652,f694])).
fof(f35706,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1018,f977])).
fof(f35638,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7295,f977])).
fof(f35617,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f35587,f7248])).
fof(f7248,plain,(
  ( ! [X57,X56] : (sP1(X56,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X56))) ) | $spl19),
  inference(resolution,[],[f1175,f692])).
fof(f692,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f508])).
fof(f35587,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1017,f977])).
fof(f1017,plain,(
  ( ! [X57,X58] : (~sP1(X58,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X58)) | p(s(bool,X57))) )),
  inference(resolution,[],[f652,f694])).
fof(f35571,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1017,f977])).
fof(f35545,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f35497,f7239])).
fof(f7239,plain,(
  ( ! [X39,X38] : (~sP1(X38,X39,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X38))) ) | $spl19),
  inference(resolution,[],[f1175,f680])).
fof(f35497,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1013,f977])).
fof(f1013,plain,(
  ( ! [X50,X49] : (sP1(X49,X50,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | p(s(bool,X50)) | p(s(bool,X49))) )),
  inference(resolution,[],[f652,f691])).
fof(f35481,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1013,f977])).
fof(f35446,plain,(
  ( ! [X26,X27] : (p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f35416,f5193])).
fof(f5193,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f688])).
fof(f35416,plain,(
  ( ! [X26,X27] : (~sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1011,f977])).
fof(f1011,plain,(
  ( ! [X45,X46] : (~sP0(X45,X46,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | p(s(bool,X46)) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f652,f687])).
fof(f35400,plain,(
  ( ! [X26,X27] : (~sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X27)) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1011,f977])).
fof(f35333,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5032,f977])).
fof(f5032,plain,(
  ( ! [X152,X153] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X153,X152) | s(bool,X152) != s(bool,X153)) ) | $spl17),
  inference(resolution,[],[f1174,f753])).
fof(f35272,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5031,f977])).
fof(f35254,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f35224,f5194])).
fof(f5194,plain,(
  ( ! [X50,X51] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X50,X51)) ) | $spl19),
  inference(resolution,[],[f1175,f690])).
fof(f35224,plain,(
  ( ! [X26,X27] : (~sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1010,f977])).
fof(f1010,plain,(
  ( ! [X43,X44] : (~sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X43,X44) | ~p(s(bool,X44)) | p(s(bool,X43))) )),
  inference(resolution,[],[f652,f687])).
fof(f35208,plain,(
  ( ! [X26,X27] : (~sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1010,f977])).
fof(f34685,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) | $spl760),
  inference(superposition,[],[f34464,f977])).
fof(f34466,plain,(
  s('type.list.list'(sK12),sK35(sK14,sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK35(sK14,sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK35(sK14,sK13,sK12),sK12)))) | $spl762),
  inference(cnf_transformation,[],[f34466_D])).
fof(f34466_D,plain,(
  s('type.list.list'(sK12),sK35(sK14,sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK35(sK14,sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK35(sK14,sK13,sK12),sK12)))) <=> ~$spl762),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl762])])).
fof(f34464,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl760),
  inference(cnf_transformation,[],[f34464_D])).
fof(f34464_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl760),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl760])])).
fof(f34202,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK35(sK14,sK13,sK12)) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK35(sK14,sK13,sK12),sK12)),s('type.list.list'(sK12),sK39(sK35(sK14,sK13,sK12),sK12))))),
  inference(superposition,[],[f995,f977])).
fof(f34462,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK35(sK14,sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl759),
  inference(cnf_transformation,[],[f34462_D])).
fof(f34462_D,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK35(sK14,sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0')))) <=> ~$spl759),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl759])])).
fof(f34197,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK35(sK14,sK13,sK12)),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f995,f863])).
fof(f34460,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK35(sK14,sK13,sK12)))) | $spl756),
  inference(cnf_transformation,[],[f34460_D])).
fof(f34460_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK35(sK14,sK13,sK12)))) <=> ~$spl756),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl756])])).
fof(f34186,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK35(sK14,sK13,sK12)))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f995,f977])).
fof(f34177,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl754),
  inference(cnf_transformation,[],[f34177_D])).
fof(f34177_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl754),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl754])])).
fof(f34156,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4844,f977])).
fof(f31237,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl646),
  inference(cnf_transformation,[],[f31237_D])).
fof(f31237_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl646),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl646])])).
fof(f34140,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4844,f977])).
fof(f34120,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl752),
  inference(cnf_transformation,[],[f34120_D])).
fof(f34120_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl752),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl752])])).
fof(f34099,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4843,f977])).
fof(f31179,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl642),
  inference(cnf_transformation,[],[f31179_D])).
fof(f31179_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl642),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl642])])).
fof(f34083,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4843,f977])).
fof(f2670,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X4)) ) | $spl62),
  inference(cnf_transformation,[],[f2670_D])).
fof(f2670_D,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X4)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f33898,plain,(
  ( ! [X31,X32] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X32) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29602,f977])).
fof(f29602,plain,(
  ( ! [X248,X247] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X247),s('type.list.list'(sK12),sK13)))),X248)) )),
  inference(resolution,[],[f1189,f800])).
fof(f800,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f555])).
fof(f555,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f496])).
fof(f496,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f33892,plain,(
  ( ! [X10,X9] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X10) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X9),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f29602,f863])).
fof(f33911,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) | $spl750),
  inference(cnf_transformation,[],[f33911_D])).
fof(f33911_D,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) <=> ~$spl750),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl750])])).
fof(f33880,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),'const.list.NIL_0')))),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29602,f977])).
fof(f33754,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),sK13)))),X27)) ) | $spl744),
  inference(cnf_transformation,[],[f33754_D])).
fof(f33754_D,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),sK13)))),X27)) ) <=> ~$spl744),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl744])])).
fof(f33851,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),sK13)))),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29602,f977])).
fof(f2633,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl60),
  inference(cnf_transformation,[],[f2633_D])).
fof(f2633_D,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f33821,plain,(
  ( ! [X31,X32] : (sP11(X32,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f29601,f977])).
fof(f29601,plain,(
  ( ! [X246,X245] : (sP11(X245,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X246),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f799])).
fof(f799,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f556])).
fof(f33815,plain,(
  ( ! [X10,X9] : (sP11(X10,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X9),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f29601,f863])).
fof(f33834,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl748),
  inference(cnf_transformation,[],[f33834_D])).
fof(f33834_D,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl748),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl748])])).
fof(f33803,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f29601,f977])).
fof(f33606,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),sK13)))))) ) | $spl736),
  inference(cnf_transformation,[],[f33606_D])).
fof(f33606_D,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl736),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl736])])).
fof(f33774,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f29601,f977])).
fof(f1690,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X4)) ) | $spl40),
  inference(cnf_transformation,[],[f1690_D])).
fof(f1690_D,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X4)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f33742,plain,(
  ( ! [X31,X32] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X32) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27753,f977])).
fof(f27753,plain,(
  ( ! [X248,X247] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X247),s('type.list.list'(sK12),sK13)))),X248)) )),
  inference(resolution,[],[f993,f800])).
fof(f33736,plain,(
  ( ! [X10,X9] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X10) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X9),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27753,f863])).
fof(f33756,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) | $spl746),
  inference(cnf_transformation,[],[f33756_D])).
fof(f33756_D,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) <=> ~$spl746),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl746])])).
fof(f33724,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),'const.list.NIL_0')))),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27753,f977])).
fof(f33695,plain,(
  ( ! [X26,X27] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X26),s('type.list.list'(sK12),sK13)))),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27753,f977])).
fof(f33674,plain,(
  $false | ($spl19 | $spl23 | $spl36)),
  inference(subsumption_resolution,[],[f33673,f942])).
fof(f33673,plain,(
  p(s(bool,'F_0')) | ($spl19 | $spl23 | $spl36)),
  inference(forward_demodulation,[],[f33620,f7317])).
fof(f33620,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ($spl23 | $spl36)),
  inference(resolution,[],[f1652,f9596])).
fof(f9596,plain,(
  ( ! [X162] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X162) | p(s(bool,X162))) ) | $spl23),
  inference(subsumption_resolution,[],[f9503,f900])).
fof(f900,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f618])).
fof(f618,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f617])).
fof(f617,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X19] : ((p(s(bool,X19)) | p(s(bool,'F_0'))) <=> p(s(bool,X19)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.OR_CLAUSES|split|3')).
fof(f9503,plain,(
  ( ! [X162] : (p(s(bool,'F_0')) | p(s(bool,X162)) | ~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X162)) ) | $spl23),
  inference(resolution,[],[f1357,f798])).
fof(f798,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f556])).
fof(f33672,plain,(
  $false | ($spl17 | $spl19 | $spl36)),
  inference(subsumption_resolution,[],[f33671,f942])).
fof(f33671,plain,(
  p(s(bool,'F_0')) | ($spl17 | $spl19 | $spl36)),
  inference(forward_demodulation,[],[f33619,f7317])).
fof(f33619,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ($spl17 | $spl36)),
  inference(resolution,[],[f1652,f5139])).
fof(f5139,plain,(
  ( ! [X162] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X162) | p(s(bool,X162))) ) | $spl17),
  inference(subsumption_resolution,[],[f5041,f900])).
fof(f5041,plain,(
  ( ! [X162] : (p(s(bool,'F_0')) | p(s(bool,X162)) | ~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X162)) ) | $spl17),
  inference(resolution,[],[f1174,f798])).
fof(f33670,plain,(
  $false | ($spl1 | $spl19 | $spl36)),
  inference(subsumption_resolution,[],[f33669,f942])).
fof(f33669,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl19 | $spl36)),
  inference(forward_demodulation,[],[f33618,f7317])).
fof(f33618,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ($spl1 | $spl36)),
  inference(resolution,[],[f1652,f4946])).
fof(f4946,plain,(
  ( ! [X162] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X162) | p(s(bool,X162))) ) | $spl1),
  inference(subsumption_resolution,[],[f4853,f900])).
fof(f4853,plain,(
  ( ! [X162] : (p(s(bool,'F_0')) | p(s(bool,X162)) | ~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X162)) ) | $spl1),
  inference(resolution,[],[f1161,f798])).
fof(f33668,plain,(
  $false | ($spl19 | $spl31 | $spl36)),
  inference(subsumption_resolution,[],[f33667,f942])).
fof(f33667,plain,(
  p(s(bool,'F_0')) | ($spl19 | $spl31 | $spl36)),
  inference(forward_demodulation,[],[f33617,f7317])).
fof(f33617,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ($spl31 | $spl36)),
  inference(resolution,[],[f1652,f10194])).
fof(f10194,plain,(
  ( ! [X162] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X162) | p(s(bool,X162))) ) | $spl31),
  inference(subsumption_resolution,[],[f10097,f900])).
fof(f10097,plain,(
  ( ! [X162] : (p(s(bool,'F_0')) | p(s(bool,X162)) | ~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X162)) ) | $spl31),
  inference(resolution,[],[f1365,f798])).
fof(f33666,plain,(
  $false | ($spl19 | $spl36)),
  inference(subsumption_resolution,[],[f33665,f942])).
fof(f33665,plain,(
  p(s(bool,'F_0')) | ($spl19 | $spl36)),
  inference(forward_demodulation,[],[f33616,f7317])).
fof(f33616,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl36),
  inference(resolution,[],[f1652,f2588])).
fof(f2588,plain,(
  ( ! [X162] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X162) | p(s(bool,X162))) )),
  inference(subsumption_resolution,[],[f2492,f900])).
fof(f2492,plain,(
  ( ! [X162] : (p(s(bool,'F_0')) | p(s(bool,X162)) | ~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X162)) )),
  inference(resolution,[],[f655,f798])).
fof(f33664,plain,(
  $false | ($spl19 | $spl36)),
  inference(subsumption_resolution,[],[f33663,f942])).
fof(f33663,plain,(
  p(s(bool,'F_0')) | ($spl19 | $spl36)),
  inference(forward_demodulation,[],[f33615,f7317])).
fof(f33615,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ($spl19 | $spl36)),
  inference(resolution,[],[f1652,f7402])).
fof(f7402,plain,(
  ( ! [X162] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X162) | p(s(bool,X162))) ) | $spl19),
  inference(subsumption_resolution,[],[f7305,f900])).
fof(f7305,plain,(
  ( ! [X162] : (p(s(bool,'F_0')) | p(s(bool,X162)) | ~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X162)) ) | $spl19),
  inference(resolution,[],[f1175,f798])).
fof(f33662,plain,(
  $false | ($spl19 | $spl36)),
  inference(subsumption_resolution,[],[f33661,f942])).
fof(f33661,plain,(
  p(s(bool,'F_0')) | ($spl19 | $spl36)),
  inference(forward_demodulation,[],[f33614,f7317])).
fof(f33614,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl36),
  inference(resolution,[],[f1652,f1982])).
fof(f1982,plain,(
  ( ! [X162] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X162) | p(s(bool,X162))) )),
  inference(subsumption_resolution,[],[f1884,f900])).
fof(f1884,plain,(
  ( ! [X162] : (p(s(bool,'F_0')) | p(s(bool,X162)) | ~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X162)) )),
  inference(resolution,[],[f654,f798])).
fof(f1652,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl36),
  inference(cnf_transformation,[],[f1652_D])).
fof(f1652_D,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f33594,plain,(
  ( ! [X31,X32] : (sP11(X32,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13))) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12)),s('type.list.list'(sK12),sK39('const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13)),sK12))))) )),
  inference(superposition,[],[f27752,f977])).
fof(f27752,plain,(
  ( ! [X246,X245] : (sP11(X245,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X246),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f799])).
fof(f33588,plain,(
  ( ! [X10,X9] : (sP11(X10,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X9),s('type.list.list'(sK12),sK13))),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f27752,f863])).
fof(f33608,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl738),
  inference(cnf_transformation,[],[f33608_D])).
fof(f33608_D,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl738),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl738])])).
fof(f33576,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27752,f977])).
fof(f33547,plain,(
  ( ! [X26,X27] : (sP11(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X27),s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27752,f977])).
fof(f33530,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X26,X27)) ) | $spl734),
  inference(cnf_transformation,[],[f33530_D])).
fof(f33530_D,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X26,X27)) ) <=> ~$spl734),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl734])])).
fof(f33504,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27534,f977])).
fof(f27534,plain,(
  ( ! [X136,X137] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X136,X137)) )),
  inference(resolution,[],[f1192,f749])).
fof(f1192,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))))),
  inference(resolution,[],[f653,f960])).
fof(f960,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f479])).
fof(f479,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))))),
  inference(ennf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))))),
  inference(flattening,[],[f391])).
fof(f391,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X11,X18] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X11),s('type.list.list'(X6),X18)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X11))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X18)))))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_LENGTH')).
fof(f32969,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X26,X27)) ) | $spl712),
  inference(cnf_transformation,[],[f32969_D])).
fof(f32969_D,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X26,X27)) ) <=> ~$spl712),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl712])])).
fof(f33478,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27534,f977])).
fof(f33527,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | $spl733),
  inference(cnf_transformation,[],[f33527_D])).
fof(f33527_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) <=> ~$spl733),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl733])])).
fof(f33470,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f27534,f23864])).
fof(f33464,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) | $spl730),
  inference(cnf_transformation,[],[f33464_D])).
fof(f33464_D,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) <=> ~$spl730),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl730])])).
fof(f33440,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27533,f977])).
fof(f27533,plain,(
  ( ! [X134,X135] : (sP6(X134,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X135)) )),
  inference(resolution,[],[f1192,f748])).
fof(f32903,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27)) ) | $spl706),
  inference(cnf_transformation,[],[f32903_D])).
fof(f32903_D,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27)) ) <=> ~$spl706),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl706])])).
fof(f33414,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27533,f977])).
fof(f33402,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl728),
  inference(cnf_transformation,[],[f33402_D])).
fof(f33402_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl728),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl728])])).
fof(f33378,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27532,f977])).
fof(f27532,plain,(
  ( ! [X132,X133] : (sP6(X132,X133,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f747])).
fof(f32840,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl702),
  inference(cnf_transformation,[],[f32840_D])).
fof(f32840_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl702),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl702])])).
fof(f33352,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27532,f977])).
fof(f33339,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl726),
  inference(cnf_transformation,[],[f33339_D])).
fof(f33339_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl726),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl726])])).
fof(f33316,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27531,f977])).
fof(f27531,plain,(
  ( ! [X130,X131] : (sP7(X130,X131,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f743])).
fof(f32776,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl698),
  inference(cnf_transformation,[],[f32776_D])).
fof(f32776_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl698),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl698])])).
fof(f33290,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27531,f977])).
fof(f33278,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl724),
  inference(cnf_transformation,[],[f33278_D])).
fof(f33278_D,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl724),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl724])])).
fof(f33255,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27528,f977])).
fof(f27528,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f739])).
fof(f32714,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl694),
  inference(cnf_transformation,[],[f32714_D])).
fof(f32714_D,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl694),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl694])])).
fof(f33229,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27528,f977])).
fof(f33217,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl722),
  inference(cnf_transformation,[],[f33217_D])).
fof(f33217_D,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl722),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl722])])).
fof(f33194,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27515,f977])).
fof(f27515,plain,(
  ( ! [X99,X98] : (sP4(X98,X99,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f722])).
fof(f32652,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl690),
  inference(cnf_transformation,[],[f32652_D])).
fof(f32652_D,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl690),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl690])])).
fof(f33168,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27515,f977])).
fof(f33154,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) | $spl720),
  inference(cnf_transformation,[],[f33154_D])).
fof(f33154_D,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) <=> ~$spl720),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl720])])).
fof(f33132,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27499,f977])).
fof(f27499,plain,(
  ( ! [X66,X67] : (sP2(X66,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X67)) )),
  inference(resolution,[],[f1192,f706])).
fof(f32589,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27)) ) | $spl686),
  inference(cnf_transformation,[],[f32589_D])).
fof(f32589_D,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27)) ) <=> ~$spl686),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl686])])).
fof(f33106,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27499,f977])).
fof(f33093,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl718),
  inference(cnf_transformation,[],[f33093_D])).
fof(f33093_D,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl718),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl718])])).
fof(f33071,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27498,f977])).
fof(f27498,plain,(
  ( ! [X64,X65] : (sP2(X64,X65,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f705])).
fof(f32527,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl682),
  inference(cnf_transformation,[],[f32527_D])).
fof(f32527_D,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl682),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl682])])).
fof(f33045,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27498,f977])).
fof(f33032,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) | $spl716),
  inference(cnf_transformation,[],[f33032_D])).
fof(f33032_D,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) <=> ~$spl716),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl716])])).
fof(f33010,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27485,f977])).
fof(f27485,plain,(
  ( ! [X39,X38] : (sP0(X38,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X39)) )),
  inference(resolution,[],[f1192,f689])).
fof(f32465,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27)) ) | $spl678),
  inference(cnf_transformation,[],[f32465_D])).
fof(f32465_D,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27)) ) <=> ~$spl678),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl678])])).
fof(f32984,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27485,f977])).
fof(f32971,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X26,X27)) ) | $spl714),
  inference(cnf_transformation,[],[f32971_D])).
fof(f32971_D,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X26,X27)) ) <=> ~$spl714),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl714])])).
fof(f32944,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27132,f977])).
fof(f27132,plain,(
  ( ! [X136,X137] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X136,X137)) )),
  inference(resolution,[],[f996,f749])).
fof(f996,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))))),
  inference(resolution,[],[f652,f960])).
fof(f32918,plain,(
  ( ! [X26,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27132,f977])).
fof(f32967,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | $spl711),
  inference(cnf_transformation,[],[f32967_D])).
fof(f32967_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) <=> ~$spl711),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl711])])).
fof(f32910,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f27132,f23864])).
fof(f32905,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) | $spl708),
  inference(cnf_transformation,[],[f32905_D])).
fof(f32905_D,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) <=> ~$spl708),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl708])])).
fof(f32880,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27131,f977])).
fof(f27131,plain,(
  ( ! [X134,X135] : (sP6(X134,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X135)) )),
  inference(resolution,[],[f996,f748])).
fof(f32854,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27131,f977])).
fof(f32842,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl704),
  inference(cnf_transformation,[],[f32842_D])).
fof(f32842_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl704),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl704])])).
fof(f32817,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27130,f977])).
fof(f27130,plain,(
  ( ! [X132,X133] : (sP6(X132,X133,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f747])).
fof(f32791,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27130,f977])).
fof(f32778,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl700),
  inference(cnf_transformation,[],[f32778_D])).
fof(f32778_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl700),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl700])])).
fof(f32754,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27129,f977])).
fof(f27129,plain,(
  ( ! [X130,X131] : (sP7(X130,X131,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f743])).
fof(f32728,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27129,f977])).
fof(f32716,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl696),
  inference(cnf_transformation,[],[f32716_D])).
fof(f32716_D,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl696),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl696])])).
fof(f32692,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27126,f977])).
fof(f27126,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f739])).
fof(f32666,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27126,f977])).
fof(f32654,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl692),
  inference(cnf_transformation,[],[f32654_D])).
fof(f32654_D,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl692),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl692])])).
fof(f32630,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27113,f977])).
fof(f27113,plain,(
  ( ! [X99,X98] : (sP4(X98,X99,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f722])).
fof(f32604,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27113,f977])).
fof(f32591,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) | $spl688),
  inference(cnf_transformation,[],[f32591_D])).
fof(f32591_D,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) <=> ~$spl688),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl688])])).
fof(f32568,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27097,f977])).
fof(f27097,plain,(
  ( ! [X66,X67] : (sP2(X66,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X67)) )),
  inference(resolution,[],[f996,f706])).
fof(f32542,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27097,f977])).
fof(f32529,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl684),
  inference(cnf_transformation,[],[f32529_D])).
fof(f32529_D,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl684),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl684])])).
fof(f32506,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27096,f977])).
fof(f27096,plain,(
  ( ! [X64,X65] : (sP2(X64,X65,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f705])).
fof(f32480,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27096,f977])).
fof(f32467,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) | $spl680),
  inference(cnf_transformation,[],[f32467_D])).
fof(f32467_D,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27)) ) <=> ~$spl680),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl680])])).
fof(f32444,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27083,f977])).
fof(f27083,plain,(
  ( ! [X39,X38] : (sP0(X38,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X39)) )),
  inference(resolution,[],[f996,f689])).
fof(f32418,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27083,f977])).
fof(f32404,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X19)) ) | $spl676),
  inference(cnf_transformation,[],[f32404_D])).
fof(f32404_D,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X19)) ) <=> ~$spl676),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl676])])).
fof(f32381,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27546,f977])).
fof(f27546,plain,(
  ( ! [X151] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X151)) )),
  inference(resolution,[],[f1192,f800])).
fof(f32269,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X19)) ) | $spl670),
  inference(cnf_transformation,[],[f32269_D])).
fof(f32269_D,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X19)) ) <=> ~$spl670),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl670])])).
fof(f32355,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X19) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27546,f977])).
fof(f32338,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl674),
  inference(cnf_transformation,[],[f32338_D])).
fof(f32338_D,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl674),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl674])])).
fof(f32315,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27545,f977])).
fof(f27545,plain,(
  ( ! [X150] : (sP11(X150,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f799])).
fof(f32202,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl666),
  inference(cnf_transformation,[],[f32202_D])).
fof(f32202_D,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl666),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl666])])).
fof(f32289,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f27545,f977])).
fof(f32271,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X19)) ) | $spl672),
  inference(cnf_transformation,[],[f32271_D])).
fof(f32271_D,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X19)) ) <=> ~$spl672),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl672])])).
fof(f32247,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))),X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27144,f977])).
fof(f27144,plain,(
  ( ! [X151] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X151)) )),
  inference(resolution,[],[f996,f800])).
fof(f32221,plain,(
  ( ! [X19] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))),X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27144,f977])).
fof(f32204,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) | $spl668),
  inference(cnf_transformation,[],[f32204_D])).
fof(f32204_D,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) ) <=> ~$spl668),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl668])])).
fof(f32180,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f27143,f977])).
fof(f27143,plain,(
  ( ! [X150] : (sP11(X150,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f799])).
fof(f32154,plain,(
  ( ! [X19] : (sP11(X19,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f27143,f977])).
fof(f31940,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19)))) ) | $spl664),
  inference(cnf_transformation,[],[f31940_D])).
fof(f31940_D,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19)))) ) <=> ~$spl664),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl664])])).
fof(f32079,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10021,f977])).
fof(f10021,plain,(
  ( ! [X11] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X11)))) ) | $spl31),
  inference(resolution,[],[f1365,f959])).
fof(f959,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) )),
  inference(cnf_transformation,[],[f645])).
fof(f25618,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),sK15) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19)))) ) | $spl552),
  inference(cnf_transformation,[],[f25618_D])).
fof(f25618_D,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),sK15) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19)))) ) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f32014,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),sK15) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9427,f977])).
fof(f9427,plain,(
  ( ! [X11] : (s('type.list.list'(sK12),sK15) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),X11)))) ) | $spl23),
  inference(resolution,[],[f1357,f959])).
fof(f31948,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7229,f977])).
fof(f7229,plain,(
  ( ! [X11] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X11)))) ) | $spl19),
  inference(resolution,[],[f1175,f959])).
fof(f31881,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4965,f977])).
fof(f4965,plain,(
  ( ! [X11] : (s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),X11)))) ) | $spl17),
  inference(resolution,[],[f1174,f959])).
fof(f25683,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19)))) ) | $spl554),
  inference(cnf_transformation,[],[f25683_D])).
fof(f25683_D,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19)))) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f31817,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4777,f977])).
fof(f4777,plain,(
  ( ! [X11] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),X11)))) ) | $spl1),
  inference(resolution,[],[f1161,f959])).
fof(f31808,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | ~p(s(bool,X19))) ) | $spl662),
  inference(cnf_transformation,[],[f31808_D])).
fof(f31808_D,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | ~p(s(bool,X19))) ) <=> ~$spl662),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl662])])).
fof(f31782,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f31632,f977])).
fof(f31632,plain,(
  ( ! [X3] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),'F_0',X3) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f30352,f733])).
fof(f30352,plain,(
  ( ! [X16] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X16,'F_0') | ~p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f30301,f923])).
fof(f30301,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,X16) | ~p(s(bool,X16)) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X16,'F_0')) )),
  inference(resolution,[],[f1259,f23629])).
fof(f31744,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),'F_0',X19) | ~p(s(bool,X19))) ) | $spl658),
  inference(cnf_transformation,[],[f31744_D])).
fof(f31744_D,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),'F_0',X19) | ~p(s(bool,X19))) ) <=> ~$spl658),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl658])])).
fof(f31766,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),'F_0',X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f31632,f977])).
fof(f31746,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | ~p(s(bool,X19))) ) | $spl660),
  inference(cnf_transformation,[],[f31746_D])).
fof(f31746_D,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | ~p(s(bool,X19))) ) <=> ~$spl660),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl660])])).
fof(f31719,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),'F_0',X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f31576,f977])).
fof(f31576,plain,(
  ( ! [X3] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),'F_0',X3) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f28277,f733])).
fof(f28277,plain,(
  ( ! [X16] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X16,'F_0') | ~p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f28226,f923])).
fof(f28226,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,X16) | ~p(s(bool,X16)) | ~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X16,'F_0')) )),
  inference(resolution,[],[f1063,f23629])).
fof(f31703,plain,(
  ( ! [X19] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),'F_0',X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f31576,f977])).
fof(f31683,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl31),
  inference(subsumption_resolution,[],[f31663,f9673])).
fof(f31663,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19,'F_0') | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f30352,f977])).
fof(f31625,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | ~p(s(bool,X19))) ) | $spl656),
  inference(cnf_transformation,[],[f31625_D])).
fof(f31625_D,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | ~p(s(bool,X19))) ) <=> ~$spl656),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl656])])).
fof(f31647,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f30352,f977])).
fof(f31628,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f31607,f5237])).
fof(f31607,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19,'F_0') | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f28277,f977])).
fof(f31591,plain,(
  ( ! [X19] : (~sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19,'F_0') | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f28277,f977])).
fof(f31502,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl654),
  inference(cnf_transformation,[],[f31502_D])).
fof(f31502_D,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl654),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl654])])).
fof(f31467,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl76),
  inference(superposition,[],[f11063,f977])).
fof(f11063,plain,(
  ( ! [X2,X3] : (sP3(X2,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X3) | ~p(s(bool,X2))) ) | $spl76),
  inference(resolution,[],[f2873,f980])).
fof(f980,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f700])).
fof(f700,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f511])).
fof(f31419,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl652),
  inference(cnf_transformation,[],[f31419_D])).
fof(f31419_D,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl652),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl652])])).
fof(f31384,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl72),
  inference(superposition,[],[f10979,f977])).
fof(f10979,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X2))) ) | $spl72),
  inference(resolution,[],[f2833,f980])).
fof(f31335,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl650),
  inference(cnf_transformation,[],[f31335_D])).
fof(f31335_D,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl650),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl650])])).
fof(f31293,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl66),
  inference(superposition,[],[f10935,f977])).
fof(f10935,plain,(
  ( ! [X2,X3] : (sP1(X2,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X3) | ~p(s(bool,X2))) ) | $spl66),
  inference(resolution,[],[f2751,f979])).
fof(f979,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f683])).
fof(f683,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f503])).
fof(f31239,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl648),
  inference(cnf_transformation,[],[f31239_D])).
fof(f31239_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl648),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl648])])).
fof(f31217,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2483,f977])).
fof(f31201,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2483,f977])).
fof(f31181,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl644),
  inference(cnf_transformation,[],[f31181_D])).
fof(f31181_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl644),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl644])])).
fof(f31159,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2482,f977])).
fof(f31143,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2482,f977])).
fof(f31119,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl640),
  inference(cnf_transformation,[],[f31119_D])).
fof(f31119_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl640),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl640])])).
fof(f31097,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1875,f977])).
fof(f31081,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1875,f977])).
fof(f31061,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl636),
  inference(cnf_transformation,[],[f31061_D])).
fof(f31061_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl636),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl636])])).
fof(f31039,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1874,f977])).
fof(f31023,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1874,f977])).
fof(f27466,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) | $spl608),
  inference(cnf_transformation,[],[f27466_D])).
fof(f27466_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) <=> ~$spl608),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl608])])).
fof(f30926,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10071,f977])).
fof(f10071,plain,(
  ( ! [X118,X119] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X118,X119) | p(s(bool,X118))) ) | $spl31),
  inference(resolution,[],[f1365,f726])).
fof(f27383,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) | $spl606),
  inference(cnf_transformation,[],[f27383_D])).
fof(f27383_D,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) <=> ~$spl606),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl606])])).
fof(f30861,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10070,f977])).
fof(f10070,plain,(
  ( ! [X116,X117] : (sP5(X116,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X117) | p(s(bool,X116))) ) | $spl31),
  inference(resolution,[],[f1365,f726])).
fof(f27322,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl604),
  inference(cnf_transformation,[],[f27322_D])).
fof(f27322_D,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f30797,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10069,f977])).
fof(f10069,plain,(
  ( ! [X114,X115] : (sP5(X114,X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X114))) ) | $spl31),
  inference(resolution,[],[f1365,f725])).
fof(f27064,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl596),
  inference(cnf_transformation,[],[f27064_D])).
fof(f27064_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f30696,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10068,f977])).
fof(f10068,plain,(
  ( ! [X112,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X112,X113) | p(s(bool,X113))) ) | $spl31),
  inference(resolution,[],[f1365,f725])).
fof(f26981,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl594),
  inference(cnf_transformation,[],[f26981_D])).
fof(f26981_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl594])])).
fof(f30610,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10062,f977])).
fof(f10062,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X101) | ~p(s(bool,X100))) ) | $spl31),
  inference(resolution,[],[f1365,f716])).
fof(f26910,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl592),
  inference(cnf_transformation,[],[f26910_D])).
fof(f26910_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f30519,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10061,f977])).
fof(f10061,plain,(
  ( ! [X99,X98] : (~sP5(X98,X99,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X98))) ) | $spl31),
  inference(resolution,[],[f1365,f715])).
fof(f26839,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl590),
  inference(cnf_transformation,[],[f26839_D])).
fof(f26839_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl590),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl590])])).
fof(f30440,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10046,f977])).
fof(f10046,plain,(
  ( ! [X68,X69] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X68,X69) | ~p(s(bool,X68))) ) | $spl31),
  inference(resolution,[],[f1365,f698])).
fof(f26774,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) | $spl588),
  inference(cnf_transformation,[],[f26774_D])).
fof(f26774_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f30376,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10045,f977])).
fof(f10045,plain,(
  ( ! [X66,X67] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X67,X66) | ~p(s(bool,X66))) ) | $spl31),
  inference(resolution,[],[f1365,f697])).
fof(f30357,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl632),
  inference(cnf_transformation,[],[f30357_D])).
fof(f30357_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl632),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl632])])).
fof(f30326,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1259,f977])).
fof(f28281,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl622),
  inference(cnf_transformation,[],[f28281_D])).
fof(f28281_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl622),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl622])])).
fof(f30310,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1259,f977])).
fof(f26709,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl586),
  inference(cnf_transformation,[],[f26709_D])).
fof(f26709_D,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl586),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl586])])).
fof(f30239,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10040,f977])).
fof(f10040,plain,(
  ( ! [X57,X56] : (sP1(X56,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X56))) ) | $spl31),
  inference(resolution,[],[f1365,f692])).
fof(f26641,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl584),
  inference(cnf_transformation,[],[f26641_D])).
fof(f26641_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl584),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl584])])).
fof(f30159,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10032,f977])).
fof(f10032,plain,(
  ( ! [X41,X40] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X40,X41) | ~p(s(bool,X40))) ) | $spl31),
  inference(resolution,[],[f1365,f681])).
fof(f26574,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl582),
  inference(cnf_transformation,[],[f26574_D])).
fof(f26574_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl582),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl582])])).
fof(f30093,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10031,f977])).
fof(f10031,plain,(
  ( ! [X39,X38] : (~sP1(X38,X39,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X38))) ) | $spl31),
  inference(resolution,[],[f1365,f680])).
fof(f26509,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl580),
  inference(cnf_transformation,[],[f26509_D])).
fof(f26509_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl580),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl580])])).
fof(f30008,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10030,f977])).
fof(f10030,plain,(
  ( ! [X37,X36] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X37,X36) | p(s(bool,X36))) ) | $spl31),
  inference(resolution,[],[f1365,f680])).
fof(f29985,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl630),
  inference(cnf_transformation,[],[f29985_D])).
fof(f29985_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl630),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl630])])).
fof(f29971,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1258,f977])).
fof(f27930,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl618),
  inference(cnf_transformation,[],[f27930_D])).
fof(f27930_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl618),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl618])])).
fof(f29955,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1258,f977])).
fof(f24826,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X26))) ) | $spl500),
  inference(cnf_transformation,[],[f24826_D])).
fof(f24826_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X26))) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f29873,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9477,f977])).
fof(f9477,plain,(
  ( ! [X118,X119] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X118,X119) | p(s(bool,X118))) ) | $spl23),
  inference(resolution,[],[f1357,f726])).
fof(f24766,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26))) ) | $spl496),
  inference(cnf_transformation,[],[f24766_D])).
fof(f24766_D,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f29815,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9476,f977])).
fof(f9476,plain,(
  ( ! [X116,X117] : (sP5(X116,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X117) | p(s(bool,X116))) ) | $spl23),
  inference(resolution,[],[f1357,f726])).
fof(f24710,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26))) ) | $spl492),
  inference(cnf_transformation,[],[f24710_D])).
fof(f24710_D,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f29758,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9475,f977])).
fof(f9475,plain,(
  ( ! [X114,X115] : (sP5(X114,X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | p(s(bool,X114))) ) | $spl23),
  inference(resolution,[],[f1357,f725])).
fof(f29726,plain,(
  ( ! [X19] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0'))))))) ) | $spl628),
  inference(cnf_transformation,[],[f29726_D])).
fof(f29726_D,plain,(
  ( ! [X19] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0'))))))) ) <=> ~$spl628),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl628])])).
fof(f29663,plain,(
  ( ! [X19] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0')))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1189,f977])).
fof(f29721,plain,(
  ( ! [X24,X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X24),s('type.list.list'(sK12),sK13)))))) ) | $spl626),
  inference(cnf_transformation,[],[f29721_D])).
fof(f29721_D,plain,(
  ( ! [X24,X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X24),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl626),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl626])])).
fof(f29624,plain,(
  ( ! [X271,X270,X272] : (p(s(bool,X270)) | p(s(bool,X271)) | s(bool,X271) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X272),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f978])).
fof(f29533,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X50),s('type.list.list'(sK12),sK13))))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f1189,f679])).
fof(f29530,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f678])).
fof(f29529,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X36),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f677])).
fof(f29528,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13))))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f1189,f675])).
fof(f29526,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X24),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1189,f673])).
fof(f24654,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27))) ) | $spl488),
  inference(cnf_transformation,[],[f24654_D])).
fof(f24654_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f29448,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9474,f977])).
fof(f9474,plain,(
  ( ! [X112,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X112,X113) | p(s(bool,X113))) ) | $spl23),
  inference(resolution,[],[f1357,f725])).
fof(f24595,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26))) ) | $spl484),
  inference(cnf_transformation,[],[f24595_D])).
fof(f24595_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f29373,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9468,f977])).
fof(f9468,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X101) | ~p(s(bool,X100))) ) | $spl23),
  inference(resolution,[],[f1357,f716])).
fof(f24548,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26))) ) | $spl480),
  inference(cnf_transformation,[],[f24548_D])).
fof(f24548_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26))) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f29294,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9467,f977])).
fof(f9467,plain,(
  ( ! [X99,X98] : (~sP5(X98,X99,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X98))) ) | $spl23),
  inference(resolution,[],[f1357,f715])).
fof(f24501,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26))) ) | $spl476),
  inference(cnf_transformation,[],[f24501_D])).
fof(f24501_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f29223,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9452,f977])).
fof(f9452,plain,(
  ( ! [X68,X69] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X68,X69) | ~p(s(bool,X68))) ) | $spl23),
  inference(resolution,[],[f1357,f698])).
fof(f24439,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X27))) ) | $spl472),
  inference(cnf_transformation,[],[f24439_D])).
fof(f24439_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f29164,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9451,f977])).
fof(f9451,plain,(
  ( ! [X66,X67] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X67,X66) | ~p(s(bool,X66))) ) | $spl23),
  inference(resolution,[],[f1357,f697])).
fof(f24377,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26))) ) | $spl468),
  inference(cnf_transformation,[],[f24377_D])).
fof(f24377_D,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f29101,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9446,f977])).
fof(f9446,plain,(
  ( ! [X57,X56] : (sP1(X56,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X56))) ) | $spl23),
  inference(resolution,[],[f1357,f692])).
fof(f24327,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26))) ) | $spl464),
  inference(cnf_transformation,[],[f24327_D])).
fof(f24327_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f29029,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9438,f977])).
fof(f9438,plain,(
  ( ! [X41,X40] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X40,X41) | ~p(s(bool,X40))) ) | $spl23),
  inference(resolution,[],[f1357,f681])).
fof(f24268,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26))) ) | $spl460),
  inference(cnf_transformation,[],[f24268_D])).
fof(f24268_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f28969,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9437,f977])).
fof(f9437,plain,(
  ( ! [X39,X38] : (~sP1(X38,X39,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | p(s(bool,X38))) ) | $spl23),
  inference(resolution,[],[f1357,f680])).
fof(f24205,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27))) ) | $spl456),
  inference(cnf_transformation,[],[f24205_D])).
fof(f24205_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f28892,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9436,f977])).
fof(f9436,plain,(
  ( ! [X37,X36] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X37,X36) | p(s(bool,X36))) ) | $spl23),
  inference(resolution,[],[f1357,f680])).
fof(f28802,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7279,f977])).
fof(f28741,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7278,f977])).
fof(f7278,plain,(
  ( ! [X116,X117] : (sP5(X116,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X117) | p(s(bool,X116))) ) | $spl19),
  inference(resolution,[],[f1175,f726])).
fof(f28679,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7277,f977])).
fof(f7277,plain,(
  ( ! [X114,X115] : (sP5(X114,X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X114))) ) | $spl19),
  inference(resolution,[],[f1175,f725])).
fof(f28590,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7276,f977])).
fof(f28516,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7270,f977])).
fof(f28439,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7269,f977])).
fof(f28366,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7254,f977])).
fof(f28302,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7253,f977])).
fof(f28283,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) | $spl624),
  inference(cnf_transformation,[],[f28283_D])).
fof(f28283_D,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27)) ) <=> ~$spl624),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl624])])).
fof(f28251,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1063,f977])).
fof(f28235,plain,(
  ( ! [X26,X27] : (~sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) = s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1063,f977])).
fof(f28170,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7248,f977])).
fof(f28096,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7240,f977])).
fof(f28032,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7239,f977])).
fof(f27953,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7238,f977])).
fof(f27932,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) | $spl620),
  inference(cnf_transformation,[],[f27932_D])).
fof(f27932_D,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl620),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl620])])).
fof(f27917,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1062,f977])).
fof(f27901,plain,(
  ( ! [X26,X27] : (sP9('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s(bool,X26) != s(bool,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1062,f977])).
fof(f27877,plain,(
  ( ! [X19] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0'))))))) ) | $spl616),
  inference(cnf_transformation,[],[f27877_D])).
fof(f27877_D,plain,(
  ( ! [X19] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0'))))))) ) <=> ~$spl616),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl616])])).
fof(f27814,plain,(
  ( ! [X19] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X19),s('type.list.list'(sK12),'const.list.NIL_0')))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f993,f977])).
fof(f27872,plain,(
  ( ! [X24,X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X24),s('type.list.list'(sK12),sK13)))))) ) | $spl614),
  inference(cnf_transformation,[],[f27872_D])).
fof(f27872_D,plain,(
  ( ! [X24,X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X24),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl614),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl614])])).
fof(f27775,plain,(
  ( ! [X271,X270,X272] : (p(s(bool,X270)) | p(s(bool,X271)) | s(bool,X271) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X272),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f978])).
fof(f27684,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X50),s('type.list.list'(sK12),sK13))))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f993,f679])).
fof(f27681,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X40),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f678])).
fof(f27680,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X36),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f677])).
fof(f27679,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X31),s('type.list.list'(sK12),sK13))))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f993,f675])).
fof(f27677,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.SNOC_2'(s(sK12,X24),s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f993,f673])).
fof(f27661,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) | $spl612),
  inference(cnf_transformation,[],[f27661_D])).
fof(f27661_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) <=> ~$spl612),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl612])])).
fof(f27604,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f1192,f977])).
fof(f27257,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) | $spl600),
  inference(cnf_transformation,[],[f27257_D])).
fof(f27257_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) <=> ~$spl600),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl600])])).
fof(f27578,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f1192,f977])).
fof(f27653,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl610),
  inference(cnf_transformation,[],[f27653_D])).
fof(f27653_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl610),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl610])])).
fof(f27568,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X161) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f978])).
fof(f27477,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1192,f679])).
fof(f27474,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f678])).
fof(f27473,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f677])).
fof(f27472,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1192,f675])).
fof(f27470,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1192,f673])).
fof(f27408,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5015,f977])).
fof(f5015,plain,(
  ( ! [X118,X119] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X118,X119) | p(s(bool,X118))) ) | $spl17),
  inference(resolution,[],[f1174,f726])).
fof(f27348,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5014,f977])).
fof(f5014,plain,(
  ( ! [X116,X117] : (sP5(X116,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X117) | p(s(bool,X116))) ) | $spl17),
  inference(resolution,[],[f1174,f726])).
fof(f27287,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5013,f977])).
fof(f5013,plain,(
  ( ! [X114,X115] : (sP5(X114,X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X114))) ) | $spl17),
  inference(resolution,[],[f1174,f725])).
fof(f27261,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) | $spl602),
  inference(cnf_transformation,[],[f27261_D])).
fof(f27261_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) <=> ~$spl602),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl602])])).
fof(f27202,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0')))))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f996,f977])).
fof(f27176,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f996,f977])).
fof(f27251,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl598),
  inference(cnf_transformation,[],[f27251_D])).
fof(f27251_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl598])])).
fof(f27166,plain,(
  ( ! [X161,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X161) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f978])).
fof(f27075,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f996,f679])).
fof(f27072,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f678])).
fof(f27071,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f677])).
fof(f27070,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f996,f675])).
fof(f27068,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f996,f673])).
fof(f27006,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5012,f977])).
fof(f5012,plain,(
  ( ! [X112,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X112,X113) | p(s(bool,X113))) ) | $spl17),
  inference(resolution,[],[f1174,f725])).
fof(f26937,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5006,f977])).
fof(f5006,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X101) | ~p(s(bool,X100))) ) | $spl17),
  inference(resolution,[],[f1174,f716])).
fof(f26866,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5005,f977])).
fof(f5005,plain,(
  ( ! [X99,X98] : (~sP5(X98,X99,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X98))) ) | $spl17),
  inference(resolution,[],[f1174,f715])).
fof(f26796,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4990,f977])).
fof(f4990,plain,(
  ( ! [X68,X69] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X68,X69) | ~p(s(bool,X68))) ) | $spl17),
  inference(resolution,[],[f1174,f698])).
fof(f26731,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4989,f977])).
fof(f4989,plain,(
  ( ! [X66,X67] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X67,X66) | ~p(s(bool,X66))) ) | $spl17),
  inference(resolution,[],[f1174,f697])).
fof(f26669,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4984,f977])).
fof(f4984,plain,(
  ( ! [X57,X56] : (sP1(X56,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X56))) ) | $spl17),
  inference(resolution,[],[f1174,f692])).
fof(f26597,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4976,f977])).
fof(f4976,plain,(
  ( ! [X41,X40] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X40,X41) | ~p(s(bool,X40))) ) | $spl17),
  inference(resolution,[],[f1174,f681])).
fof(f26533,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4975,f977])).
fof(f4975,plain,(
  ( ! [X39,X38] : (~sP1(X38,X39,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X38))) ) | $spl17),
  inference(resolution,[],[f1174,f680])).
fof(f26456,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4974,f977])).
fof(f4974,plain,(
  ( ! [X37,X36] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X37,X36) | p(s(bool,X36))) ) | $spl17),
  inference(resolution,[],[f1174,f680])).
fof(f26436,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) | $spl578),
  inference(cnf_transformation,[],[f26436_D])).
fof(f26436_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f26399,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4827,f977])).
fof(f4827,plain,(
  ( ! [X118,X119] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X118,X119) | p(s(bool,X118))) ) | $spl1),
  inference(resolution,[],[f1161,f726])).
fof(f25550,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X26))) ) | $spl548),
  inference(cnf_transformation,[],[f25550_D])).
fof(f25550_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X26))) ) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f26383,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4827,f977])).
fof(f26363,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) | $spl576),
  inference(cnf_transformation,[],[f26363_D])).
fof(f26363_D,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f26344,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4826,f977])).
fof(f4826,plain,(
  ( ! [X116,X117] : (sP5(X116,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X117) | p(s(bool,X116))) ) | $spl1),
  inference(resolution,[],[f1161,f726])).
fof(f25483,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | p(s(bool,X26))) ) | $spl544),
  inference(cnf_transformation,[],[f25483_D])).
fof(f25483_D,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | p(s(bool,X26))) ) <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f26328,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4826,f977])).
fof(f26306,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl574),
  inference(cnf_transformation,[],[f26306_D])).
fof(f26306_D,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f26287,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4825,f977])).
fof(f4825,plain,(
  ( ! [X114,X115] : (sP5(X114,X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X114))) ) | $spl1),
  inference(resolution,[],[f1161,f725])).
fof(f25426,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X26))) ) | $spl540),
  inference(cnf_transformation,[],[f25426_D])).
fof(f25426_D,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X26))) ) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f26271,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4825,f977])).
fof(f26249,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl572),
  inference(cnf_transformation,[],[f26249_D])).
fof(f26249_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f26212,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4824,f977])).
fof(f4824,plain,(
  ( ! [X112,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X112,X113) | p(s(bool,X113))) ) | $spl1),
  inference(resolution,[],[f1161,f725])).
fof(f25369,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X27))) ) | $spl536),
  inference(cnf_transformation,[],[f25369_D])).
fof(f25369_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X27))) ) <=> ~$spl536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl536])])).
fof(f26196,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4824,f977])).
fof(f26176,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl570),
  inference(cnf_transformation,[],[f26176_D])).
fof(f26176_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f26153,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4818,f977])).
fof(f4818,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X101) | ~p(s(bool,X100))) ) | $spl1),
  inference(resolution,[],[f1161,f716])).
fof(f25303,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | ~p(s(bool,X26))) ) | $spl532),
  inference(cnf_transformation,[],[f25303_D])).
fof(f25303_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | ~p(s(bool,X26))) ) <=> ~$spl532),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl532])])).
fof(f26137,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4818,f977])).
fof(f26116,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl568),
  inference(cnf_transformation,[],[f26116_D])).
fof(f26116_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f26093,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4817,f977])).
fof(f4817,plain,(
  ( ! [X99,X98] : (~sP5(X98,X99,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X98))) ) | $spl1),
  inference(resolution,[],[f1161,f715])).
fof(f25249,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X26))) ) | $spl528),
  inference(cnf_transformation,[],[f25249_D])).
fof(f25249_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X26))) ) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f26077,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4817,f977])).
fof(f26056,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl566),
  inference(cnf_transformation,[],[f26056_D])).
fof(f26056_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f26030,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4802,f977])).
fof(f4802,plain,(
  ( ! [X68,X69] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X68,X69) | ~p(s(bool,X68))) ) | $spl1),
  inference(resolution,[],[f1161,f698])).
fof(f25195,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X26))) ) | $spl524),
  inference(cnf_transformation,[],[f25195_D])).
fof(f25195_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl524),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl524])])).
fof(f26014,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4802,f977])).
fof(f25995,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) | $spl564),
  inference(cnf_transformation,[],[f25995_D])).
fof(f25995_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl564),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl564])])).
fof(f25969,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4801,f977])).
fof(f4801,plain,(
  ( ! [X66,X67] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X67,X66) | ~p(s(bool,X66))) ) | $spl1),
  inference(resolution,[],[f1161,f697])).
fof(f25133,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X27))) ) | $spl520),
  inference(cnf_transformation,[],[f25133_D])).
fof(f25133_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f25953,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4801,f977])).
fof(f25934,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl562),
  inference(cnf_transformation,[],[f25934_D])).
fof(f25934_D,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl562),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl562])])).
fof(f25914,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4796,f977])).
fof(f4796,plain,(
  ( ! [X57,X56] : (sP1(X56,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X56))) ) | $spl1),
  inference(resolution,[],[f1161,f692])).
fof(f25071,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X26))) ) | $spl516),
  inference(cnf_transformation,[],[f25071_D])).
fof(f25071_D,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X26))) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f25898,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4796,f977])).
fof(f25875,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl560),
  inference(cnf_transformation,[],[f25875_D])).
fof(f25875_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl560),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl560])])).
fof(f25849,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4788,f977])).
fof(f4788,plain,(
  ( ! [X41,X40] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X40,X41) | ~p(s(bool,X40))) ) | $spl1),
  inference(resolution,[],[f1161,f681])).
fof(f25016,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X26))) ) | $spl512),
  inference(cnf_transformation,[],[f25016_D])).
fof(f25016_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f25833,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4788,f977])).
fof(f25814,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl558),
  inference(cnf_transformation,[],[f25814_D])).
fof(f25814_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl558),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl558])])).
fof(f25789,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4787,f977])).
fof(f4787,plain,(
  ( ! [X39,X38] : (~sP1(X38,X39,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X38))) ) | $spl1),
  inference(resolution,[],[f1161,f680])).
fof(f24955,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X26))) ) | $spl508),
  inference(cnf_transformation,[],[f24955_D])).
fof(f24955_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X26))) ) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f25773,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4787,f977])).
fof(f25752,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl556),
  inference(cnf_transformation,[],[f25752_D])).
fof(f25752_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl556),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl556])])).
fof(f25718,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4786,f977])).
fof(f4786,plain,(
  ( ! [X37,X36] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X37,X36) | p(s(bool,X36))) ) | $spl1),
  inference(resolution,[],[f1161,f680])).
fof(f24892,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X27))) ) | $spl504),
  inference(cnf_transformation,[],[f24892_D])).
fof(f24892_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X27))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f25702,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4786,f977])).
fof(f25626,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2416,f977])).
fof(f2416,plain,(
  ( ! [X11] : (s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X11)))) )),
  inference(resolution,[],[f655,f959])).
fof(f25561,plain,(
  ( ! [X19] : (s('type.list.list'(sK12),sK15) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),X19))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1808,f977])).
fof(f1808,plain,(
  ( ! [X11] : (s('type.list.list'(sK12),sK15) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X11)))) )),
  inference(resolution,[],[f654,f959])).
fof(f25553,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) | $spl550),
  inference(cnf_transformation,[],[f25553_D])).
fof(f25553_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f25520,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2466,f977])).
fof(f2466,plain,(
  ( ! [X118,X119] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X118,X119) | p(s(bool,X118))) )),
  inference(resolution,[],[f655,f726])).
fof(f25504,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2466,f977])).
fof(f25486,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) | $spl546),
  inference(cnf_transformation,[],[f25486_D])).
fof(f25486_D,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f25466,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2465,f977])).
fof(f2465,plain,(
  ( ! [X116,X117] : (sP5(X116,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X117) | p(s(bool,X116))) )),
  inference(resolution,[],[f655,f726])).
fof(f25450,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2465,f977])).
fof(f25429,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl542),
  inference(cnf_transformation,[],[f25429_D])).
fof(f25429_D,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f25409,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2464,f977])).
fof(f2464,plain,(
  ( ! [X114,X115] : (sP5(X114,X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | p(s(bool,X114))) )),
  inference(resolution,[],[f655,f725])).
fof(f25393,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2464,f977])).
fof(f25372,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl538),
  inference(cnf_transformation,[],[f25372_D])).
fof(f25372_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f25339,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2463,f977])).
fof(f2463,plain,(
  ( ! [X112,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X112,X113) | p(s(bool,X113))) )),
  inference(resolution,[],[f655,f725])).
fof(f25323,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2463,f977])).
fof(f25305,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl534),
  inference(cnf_transformation,[],[f25305_D])).
fof(f25305_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl534])])).
fof(f25285,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2457,f977])).
fof(f2457,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X101) | ~p(s(bool,X100))) )),
  inference(resolution,[],[f655,f716])).
fof(f25269,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2457,f977])).
fof(f25251,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl530),
  inference(cnf_transformation,[],[f25251_D])).
fof(f25251_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl530),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl530])])).
fof(f25231,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2456,f977])).
fof(f2456,plain,(
  ( ! [X99,X98] : (~sP5(X98,X99,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X98))) )),
  inference(resolution,[],[f655,f715])).
fof(f25215,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2456,f977])).
fof(f25197,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl526),
  inference(cnf_transformation,[],[f25197_D])).
fof(f25197_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f25170,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2441,f977])).
fof(f2441,plain,(
  ( ! [X68,X69] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X68,X69) | ~p(s(bool,X68))) )),
  inference(resolution,[],[f655,f698])).
fof(f25154,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2441,f977])).
fof(f25135,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) | $spl522),
  inference(cnf_transformation,[],[f25135_D])).
fof(f25135_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f25108,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2440,f977])).
fof(f2440,plain,(
  ( ! [X66,X67] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X67,X66) | ~p(s(bool,X66))) )),
  inference(resolution,[],[f655,f697])).
fof(f25092,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2440,f977])).
fof(f25073,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl518),
  inference(cnf_transformation,[],[f25073_D])).
fof(f25073_D,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f25055,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2435,f977])).
fof(f2435,plain,(
  ( ! [X57,X56] : (sP1(X56,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X56))) )),
  inference(resolution,[],[f655,f692])).
fof(f25039,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2435,f977])).
fof(f25018,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl514),
  inference(cnf_transformation,[],[f25018_D])).
fof(f25018_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f24992,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2427,f977])).
fof(f2427,plain,(
  ( ! [X41,X40] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X40,X41) | ~p(s(bool,X40))) )),
  inference(resolution,[],[f655,f681])).
fof(f24976,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2427,f977])).
fof(f24958,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl510),
  inference(cnf_transformation,[],[f24958_D])).
fof(f24958_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f24932,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2426,f977])).
fof(f2426,plain,(
  ( ! [X39,X38] : (~sP1(X38,X39,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | p(s(bool,X38))) )),
  inference(resolution,[],[f655,f680])).
fof(f24916,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2426,f977])).
fof(f24895,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl506),
  inference(cnf_transformation,[],[f24895_D])).
fof(f24895_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f24863,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2425,f977])).
fof(f2425,plain,(
  ( ! [X37,X36] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X37,X36) | p(s(bool,X36))) )),
  inference(resolution,[],[f655,f680])).
fof(f24847,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2425,f977])).
fof(f24829,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) | $spl502),
  inference(cnf_transformation,[],[f24829_D])).
fof(f24829_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f24801,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1858,f977])).
fof(f1858,plain,(
  ( ! [X118,X119] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X118,X119) | p(s(bool,X118))) )),
  inference(resolution,[],[f654,f726])).
fof(f24785,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1858,f977])).
fof(f24769,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) | $spl498),
  inference(cnf_transformation,[],[f24769_D])).
fof(f24769_D,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f24749,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1857,f977])).
fof(f1857,plain,(
  ( ! [X116,X117] : (sP5(X116,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X117) | p(s(bool,X116))) )),
  inference(resolution,[],[f654,f726])).
fof(f24733,plain,(
  ( ! [X26,X27] : (sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1857,f977])).
fof(f24713,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl494),
  inference(cnf_transformation,[],[f24713_D])).
fof(f24713_D,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f24693,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1856,f977])).
fof(f1856,plain,(
  ( ! [X114,X115] : (sP5(X114,X115,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | p(s(bool,X114))) )),
  inference(resolution,[],[f654,f725])).
fof(f24677,plain,(
  ( ! [X26,X27] : (sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1856,f977])).
fof(f24657,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl490),
  inference(cnf_transformation,[],[f24657_D])).
fof(f24657_D,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f24629,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1855,f977])).
fof(f1855,plain,(
  ( ! [X112,X113] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X112,X113) | p(s(bool,X113))) )),
  inference(resolution,[],[f654,f725])).
fof(f24613,plain,(
  ( ! [X26,X27] : (sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1855,f977])).
fof(f24597,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl486),
  inference(cnf_transformation,[],[f24597_D])).
fof(f24597_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f24581,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1849,f977])).
fof(f1849,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X101) | ~p(s(bool,X100))) )),
  inference(resolution,[],[f654,f716])).
fof(f24565,plain,(
  ( ! [X26,X27] : (~sP5(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1849,f977])).
fof(f24550,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl482),
  inference(cnf_transformation,[],[f24550_D])).
fof(f24550_D,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f24534,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1848,f977])).
fof(f1848,plain,(
  ( ! [X99,X98] : (~sP5(X98,X99,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X98))) )),
  inference(resolution,[],[f654,f715])).
fof(f24518,plain,(
  ( ! [X26,X27] : (~sP5(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1848,f977])).
fof(f24503,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl478),
  inference(cnf_transformation,[],[f24503_D])).
fof(f24503_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f24476,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1833,f977])).
fof(f1833,plain,(
  ( ! [X68,X69] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X68,X69) | ~p(s(bool,X68))) )),
  inference(resolution,[],[f654,f698])).
fof(f24460,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1833,f977])).
fof(f24441,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) | $spl474),
  inference(cnf_transformation,[],[f24441_D])).
fof(f24441_D,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f24414,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1832,f977])).
fof(f1832,plain,(
  ( ! [X66,X67] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X67,X66) | ~p(s(bool,X66))) )),
  inference(resolution,[],[f654,f697])).
fof(f24398,plain,(
  ( ! [X26,X27] : (~sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1832,f977])).
fof(f24379,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl470),
  inference(cnf_transformation,[],[f24379_D])).
fof(f24379_D,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f24364,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1827,f977])).
fof(f1827,plain,(
  ( ! [X57,X56] : (sP1(X56,X57,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X56))) )),
  inference(resolution,[],[f654,f692])).
fof(f24348,plain,(
  ( ! [X26,X27] : (sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1827,f977])).
fof(f24329,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl466),
  inference(cnf_transformation,[],[f24329_D])).
fof(f24329_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f24304,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1819,f977])).
fof(f1819,plain,(
  ( ! [X41,X40] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X40,X41) | ~p(s(bool,X40))) )),
  inference(resolution,[],[f654,f681])).
fof(f24288,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1819,f977])).
fof(f24271,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl462),
  inference(cnf_transformation,[],[f24271_D])).
fof(f24271_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f24245,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1818,f977])).
fof(f1818,plain,(
  ( ! [X39,X38] : (~sP1(X38,X39,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | p(s(bool,X38))) )),
  inference(resolution,[],[f654,f680])).
fof(f24229,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1818,f977])).
fof(f24208,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl458),
  inference(cnf_transformation,[],[f24208_D])).
fof(f24208_D,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f24179,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1817,f977])).
fof(f1817,plain,(
  ( ! [X37,X36] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X37,X36) | p(s(bool,X36))) )),
  inference(resolution,[],[f654,f680])).
fof(f24163,plain,(
  ( ! [X26,X27] : (~sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1817,f977])).
fof(f24065,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) | $spl454),
  inference(cnf_transformation,[],[f24065_D])).
fof(f24065_D,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f24044,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1231,f977])).
fof(f1231,plain,(
  ( ! [X94,X93] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X93,X94) | p(s(bool,X93))) )),
  inference(resolution,[],[f653,f716])).
fof(f21194,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26))) ) | $spl426),
  inference(cnf_transformation,[],[f21194_D])).
fof(f21194_D,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f24028,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1231,f977])).
fof(f24013,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl452),
  inference(cnf_transformation,[],[f24013_D])).
fof(f24013_D,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f23992,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1230,f977])).
fof(f1230,plain,(
  ( ! [X92,X91] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X92,X91) | p(s(bool,X91))) )),
  inference(resolution,[],[f653,f715])).
fof(f20974,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27))) ) | $spl422),
  inference(cnf_transformation,[],[f20974_D])).
fof(f20974_D,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f23976,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1230,f977])).
fof(f23961,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl450),
  inference(cnf_transformation,[],[f23961_D])).
fof(f23961_D,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f23939,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1226,f977])).
fof(f1226,plain,(
  ( ! [X83,X84] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X83,X84) | ~p(s(bool,X83))) )),
  inference(resolution,[],[f653,f710])).
fof(f710,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f516])).
fof(f19929,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26))) ) | $spl418),
  inference(cnf_transformation,[],[f19929_D])).
fof(f19929_D,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f23923,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1226,f977])).
fof(f23844,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl448),
  inference(cnf_transformation,[],[f23844_D])).
fof(f23844_D,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f23831,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1225,f977])).
fof(f1225,plain,(
  ( ! [X81,X82] : (sP3(X81,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X82) | ~p(s(bool,X81))) )),
  inference(resolution,[],[f653,f710])).
fof(f19714,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26))) ) | $spl414),
  inference(cnf_transformation,[],[f19714_D])).
fof(f19714_D,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f23815,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1225,f977])).
fof(f23798,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) | $spl446),
  inference(cnf_transformation,[],[f23798_D])).
fof(f23798_D,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f23776,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1224,f977])).
fof(f1224,plain,(
  ( ! [X80,X79] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X79,X80) | ~p(s(bool,X80))) )),
  inference(resolution,[],[f653,f709])).
fof(f709,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f516])).
fof(f19669,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27))) ) | $spl410),
  inference(cnf_transformation,[],[f19669_D])).
fof(f19669_D,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f23760,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1224,f977])).
fof(f23685,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl444),
  inference(cnf_transformation,[],[f23685_D])).
fof(f23685_D,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f23672,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1223,f977])).
fof(f1223,plain,(
  ( ! [X78,X77] : (sP3(X77,X78,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X77))) )),
  inference(resolution,[],[f653,f709])).
fof(f19619,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26))) ) | $spl406),
  inference(cnf_transformation,[],[f19619_D])).
fof(f19619_D,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f23656,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1223,f977])).
fof(f23510,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) | $spl442),
  inference(cnf_transformation,[],[f23510_D])).
fof(f23510_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f23490,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1217,f977])).
fof(f1217,plain,(
  ( ! [X66,X65] : (~sP3(X65,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X66) | p(s(bool,X65))) )),
  inference(resolution,[],[f653,f698])).
fof(f19573,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26))) ) | $spl402),
  inference(cnf_transformation,[],[f19573_D])).
fof(f19573_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f23474,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1217,f977])).
fof(f23443,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl440),
  inference(cnf_transformation,[],[f23443_D])).
fof(f23443_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f23442,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f23437])).
fof(f23437,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f12503,f753])).
fof(f12503,plain,(
  ( ! [X4,X5] : (sP9(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) )),
  inference(superposition,[],[f12497,f750])).
fof(f23435,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl438),
  inference(cnf_transformation,[],[f23435_D])).
fof(f23435_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f23415,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1216,f977])).
fof(f1216,plain,(
  ( ! [X64,X63] : (~sP3(X63,X64,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X63))) )),
  inference(resolution,[],[f653,f697])).
fof(f19410,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26))) ) | $spl394),
  inference(cnf_transformation,[],[f19410_D])).
fof(f19410_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f23399,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1216,f977])).
fof(f23381,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl436),
  inference(cnf_transformation,[],[f23381_D])).
fof(f23381_D,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f23361,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1212,f977])).
fof(f1212,plain,(
  ( ! [X56,X55] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X55,X56) | ~p(s(bool,X55))) )),
  inference(resolution,[],[f653,f693])).
fof(f693,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f508])).
fof(f18480,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26))) ) | $spl386),
  inference(cnf_transformation,[],[f18480_D])).
fof(f18480_D,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f23345,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1212,f977])).
fof(f19525,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl400),
  inference(cnf_transformation,[],[f19525_D])).
fof(f19525_D,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f23289,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10195,f977])).
fof(f10195,plain,(
  ( ! [X163] : (~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X163))) ) | $spl31),
  inference(subsumption_resolution,[],[f10098,f900])).
fof(f10098,plain,(
  ( ! [X163] : (p(s(bool,'F_0')) | p(s(bool,X163)) | ~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f798])).
fof(f23272,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl434),
  inference(cnf_transformation,[],[f23272_D])).
fof(f23272_D,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f23259,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1211,f977])).
fof(f1211,plain,(
  ( ! [X54,X53] : (sP1(X53,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X54) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f653,f693])).
fof(f18276,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26))) ) | $spl376),
  inference(cnf_transformation,[],[f18276_D])).
fof(f18276_D,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f23243,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1211,f977])).
fof(f19467,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl398),
  inference(cnf_transformation,[],[f19467_D])).
fof(f19467_D,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f23189,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10194,f977])).
fof(f23084,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl432),
  inference(cnf_transformation,[],[f23084_D])).
fof(f23084_D,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f23063,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1210,f977])).
fof(f1210,plain,(
  ( ! [X52,X51] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X51,X52) | p(s(bool,X52))) )),
  inference(resolution,[],[f653,f692])).
fof(f18100,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27))) ) | $spl368),
  inference(cnf_transformation,[],[f18100_D])).
fof(f18100_D,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f23047,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1210,f977])).
fof(f22349,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) | $spl430),
  inference(cnf_transformation,[],[f22349_D])).
fof(f22349_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f22331,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1201,f977])).
fof(f1201,plain,(
  ( ! [X33,X34] : (~sP1(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X34) | p(s(bool,X33))) )),
  inference(resolution,[],[f653,f681])).
fof(f17378,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26))) ) | $spl362),
  inference(cnf_transformation,[],[f17378_D])).
fof(f17378_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f22315,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1201,f977])).
fof(f18592,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) | $spl392),
  inference(cnf_transformation,[],[f18592_D])).
fof(f18592_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f22255,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10092,f977])).
fof(f10092,plain,(
  ( ! [X157] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X157) | ~p(s(bool,X157))) ) | $spl31),
  inference(resolution,[],[f1365,f787])).
fof(f787,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f550])).
fof(f550,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f549])).
fof(f549,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f494])).
fof(f494,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f18535,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) | $spl390),
  inference(cnf_transformation,[],[f18535_D])).
fof(f18535_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f22190,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10091,f977])).
fof(f10091,plain,(
  ( ! [X156] : (sP10(X156,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X156))) ) | $spl31),
  inference(resolution,[],[f1365,f784])).
fof(f784,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f551])).
fof(f18431,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl384),
  inference(cnf_transformation,[],[f18431_D])).
fof(f18431_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f22133,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10090,f977])).
fof(f10090,plain,(
  ( ! [X155] : (~sP10(X155,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X155))) ) | $spl31),
  inference(resolution,[],[f1365,f782])).
fof(f782,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f551])).
fof(f18378,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl382),
  inference(cnf_transformation,[],[f18378_D])).
fof(f18378_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f22078,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f10089,f977])).
fof(f10089,plain,(
  ( ! [X154] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X154) | p(s(bool,X154))) ) | $spl31),
  inference(resolution,[],[f1365,f782])).
fof(f16085,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X19))) ) | $spl326),
  inference(cnf_transformation,[],[f16085_D])).
fof(f16085_D,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X19))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f22023,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9597,f977])).
fof(f9597,plain,(
  ( ! [X163] : (~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | p(s(bool,X163))) ) | $spl23),
  inference(subsumption_resolution,[],[f9504,f900])).
fof(f9504,plain,(
  ( ! [X163] : (p(s(bool,'F_0')) | p(s(bool,X163)) | ~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f798])).
fof(f16033,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | p(s(bool,X19))) ) | $spl322),
  inference(cnf_transformation,[],[f16033_D])).
fof(f16033_D,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | p(s(bool,X19))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f21972,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9596,f977])).
fof(f21197,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) | $spl428),
  inference(cnf_transformation,[],[f21197_D])).
fof(f21197_D,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f21175,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1035,f977])).
fof(f1035,plain,(
  ( ! [X94,X93] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X93,X94) | p(s(bool,X93))) )),
  inference(resolution,[],[f652,f716])).
fof(f21159,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1035,f977])).
fof(f15229,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | ~p(s(bool,X19))) ) | $spl318),
  inference(cnf_transformation,[],[f15229_D])).
fof(f15229_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | ~p(s(bool,X19))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f21107,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9498,f977])).
fof(f9498,plain,(
  ( ! [X157] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X157) | ~p(s(bool,X157))) ) | $spl23),
  inference(resolution,[],[f1357,f787])).
fof(f15185,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X19))) ) | $spl314),
  inference(cnf_transformation,[],[f15185_D])).
fof(f15185_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X19))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f21049,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9497,f977])).
fof(f9497,plain,(
  ( ! [X156] : (sP10(X156,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X156))) ) | $spl23),
  inference(resolution,[],[f1357,f784])).
fof(f15141,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X19))) ) | $spl310),
  inference(cnf_transformation,[],[f15141_D])).
fof(f15141_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X19))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f20996,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9496,f977])).
fof(f9496,plain,(
  ( ! [X155] : (~sP10(X155,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | p(s(bool,X155))) ) | $spl23),
  inference(resolution,[],[f1357,f782])).
fof(f20977,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl424),
  inference(cnf_transformation,[],[f20977_D])).
fof(f20977_D,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f20955,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1034,f977])).
fof(f1034,plain,(
  ( ! [X92,X91] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X92,X91) | p(s(bool,X91))) )),
  inference(resolution,[],[f652,f715])).
fof(f20939,plain,(
  ( ! [X26,X27] : (~sP5('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1034,f977])).
fof(f15095,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | p(s(bool,X19))) ) | $spl306),
  inference(cnf_transformation,[],[f15095_D])).
fof(f15095_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | p(s(bool,X19))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f20892,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f9495,f977])).
fof(f9495,plain,(
  ( ! [X154] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X154) | p(s(bool,X154))) ) | $spl23),
  inference(resolution,[],[f1357,f782])).
fof(f20833,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7403,f977])).
fof(f7403,plain,(
  ( ! [X163] : (~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X163))) ) | $spl19),
  inference(subsumption_resolution,[],[f7306,f900])).
fof(f7306,plain,(
  ( ! [X163] : (p(s(bool,'F_0')) | p(s(bool,X163)) | ~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f798])).
fof(f20777,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7402,f977])).
fof(f19951,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7300,f977])).
fof(f7300,plain,(
  ( ! [X157] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X157) | ~p(s(bool,X157))) ) | $spl19),
  inference(resolution,[],[f1175,f787])).
fof(f19931,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl420),
  inference(cnf_transformation,[],[f19931_D])).
fof(f19931_D,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f19913,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1030,f977])).
fof(f1030,plain,(
  ( ! [X83,X84] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X83,X84) | ~p(s(bool,X83))) )),
  inference(resolution,[],[f652,f710])).
fof(f19897,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1030,f977])).
fof(f19842,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7299,f977])).
fof(f7299,plain,(
  ( ! [X156] : (sP10(X156,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X156))) ) | $spl19),
  inference(resolution,[],[f1175,f784])).
fof(f19787,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7298,f977])).
fof(f7298,plain,(
  ( ! [X155] : (~sP10(X155,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X155))) ) | $spl19),
  inference(resolution,[],[f1175,f782])).
fof(f19734,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f7297,f977])).
fof(f7297,plain,(
  ( ! [X154] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X154) | p(s(bool,X154))) ) | $spl19),
  inference(resolution,[],[f1175,f782])).
fof(f19716,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl416),
  inference(cnf_transformation,[],[f19716_D])).
fof(f19716_D,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f19703,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1029,f977])).
fof(f1029,plain,(
  ( ! [X81,X82] : (sP3(X81,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X82) | ~p(s(bool,X81))) )),
  inference(resolution,[],[f652,f710])).
fof(f19687,plain,(
  ( ! [X26,X27] : (sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1029,f977])).
fof(f19671,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) | $spl412),
  inference(cnf_transformation,[],[f19671_D])).
fof(f19671_D,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f19653,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1028,f977])).
fof(f1028,plain,(
  ( ! [X80,X79] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X79,X80) | ~p(s(bool,X80))) )),
  inference(resolution,[],[f652,f709])).
fof(f19637,plain,(
  ( ! [X26,X27] : (sP3('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1028,f977])).
fof(f19621,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) | $spl408),
  inference(cnf_transformation,[],[f19621_D])).
fof(f19621_D,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f19608,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1027,f977])).
fof(f1027,plain,(
  ( ! [X78,X77] : (sP3(X77,X78,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X77))) )),
  inference(resolution,[],[f652,f709])).
fof(f19592,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1027,f977])).
fof(f19576,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) | $spl404),
  inference(cnf_transformation,[],[f19576_D])).
fof(f19576_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f19559,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1021,f977])).
fof(f1021,plain,(
  ( ! [X66,X65] : (~sP3(X65,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X66) | p(s(bool,X65))) )),
  inference(resolution,[],[f652,f698])).
fof(f19543,plain,(
  ( ! [X26,X27] : (~sP3(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1021,f977])).
fof(f19487,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5140,f977])).
fof(f5140,plain,(
  ( ! [X163] : (~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X163))) ) | $spl17),
  inference(subsumption_resolution,[],[f5042,f900])).
fof(f5042,plain,(
  ( ! [X163] : (p(s(bool,'F_0')) | p(s(bool,X163)) | ~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl17),
  inference(resolution,[],[f1174,f798])).
fof(f19430,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5139,f977])).
fof(f19413,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) | $spl396),
  inference(cnf_transformation,[],[f19413_D])).
fof(f19413_D,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f19396,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1020,f977])).
fof(f1020,plain,(
  ( ! [X64,X63] : (~sP3(X63,X64,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | p(s(bool,X63))) )),
  inference(resolution,[],[f652,f697])).
fof(f19380,plain,(
  ( ! [X26,X27] : (~sP3(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1020,f977])).
fof(f18557,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5036,f977])).
fof(f5036,plain,(
  ( ! [X157] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X157) | ~p(s(bool,X157))) ) | $spl17),
  inference(resolution,[],[f1174,f787])).
fof(f18500,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5035,f977])).
fof(f5035,plain,(
  ( ! [X156] : (sP10(X156,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X156))) ) | $spl17),
  inference(resolution,[],[f1174,f784])).
fof(f18481,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) | $spl388),
  inference(cnf_transformation,[],[f18481_D])).
fof(f18481_D,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f18465,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1016,f977])).
fof(f1016,plain,(
  ( ! [X56,X55] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X55,X56) | ~p(s(bool,X55))) )),
  inference(resolution,[],[f652,f693])).
fof(f18449,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1016,f977])).
fof(f18398,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5034,f977])).
fof(f5034,plain,(
  ( ! [X155] : (~sP10(X155,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X155))) ) | $spl17),
  inference(resolution,[],[f1174,f782])).
fof(f18345,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5033,f977])).
fof(f5033,plain,(
  ( ! [X154] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X154) | p(s(bool,X154))) ) | $spl17),
  inference(resolution,[],[f1174,f782])).
fof(f18328,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl380),
  inference(cnf_transformation,[],[f18328_D])).
fof(f18328_D,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f18310,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4947,f977])).
fof(f4947,plain,(
  ( ! [X163] : (~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X163))) ) | $spl1),
  inference(subsumption_resolution,[],[f4854,f900])).
fof(f4854,plain,(
  ( ! [X163] : (p(s(bool,'F_0')) | p(s(bool,X163)) | ~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f798])).
fof(f17131,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X19))) ) | $spl350),
  inference(cnf_transformation,[],[f17131_D])).
fof(f17131_D,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X19))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f18294,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4947,f977])).
fof(f18277,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) | $spl378),
  inference(cnf_transformation,[],[f18277_D])).
fof(f18277_D,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f18265,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1015,f977])).
fof(f1015,plain,(
  ( ! [X54,X53] : (sP1(X53,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X54) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f652,f693])).
fof(f18249,plain,(
  ( ! [X26,X27] : (sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | ~p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1015,f977])).
fof(f18234,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl374),
  inference(cnf_transformation,[],[f18234_D])).
fof(f18234_D,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f18217,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4946,f977])).
fof(f17079,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | p(s(bool,X19))) ) | $spl346),
  inference(cnf_transformation,[],[f17079_D])).
fof(f17079_D,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | p(s(bool,X19))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f18201,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4946,f977])).
fof(f18165,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl373),
  inference(cnf_transformation,[],[f18165_D])).
fof(f18165_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f18125,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) | $spl1),
  inference(superposition,[],[f4876,f977])).
fof(f4876,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | $spl1),
  inference(resolution,[],[f1161,f927])).
fof(f927,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f635])).
fof(f18101,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) | $spl370),
  inference(cnf_transformation,[],[f18101_D])).
fof(f18101_D,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f18084,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1014,f977])).
fof(f1014,plain,(
  ( ! [X52,X51] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X51,X52) | p(s(bool,X52))) )),
  inference(resolution,[],[f652,f692])).
fof(f18068,plain,(
  ( ! [X26,X27] : (sP1('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | p(s(bool,X27)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1014,f977])).
fof(f17799,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl366),
  inference(cnf_transformation,[],[f17799_D])).
fof(f17799_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f17404,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) | $spl1),
  inference(superposition,[],[f4865,f977])).
fof(f4865,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | $spl1),
  inference(resolution,[],[f1161,f878])).
fof(f17379,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) | $spl364),
  inference(cnf_transformation,[],[f17379_D])).
fof(f17379_D,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f17364,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1005,f977])).
fof(f1005,plain,(
  ( ! [X33,X34] : (~sP1(X33,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X34) | p(s(bool,X33))) )),
  inference(resolution,[],[f652,f681])).
fof(f17348,plain,(
  ( ! [X26,X27] : (~sP1(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | p(s(bool,X26)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1005,f977])).
fof(f17333,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) | $spl360),
  inference(cnf_transformation,[],[f17333_D])).
fof(f17333_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f17318,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4848,f977])).
fof(f4848,plain,(
  ( ! [X157] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X157) | ~p(s(bool,X157))) ) | $spl1),
  inference(resolution,[],[f1161,f787])).
fof(f16273,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | ~p(s(bool,X19))) ) | $spl342),
  inference(cnf_transformation,[],[f16273_D])).
fof(f16273_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | ~p(s(bool,X19))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f17302,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4848,f977])).
fof(f17284,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) | $spl358),
  inference(cnf_transformation,[],[f17284_D])).
fof(f17284_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f17269,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4847,f977])).
fof(f4847,plain,(
  ( ! [X156] : (sP10(X156,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X156))) ) | $spl1),
  inference(resolution,[],[f1161,f784])).
fof(f16226,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X19))) ) | $spl338),
  inference(cnf_transformation,[],[f16226_D])).
fof(f16226_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X19))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f17253,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4847,f977])).
fof(f17235,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl356),
  inference(cnf_transformation,[],[f17235_D])).
fof(f17235_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f17220,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4846,f977])).
fof(f4846,plain,(
  ( ! [X155] : (~sP10(X155,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | p(s(bool,X155))) ) | $spl1),
  inference(resolution,[],[f1161,f782])).
fof(f16179,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X19))) ) | $spl334),
  inference(cnf_transformation,[],[f16179_D])).
fof(f16179_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X19))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f17204,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4846,f977])).
fof(f17188,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl354),
  inference(cnf_transformation,[],[f17188_D])).
fof(f17188_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f17173,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f4845,f977])).
fof(f4845,plain,(
  ( ! [X154] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X154) | p(s(bool,X154))) ) | $spl1),
  inference(resolution,[],[f1161,f782])).
fof(f16132,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | p(s(bool,X19))) ) | $spl330),
  inference(cnf_transformation,[],[f16132_D])).
fof(f16132_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | p(s(bool,X19))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f17157,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f4845,f977])).
fof(f17132,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl352),
  inference(cnf_transformation,[],[f17132_D])).
fof(f17132_D,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f17113,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2589,f977])).
fof(f2589,plain,(
  ( ! [X163] : (~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | p(s(bool,X163))) )),
  inference(subsumption_resolution,[],[f2493,f900])).
fof(f2493,plain,(
  ( ! [X163] : (p(s(bool,'F_0')) | p(s(bool,X163)) | ~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f798])).
fof(f17097,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2589,f977])).
fof(f17080,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl348),
  inference(cnf_transformation,[],[f17080_D])).
fof(f17080_D,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f17062,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2588,f977])).
fof(f17046,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2588,f977])).
fof(f14805,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl305),
  inference(cnf_transformation,[],[f14805_D])).
fof(f14805_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl305),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl305])])).
fof(f16972,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f2515,f977])).
fof(f2515,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))),
  inference(resolution,[],[f655,f927])).
fof(f14382,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl302),
  inference(cnf_transformation,[],[f14382_D])).
fof(f14382_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f16299,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f2504,f977])).
fof(f2504,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))),
  inference(resolution,[],[f655,f878])).
fof(f16274,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) | $spl344),
  inference(cnf_transformation,[],[f16274_D])).
fof(f16274_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f16260,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2487,f977])).
fof(f2487,plain,(
  ( ! [X157] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X157) | ~p(s(bool,X157))) )),
  inference(resolution,[],[f655,f787])).
fof(f16244,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2487,f977])).
fof(f16227,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) | $spl340),
  inference(cnf_transformation,[],[f16227_D])).
fof(f16227_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f16213,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2486,f977])).
fof(f2486,plain,(
  ( ! [X156] : (sP10(X156,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X156))) )),
  inference(resolution,[],[f655,f784])).
fof(f16197,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2486,f977])).
fof(f16180,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl336),
  inference(cnf_transformation,[],[f16180_D])).
fof(f16180_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f16164,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2485,f977])).
fof(f2485,plain,(
  ( ! [X155] : (~sP10(X155,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | p(s(bool,X155))) )),
  inference(resolution,[],[f655,f782])).
fof(f16148,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2485,f977])).
fof(f16133,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl332),
  inference(cnf_transformation,[],[f16133_D])).
fof(f16133_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f16117,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2484,f977])).
fof(f2484,plain,(
  ( ! [X154] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X154) | p(s(bool,X154))) )),
  inference(resolution,[],[f655,f782])).
fof(f16101,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2484,f977])).
fof(f16086,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl328),
  inference(cnf_transformation,[],[f16086_D])).
fof(f16086_D,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f16067,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1983,f977])).
fof(f1983,plain,(
  ( ! [X163] : (~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | p(s(bool,X163))) )),
  inference(subsumption_resolution,[],[f1885,f900])).
fof(f1885,plain,(
  ( ! [X163] : (p(s(bool,'F_0')) | p(s(bool,X163)) | ~sP11(X163,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f798])).
fof(f16051,plain,(
  ( ! [X19] : (~sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1983,f977])).
fof(f16034,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl324),
  inference(cnf_transformation,[],[f16034_D])).
fof(f16034_D,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f16016,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1982,f977])).
fof(f16000,plain,(
  ( ! [X19] : (~sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1982,f977])).
fof(f13166,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl291),
  inference(cnf_transformation,[],[f13166_D])).
fof(f13166_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl291),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl291])])).
fof(f15926,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f1907,f977])).
fof(f1907,plain,(
  s(bool,'T_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))),
  inference(resolution,[],[f654,f927])).
fof(f12491,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl288),
  inference(cnf_transformation,[],[f12491_D])).
fof(f12491_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f15255,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f1896,f977])).
fof(f1896,plain,(
  s(bool,'F_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))),
  inference(resolution,[],[f654,f878])).
fof(f15230,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) | $spl320),
  inference(cnf_transformation,[],[f15230_D])).
fof(f15230_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f15218,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1879,f977])).
fof(f1879,plain,(
  ( ! [X157] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X157) | ~p(s(bool,X157))) )),
  inference(resolution,[],[f654,f787])).
fof(f15202,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1879,f977])).
fof(f15186,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) | $spl316),
  inference(cnf_transformation,[],[f15186_D])).
fof(f15186_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f15174,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1878,f977])).
fof(f1878,plain,(
  ( ! [X156] : (sP10(X156,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X156))) )),
  inference(resolution,[],[f654,f784])).
fof(f15158,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1878,f977])).
fof(f15142,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl312),
  inference(cnf_transformation,[],[f15142_D])).
fof(f15142_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f15126,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1877,f977])).
fof(f1877,plain,(
  ( ! [X155] : (~sP10(X155,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | p(s(bool,X155))) )),
  inference(resolution,[],[f654,f782])).
fof(f15110,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1877,f977])).
fof(f15096,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl308),
  inference(cnf_transformation,[],[f15096_D])).
fof(f15096_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f15080,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1876,f977])).
fof(f1876,plain,(
  ( ! [X154] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X154) | p(s(bool,X154))) )),
  inference(resolution,[],[f654,f782])).
fof(f15064,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1876,f977])).
fof(f13412,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) | $spl293),
  inference(cnf_transformation,[],[f13412_D])).
fof(f13412_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) <=> ~$spl293),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl293])])).
fof(f15031,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12)))) | $spl27),
  inference(superposition,[],[f1361,f977])).
fof(f14806,plain,(
  s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl304),
  inference(cnf_transformation,[],[f14806_D])).
fof(f14806_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f14413,plain,(
  s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f1288,f977])).
fof(f14383,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl303),
  inference(cnf_transformation,[],[f14383_D])).
fof(f14383_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f14342,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f1286,f977])).
fof(f1286,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))),
  inference(resolution,[],[f653,f923])).
fof(f14286,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9676,f977])).
fof(f9676,plain,(
  ( ! [X142,X143] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X142,X143)) ) | $spl31),
  inference(resolution,[],[f1365,f745])).
fof(f14242,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9675,f977])).
fof(f9675,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X141)) ) | $spl31),
  inference(resolution,[],[f1365,f744])).
fof(f14198,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9673,f977])).
fof(f14155,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9672,f977])).
fof(f9672,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X135)) ) | $spl31),
  inference(resolution,[],[f1365,f740])).
fof(f14113,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9660,f977])).
fof(f9660,plain,(
  ( ! [X111,X110] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X110,X111)) ) | $spl31),
  inference(resolution,[],[f1365,f724])).
fof(f14070,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9659,f977])).
fof(f9659,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X109)) ) | $spl31),
  inference(resolution,[],[f1365,f723])).
fof(f14026,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9646,f977])).
fof(f9646,plain,(
  ( ! [X83,X82] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X82,X83)) ) | $spl31),
  inference(resolution,[],[f1365,f707])).
fof(f14013,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl300),
  inference(cnf_transformation,[],[f14013_D])).
fof(f14013_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f14000,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1263,f977])).
fof(f1263,plain,(
  ( ! [X154] : (sP10(X154,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X154))) )),
  inference(resolution,[],[f653,f787])).
fof(f11811,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X19))) ) | $spl262),
  inference(cnf_transformation,[],[f11811_D])).
fof(f11811_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X19))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f13984,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1263,f977])).
fof(f13939,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9630,f977])).
fof(f9630,plain,(
  ( ! [X50,X51] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X50,X51)) ) | $spl31),
  inference(resolution,[],[f1365,f690])).
fof(f13896,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) ) | $spl31),
  inference(superposition,[],[f9629,f977])).
fof(f9629,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f688])).
fof(f13883,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl298),
  inference(cnf_transformation,[],[f13883_D])).
fof(f13883_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f13870,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1262,f977])).
fof(f1262,plain,(
  ( ! [X153] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X153) | p(s(bool,X153))) )),
  inference(resolution,[],[f653,f784])).
fof(f11572,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | p(s(bool,X19))) ) | $spl248),
  inference(cnf_transformation,[],[f11572_D])).
fof(f11572_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | p(s(bool,X19))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f13854,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1262,f977])).
fof(f13839,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) | $spl296),
  inference(cnf_transformation,[],[f13839_D])).
fof(f13839_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f13828,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1261,f977])).
fof(f1261,plain,(
  ( ! [X152] : (~sP10(X152,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X152))) )),
  inference(resolution,[],[f653,f783])).
fof(f783,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f551])).
fof(f11322,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X19))) ) | $spl234),
  inference(cnf_transformation,[],[f11322_D])).
fof(f11322_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X19))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f13812,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1261,f977])).
fof(f10567,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) | $spl176),
  inference(cnf_transformation,[],[f10567_D])).
fof(f10567_D,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f13772,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7839,f977])).
fof(f7839,plain,(
  ( ! [X142,X143] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X142,X143)) ) | $spl23),
  inference(resolution,[],[f1357,f745])).
fof(f10526,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27)) ) | $spl172),
  inference(cnf_transformation,[],[f10526_D])).
fof(f10526_D,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f13733,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7838,f977])).
fof(f7838,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X141)) ) | $spl23),
  inference(resolution,[],[f1357,f744])).
fof(f10485,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) | $spl168),
  inference(cnf_transformation,[],[f10485_D])).
fof(f10485_D,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f13694,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7836,f977])).
fof(f7836,plain,(
  ( ! [X136,X137] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X136,X137)) ) | $spl23),
  inference(resolution,[],[f1357,f741])).
fof(f10444,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27)) ) | $spl164),
  inference(cnf_transformation,[],[f10444_D])).
fof(f10444_D,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f13656,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7835,f977])).
fof(f7835,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X135)) ) | $spl23),
  inference(resolution,[],[f1357,f740])).
fof(f13644,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) | $spl294),
  inference(cnf_transformation,[],[f13644_D])).
fof(f13644_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f13633,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1260,f977])).
fof(f1260,plain,(
  ( ! [X151] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X151) | ~p(s(bool,X151))) )),
  inference(resolution,[],[f653,f783])).
fof(f11060,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | ~p(s(bool,X19))) ) | $spl220),
  inference(cnf_transformation,[],[f11060_D])).
fof(f11060_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | ~p(s(bool,X19))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f13617,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1260,f977])).
fof(f10404,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) | $spl160),
  inference(cnf_transformation,[],[f10404_D])).
fof(f10404_D,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f13579,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7823,f977])).
fof(f7823,plain,(
  ( ! [X111,X110] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X110,X111)) ) | $spl23),
  inference(resolution,[],[f1357,f724])).
fof(f10364,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27)) ) | $spl156),
  inference(cnf_transformation,[],[f10364_D])).
fof(f10364_D,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f13541,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7822,f977])).
fof(f7822,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X109)) ) | $spl23),
  inference(resolution,[],[f1357,f723])).
fof(f10324,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) | $spl152),
  inference(cnf_transformation,[],[f10324_D])).
fof(f10324_D,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f13502,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7809,f977])).
fof(f7809,plain,(
  ( ! [X83,X82] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X82,X83)) ) | $spl23),
  inference(resolution,[],[f1357,f707])).
fof(f10283,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) | $spl148),
  inference(cnf_transformation,[],[f10283_D])).
fof(f10283_D,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f13464,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7793,f977])).
fof(f7793,plain,(
  ( ! [X50,X51] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)),X50,X51)) ) | $spl23),
  inference(resolution,[],[f1357,f690])).
fof(f10243,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)))) ) | $spl144),
  inference(cnf_transformation,[],[f10243_D])).
fof(f10243_D,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f13426,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl23),
  inference(superposition,[],[f7792,f977])).
fof(f7792,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f688])).
fof(f13392,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) | $spl5),
  inference(superposition,[],[f1165,f977])).
fof(f13167,plain,(
  s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl290),
  inference(cnf_transformation,[],[f13167_D])).
fof(f13167_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f12774,plain,(
  s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f1092,f977])).
fof(f1092,plain,(
  s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))),
  inference(resolution,[],[f652,f928])).
fof(f12716,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5240,f977])).
fof(f12672,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5239,f977])).
fof(f12628,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5237,f977])).
fof(f12585,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5236,f977])).
fof(f12543,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5224,f977])).
fof(f12492,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | $spl289),
  inference(cnf_transformation,[],[f12492_D])).
fof(f12492_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) <=> ~$spl289),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl289])])).
fof(f12449,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f1090,f977])).
fof(f12490,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl287),
  inference(cnf_transformation,[],[f12490_D])).
fof(f12490_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f12488,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(forward_demodulation,[],[f12433,f964])).
fof(f964,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f400])).
fof(f400,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f399])).
fof(f399,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7] : s(bool,'T_0') = s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X6),X7)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.list.isPREFIX|split|0')).
fof(f12433,plain,(
  s(bool,'F_0') != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f1090,f977])).
fof(f12395,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5223,f977])).
fof(f12351,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5210,f977])).
fof(f12308,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5194,f977])).
fof(f12265,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl19),
  inference(superposition,[],[f5193,f977])).
fof(f12219,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5027,f977])).
fof(f5027,plain,(
  ( ! [X142,X143] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X142,X143)) ) | $spl17),
  inference(resolution,[],[f1174,f745])).
fof(f12174,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5026,f977])).
fof(f5026,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X141)) ) | $spl17),
  inference(resolution,[],[f1174,f744])).
fof(f12129,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5024,f977])).
fof(f5024,plain,(
  ( ! [X136,X137] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X136,X137)) ) | $spl17),
  inference(resolution,[],[f1174,f741])).
fof(f12085,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5023,f977])).
fof(f5023,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X135)) ) | $spl17),
  inference(resolution,[],[f1174,f740])).
fof(f12042,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5011,f977])).
fof(f5011,plain,(
  ( ! [X111,X110] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X110,X111)) ) | $spl17),
  inference(resolution,[],[f1174,f724])).
fof(f11998,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f5010,f977])).
fof(f5010,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X109)) ) | $spl17),
  inference(resolution,[],[f1174,f723])).
fof(f11953,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4997,f977])).
fof(f4997,plain,(
  ( ! [X83,X82] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X82,X83)) ) | $spl17),
  inference(resolution,[],[f1174,f707])).
fof(f11909,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4981,f977])).
fof(f4981,plain,(
  ( ! [X50,X51] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X50,X51)) ) | $spl17),
  inference(resolution,[],[f1174,f690])).
fof(f11865,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl17),
  inference(superposition,[],[f4980,f977])).
fof(f4980,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl17),
  inference(resolution,[],[f1174,f688])).
fof(f11852,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl266),
  inference(cnf_transformation,[],[f11852_D])).
fof(f11852_D,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f11840,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3190,f977])).
fof(f3190,plain,(
  ( ! [X142,X143] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X142,X143)) ) | $spl1),
  inference(resolution,[],[f1161,f745])).
fof(f10932,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) | $spl212),
  inference(cnf_transformation,[],[f10932_D])).
fof(f10932_D,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f11824,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3190,f977])).
fof(f11812,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) | $spl264),
  inference(cnf_transformation,[],[f11812_D])).
fof(f11812_D,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f11800,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1067,f977])).
fof(f1067,plain,(
  ( ! [X154] : (sP10(X154,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | p(s(bool,X154))) )),
  inference(resolution,[],[f652,f787])).
fof(f11784,plain,(
  ( ! [X19] : (sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1067,f977])).
fof(f11770,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl260),
  inference(cnf_transformation,[],[f11770_D])).
fof(f11770_D,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f11758,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3189,f977])).
fof(f3189,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X141)) ) | $spl1),
  inference(resolution,[],[f1161,f744])).
fof(f10891,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27)) ) | $spl208),
  inference(cnf_transformation,[],[f10891_D])).
fof(f10891_D,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f11742,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3189,f977])).
fof(f11730,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl258),
  inference(cnf_transformation,[],[f11730_D])).
fof(f11730_D,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f11718,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3187,f977])).
fof(f10850,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) | $spl204),
  inference(cnf_transformation,[],[f10850_D])).
fof(f10850_D,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f11702,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3187,f977])).
fof(f11690,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl256),
  inference(cnf_transformation,[],[f11690_D])).
fof(f11690_D,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f11679,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3186,f977])).
fof(f3186,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X135)) ) | $spl1),
  inference(resolution,[],[f1161,f740])).
fof(f10809,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27)) ) | $spl200),
  inference(cnf_transformation,[],[f10809_D])).
fof(f10809_D,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f11663,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3186,f977])).
fof(f11651,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl254),
  inference(cnf_transformation,[],[f11651_D])).
fof(f11651_D,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f11641,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3174,f977])).
fof(f3174,plain,(
  ( ! [X111,X110] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X110,X111)) ) | $spl1),
  inference(resolution,[],[f1161,f724])).
fof(f10769,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) | $spl196),
  inference(cnf_transformation,[],[f10769_D])).
fof(f10769_D,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f11625,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3174,f977])).
fof(f11612,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl252),
  inference(cnf_transformation,[],[f11612_D])).
fof(f11612_D,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f11602,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3173,f977])).
fof(f3173,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X109)) ) | $spl1),
  inference(resolution,[],[f1161,f723])).
fof(f10729,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27)) ) | $spl192),
  inference(cnf_transformation,[],[f10729_D])).
fof(f10729_D,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f11586,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3173,f977])).
fof(f11573,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) | $spl250),
  inference(cnf_transformation,[],[f11573_D])).
fof(f11573_D,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f11561,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1066,f977])).
fof(f1066,plain,(
  ( ! [X153] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X153) | p(s(bool,X153))) )),
  inference(resolution,[],[f652,f784])).
fof(f11545,plain,(
  ( ! [X19] : (sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1066,f977])).
fof(f11531,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl246),
  inference(cnf_transformation,[],[f11531_D])).
fof(f11531_D,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f11520,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3160,f977])).
fof(f3160,plain,(
  ( ! [X83,X82] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X82,X83)) ) | $spl1),
  inference(resolution,[],[f1161,f707])).
fof(f10689,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) | $spl188),
  inference(cnf_transformation,[],[f10689_D])).
fof(f10689_D,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f11504,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3160,f977])).
fof(f11491,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl244),
  inference(cnf_transformation,[],[f11491_D])).
fof(f11491_D,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f11481,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3144,f977])).
fof(f3144,plain,(
  ( ! [X50,X51] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)),X50,X51)) ) | $spl1),
  inference(resolution,[],[f1161,f690])).
fof(f10648,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) | $spl184),
  inference(cnf_transformation,[],[f10648_D])).
fof(f10648_D,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f11465,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3144,f977])).
fof(f11452,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl242),
  inference(cnf_transformation,[],[f11452_D])).
fof(f11452_D,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f11442,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) ) | $spl1),
  inference(superposition,[],[f3143,f977])).
fof(f3143,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f688])).
fof(f10608,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)))) ) | $spl180),
  inference(cnf_transformation,[],[f10608_D])).
fof(f10608_D,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f11426,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl1),
  inference(superposition,[],[f3143,f977])).
fof(f11413,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl240),
  inference(cnf_transformation,[],[f11413_D])).
fof(f11413_D,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f11398,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl100),
  inference(superposition,[],[f3117,f977])).
fof(f11368,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl238),
  inference(cnf_transformation,[],[f11368_D])).
fof(f11368_D,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f11353,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl96),
  inference(superposition,[],[f3076,f977])).
fof(f11323,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) | $spl236),
  inference(cnf_transformation,[],[f11323_D])).
fof(f11323_D,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f11311,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1065,f977])).
fof(f1065,plain,(
  ( ! [X152] : (~sP10(X152,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X152))) )),
  inference(resolution,[],[f652,f783])).
fof(f11295,plain,(
  ( ! [X19] : (~sP10(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1065,f977])).
fof(f11283,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl232),
  inference(cnf_transformation,[],[f11283_D])).
fof(f11283_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f11268,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl92),
  inference(superposition,[],[f3035,f977])).
fof(f11238,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl230),
  inference(cnf_transformation,[],[f11238_D])).
fof(f11238_D,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f11224,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl88),
  inference(superposition,[],[f2994,f977])).
fof(f11194,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl228),
  inference(cnf_transformation,[],[f11194_D])).
fof(f11194_D,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f11180,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl84),
  inference(superposition,[],[f2954,f977])).
fof(f11150,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl226),
  inference(cnf_transformation,[],[f11150_D])).
fof(f11150_D,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f11136,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl80),
  inference(superposition,[],[f2914,f977])).
fof(f11105,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl224),
  inference(cnf_transformation,[],[f11105_D])).
fof(f11105_D,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f11092,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl76),
  inference(superposition,[],[f2873,f977])).
fof(f11061,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) | $spl222),
  inference(cnf_transformation,[],[f11061_D])).
fof(f11061_D,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f11049,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1064,f977])).
fof(f1064,plain,(
  ( ! [X151] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X151) | ~p(s(bool,X151))) )),
  inference(resolution,[],[f652,f783])).
fof(f11033,plain,(
  ( ! [X19] : (~sP10('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | ~p(s(bool,X19)) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1064,f977])).
fof(f11021,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl218),
  inference(cnf_transformation,[],[f11021_D])).
fof(f11021_D,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f11008,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl72),
  inference(superposition,[],[f2833,f977])).
fof(f10977,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl216),
  inference(cnf_transformation,[],[f10977_D])).
fof(f10977_D,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f10964,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl66),
  inference(superposition,[],[f2751,f977])).
fof(f10933,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl214),
  inference(cnf_transformation,[],[f10933_D])).
fof(f10933_D,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f10920,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2478,f977])).
fof(f2478,plain,(
  ( ! [X142,X143] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X142,X143)) )),
  inference(resolution,[],[f655,f745])).
fof(f10904,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2478,f977])).
fof(f10892,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl210),
  inference(cnf_transformation,[],[f10892_D])).
fof(f10892_D,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f10879,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2477,f977])).
fof(f2477,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X141)) )),
  inference(resolution,[],[f655,f744])).
fof(f10863,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2477,f977])).
fof(f10851,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl206),
  inference(cnf_transformation,[],[f10851_D])).
fof(f10851_D,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f10838,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2475,f977])).
fof(f10822,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2475,f977])).
fof(f10810,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl202),
  inference(cnf_transformation,[],[f10810_D])).
fof(f10810_D,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f10798,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2474,f977])).
fof(f2474,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X135)) )),
  inference(resolution,[],[f655,f740])).
fof(f10782,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2474,f977])).
fof(f10770,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl198),
  inference(cnf_transformation,[],[f10770_D])).
fof(f10770_D,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f10759,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2462,f977])).
fof(f2462,plain,(
  ( ! [X111,X110] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X110,X111)) )),
  inference(resolution,[],[f655,f724])).
fof(f10743,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2462,f977])).
fof(f10730,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl194),
  inference(cnf_transformation,[],[f10730_D])).
fof(f10730_D,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f10719,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2461,f977])).
fof(f2461,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X109)) )),
  inference(resolution,[],[f655,f723])).
fof(f10703,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2461,f977])).
fof(f10690,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl190),
  inference(cnf_transformation,[],[f10690_D])).
fof(f10690_D,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f10678,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2448,f977])).
fof(f2448,plain,(
  ( ! [X83,X82] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X82,X83)) )),
  inference(resolution,[],[f655,f707])).
fof(f10662,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2448,f977])).
fof(f10649,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl186),
  inference(cnf_transformation,[],[f10649_D])).
fof(f10649_D,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f10638,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2432,f977])).
fof(f2432,plain,(
  ( ! [X50,X51] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)),X50,X51)) )),
  inference(resolution,[],[f655,f690])).
fof(f10622,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2432,f977])).
fof(f10609,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl182),
  inference(cnf_transformation,[],[f10609_D])).
fof(f10609_D,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f10598,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f2431,f977])).
fof(f2431,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f688])).
fof(f10582,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f2431,f977])).
fof(f10569,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f1978,f655])).
fof(f10568,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl178),
  inference(cnf_transformation,[],[f10568_D])).
fof(f10568_D,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f10555,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1870,f977])).
fof(f1870,plain,(
  ( ! [X142,X143] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X142,X143)) )),
  inference(resolution,[],[f654,f745])).
fof(f10539,plain,(
  ( ! [X26,X27] : (sP7('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1870,f977])).
fof(f10527,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl174),
  inference(cnf_transformation,[],[f10527_D])).
fof(f10527_D,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f10514,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1869,f977])).
fof(f1869,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X141)) )),
  inference(resolution,[],[f654,f744])).
fof(f10498,plain,(
  ( ! [X26,X27] : (sP7(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1869,f977])).
fof(f10486,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl170),
  inference(cnf_transformation,[],[f10486_D])).
fof(f10486_D,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f10473,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1867,f977])).
fof(f10457,plain,(
  ( ! [X26,X27] : (sP8('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1867,f977])).
fof(f10445,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl166),
  inference(cnf_transformation,[],[f10445_D])).
fof(f10445_D,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f10433,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1866,f977])).
fof(f1866,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X135)) )),
  inference(resolution,[],[f654,f740])).
fof(f10417,plain,(
  ( ! [X26,X27] : (sP8(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1866,f977])).
fof(f10405,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl162),
  inference(cnf_transformation,[],[f10405_D])).
fof(f10405_D,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f10394,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1854,f977])).
fof(f1854,plain,(
  ( ! [X111,X110] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X110,X111)) )),
  inference(resolution,[],[f654,f724])).
fof(f10378,plain,(
  ( ! [X26,X27] : (sP4('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1854,f977])).
fof(f10365,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) | $spl158),
  inference(cnf_transformation,[],[f10365_D])).
fof(f10365_D,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f10354,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1853,f977])).
fof(f1853,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X109)) )),
  inference(resolution,[],[f654,f723])).
fof(f10338,plain,(
  ( ! [X26,X27] : (sP4(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1853,f977])).
fof(f10325,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl154),
  inference(cnf_transformation,[],[f10325_D])).
fof(f10325_D,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f10313,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1840,f977])).
fof(f1840,plain,(
  ( ! [X83,X82] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X82,X83)) )),
  inference(resolution,[],[f654,f707])).
fof(f10297,plain,(
  ( ! [X26,X27] : (sP2('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1840,f977])).
fof(f10284,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) | $spl150),
  inference(cnf_transformation,[],[f10284_D])).
fof(f10284_D,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f10273,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1824,f977])).
fof(f1824,plain,(
  ( ! [X50,X51] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)),X50,X51)) )),
  inference(resolution,[],[f654,f690])).
fof(f10257,plain,(
  ( ! [X26,X27] : (sP0('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1824,f977])).
fof(f10244,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl146),
  inference(cnf_transformation,[],[f10244_D])).
fof(f10244_D,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f10233,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1823,f977])).
fof(f1823,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f688])).
fof(f10217,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1823,f977])).
fof(f9786,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl136),
  inference(cnf_transformation,[],[f9786_D])).
fof(f9786_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f10124,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f978])).
fof(f10029,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X33) != s(bool,X35)) ) | $spl31),
  inference(resolution,[],[f1365,f679])).
fof(f10028,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f678])).
fof(f10026,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f677])).
fof(f10024,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl31),
  inference(resolution,[],[f1365,f675])).
fof(f10022,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f673])).
fof(f9983,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl142),
  inference(cnf_transformation,[],[f9983_D])).
fof(f9983_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f9909,plain,(
  ( ! [X169,X170] : (p(s(bool,X169)) | p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl30),
  inference(resolution,[],[f1366,f978])).
fof(f9818,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X31) != s(bool,X32)) ) | $spl30),
  inference(resolution,[],[f1366,f679])).
fof(f9815,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl30),
  inference(resolution,[],[f1366,f678])).
fof(f9814,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl30),
  inference(resolution,[],[f1366,f677])).
fof(f9813,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X16) != s(bool,X17)) ) | $spl30),
  inference(resolution,[],[f1366,f675])).
fof(f9811,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl30),
  inference(resolution,[],[f1366,f673])).
fof(f9980,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) | $spl139),
  inference(cnf_transformation,[],[f9980_D])).
fof(f9980_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f9982,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK15) | $spl140),
  inference(cnf_transformation,[],[f9982_D])).
fof(f9982_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK15) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f9808,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK15) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) | $spl30),
  inference(resolution,[],[f1366,f955])).
fof(f955,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f642])).
fof(f642,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))))),
  inference(flattening,[],[f641])).
fof(f641,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))) <=> s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(flattening,[],[f387])).
fof(f387,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))) <=> s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6] : ! [X11,X18] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X11),s('type.list.list'(X6),X18)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X11))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X18)))) <=> s('type.list.list'(X6),X11) = s('type.list.list'(X6),X18))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_LENGTH_ANTI')).
fof(f9717,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f978])).
fof(f9622,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X33) != s(bool,X35)) ) | $spl31),
  inference(resolution,[],[f1365,f679])).
fof(f9621,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f678])).
fof(f9619,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f677])).
fof(f9617,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl31),
  inference(resolution,[],[f1365,f675])).
fof(f9615,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl31),
  inference(resolution,[],[f1365,f673])).
fof(f7945,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl134),
  inference(cnf_transformation,[],[f7945_D])).
fof(f7945_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f9530,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f978])).
fof(f9435,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl23),
  inference(resolution,[],[f1357,f679])).
fof(f9434,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f678])).
fof(f9432,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f677])).
fof(f9430,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | s(bool,X18) != s(bool,X20)) ) | $spl23),
  inference(resolution,[],[f1357,f675])).
fof(f9428,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f673])).
fof(f9406,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1257,f977])).
fof(f1257,plain,(
  ( ! [X146,X145] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X145,X146)) )),
  inference(resolution,[],[f653,f749])).
fof(f3117,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X6,X7)) ) | $spl100),
  inference(cnf_transformation,[],[f3117_D])).
fof(f3117_D,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X6,X7)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f9390,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1257,f977])).
fof(f1365,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl31),
  inference(cnf_transformation,[],[f1365_D])).
fof(f1365_D,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f9385,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X6,X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1257,f863])).
fof(f9367,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1256,f977])).
fof(f1256,plain,(
  ( ! [X144,X143] : (sP6(X143,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X144)) )),
  inference(resolution,[],[f653,f748])).
fof(f3076,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7)) ) | $spl96),
  inference(cnf_transformation,[],[f3076_D])).
fof(f3076_D,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f9351,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1256,f977])).
fof(f9346,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1256,f863])).
fof(f9328,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1255,f977])).
fof(f3035,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) | $spl92),
  inference(cnf_transformation,[],[f3035_D])).
fof(f3035_D,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f9312,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1255,f977])).
fof(f9307,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1255,f863])).
fof(f9290,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1254,f977])).
fof(f1254,plain,(
  ( ! [X140,X139] : (sP7(X139,X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f743])).
fof(f2994,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) | $spl88),
  inference(cnf_transformation,[],[f2994_D])).
fof(f2994_D,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f9274,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1254,f977])).
fof(f9269,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1254,f863])).
fof(f9252,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1251,f977])).
fof(f1251,plain,(
  ( ! [X134,X133] : (sP8(X133,X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f739])).
fof(f2954,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) | $spl84),
  inference(cnf_transformation,[],[f2954_D])).
fof(f2954_D,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9236,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1251,f977])).
fof(f9231,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1251,f863])).
fof(f9214,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1238,f977])).
fof(f1238,plain,(
  ( ! [X107,X108] : (sP4(X107,X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f722])).
fof(f2914,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) | $spl80),
  inference(cnf_transformation,[],[f2914_D])).
fof(f2914_D,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f9198,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1238,f977])).
fof(f9193,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1238,f863])).
fof(f9176,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1222,f977])).
fof(f1222,plain,(
  ( ! [X76,X75] : (sP2(X75,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X76)) )),
  inference(resolution,[],[f653,f706])).
fof(f2873,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7)) ) | $spl76),
  inference(cnf_transformation,[],[f2873_D])).
fof(f2873_D,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f9160,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1222,f977])).
fof(f9155,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1222,f863])).
fof(f9138,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1221,f977])).
fof(f1221,plain,(
  ( ! [X74,X73] : (sP2(X73,X74,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f705])).
fof(f2833,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) | $spl72),
  inference(cnf_transformation,[],[f2833_D])).
fof(f2833_D,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f9122,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1221,f977])).
fof(f9117,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1221,f863])).
fof(f9100,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1208,f977])).
fof(f1208,plain,(
  ( ! [X47,X48] : (sP0(X47,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X48)) )),
  inference(resolution,[],[f653,f689])).
fof(f2751,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7)) ) | $spl66),
  inference(cnf_transformation,[],[f2751_D])).
fof(f2751_D,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f9084,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1208,f977])).
fof(f9079,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1208,f863])).
fof(f9020,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f655,f977])).
fof(f2586,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) ) | $spl58),
  inference(cnf_transformation,[],[f2586_D])).
fof(f2586_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f8995,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f978])).
fof(f8900,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f655,f679])).
fof(f8899,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f678])).
fof(f8897,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f677])).
fof(f8895,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f655,f675])).
fof(f8893,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f673])).
fof(f1980,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) ) | $spl50),
  inference(cnf_transformation,[],[f1980_D])).
fof(f1980_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f8809,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f978])).
fof(f8714,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f654,f679])).
fof(f8713,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f678])).
fof(f8711,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f677])).
fof(f8709,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f654,f675])).
fof(f8707,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f673])).
fof(f1979,plain,(
  ( ! [X9] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X9))))))) ) | $spl48),
  inference(cnf_transformation,[],[f1979_D])).
fof(f1979_D,plain,(
  ( ! [X9] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X9))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1978,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | $spl46),
  inference(cnf_transformation,[],[f1978_D])).
fof(f1978_D,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f8703,plain,(
  ( ! [X9] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X9))))))) )),
  inference(resolution,[],[f654,f952])).
fof(f952,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f477])).
fof(f477,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f476])).
fof(f476,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(ennf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))) => (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(flattening,[],[f381])).
fof(f381,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))) => (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ! [X24,X25,X26] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X24),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X25),s('type.list.list'(X6),X26)))))) => (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X24),s('type.list.list'(X6),X25)))) | p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X25),s('type.list.list'(X6),X24))))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_APPEND2')).
fof(f1976,plain,(
  ( ! [X8] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X8))))))) ) | $spl44),
  inference(cnf_transformation,[],[f1976_D])).
fof(f1976_D,plain,(
  ( ! [X8] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X8))))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f8702,plain,(
  ( ! [X8] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X8))))))) )),
  inference(resolution,[],[f654,f952])).
fof(f8686,plain,(
  ( ! [X19] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1269,f977])).
fof(f1269,plain,(
  ( ! [X160] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)),X160)) )),
  inference(resolution,[],[f653,f800])).
fof(f8650,plain,(
  ( ! [X19] : (sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1268,f977])).
fof(f1268,plain,(
  ( ! [X159] : (sP11(X159,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f799])).
fof(f1366,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl30),
  inference(cnf_transformation,[],[f1366_D])).
fof(f1366_D,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f8578,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f653,f977])).
fof(f1362,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl28),
  inference(cnf_transformation,[],[f1362_D])).
fof(f1362_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8552,plain,(
  ( ! [X169,X170] : (p(s(bool,X169)) | p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f978])).
fof(f8461,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f653,f679])).
fof(f8458,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f678])).
fof(f8457,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f677])).
fof(f8456,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f653,f675])).
fof(f8454,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f673])).
fof(f8198,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f7956,f652])).
fof(f7956,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) | $spl24),
  inference(backward_demodulation,[],[f1359,f654])).
fof(f7880,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f978])).
fof(f7785,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl23),
  inference(resolution,[],[f1357,f679])).
fof(f7784,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f678])).
fof(f7782,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f677])).
fof(f7780,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))) | s(bool,X18) != s(bool,X20)) ) | $spl23),
  inference(resolution,[],[f1357,f675])).
fof(f7778,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) ) | $spl23),
  inference(resolution,[],[f1357,f673])).
fof(f7755,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1257,f977])).
fof(f7739,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1257,f977])).
fof(f7734,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X6,X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1257,f863])).
fof(f7715,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1256,f977])).
fof(f7699,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1256,f977])).
fof(f7694,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1256,f863])).
fof(f7675,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1255,f977])).
fof(f7659,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1255,f977])).
fof(f7654,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1255,f863])).
fof(f7636,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1254,f977])).
fof(f7620,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1254,f977])).
fof(f7615,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1254,f863])).
fof(f7597,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1251,f977])).
fof(f7581,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1251,f977])).
fof(f7576,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1251,f863])).
fof(f7558,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1238,f977])).
fof(f7542,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1238,f977])).
fof(f7537,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1238,f863])).
fof(f7519,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1222,f977])).
fof(f7503,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1222,f977])).
fof(f7498,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1222,f863])).
fof(f7480,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1221,f977])).
fof(f7464,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1221,f977])).
fof(f7459,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1221,f863])).
fof(f7441,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1208,f977])).
fof(f7425,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1208,f977])).
fof(f7420,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1208,f863])).
fof(f5350,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl110),
  inference(cnf_transformation,[],[f5350_D])).
fof(f5350_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f7332,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f978])).
fof(f7237,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X33) != s(bool,X35)) ) | $spl19),
  inference(resolution,[],[f1175,f679])).
fof(f7236,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f678])).
fof(f7234,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f677])).
fof(f7232,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl19),
  inference(resolution,[],[f1175,f675])).
fof(f7230,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f673])).
fof(f3296,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl104),
  inference(cnf_transformation,[],[f3296_D])).
fof(f3296_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f7145,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f978])).
fof(f7050,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1161,f679])).
fof(f7049,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f678])).
fof(f7047,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f677])).
fof(f7045,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X18) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f1161,f675])).
fof(f7043,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f673])).
fof(f7021,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1061,f977])).
fof(f1061,plain,(
  ( ! [X146,X145] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X145,X146)) )),
  inference(resolution,[],[f652,f749])).
fof(f7005,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1061,f977])).
fof(f1175,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl19),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f7000,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X6,X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1061,f863])).
fof(f6982,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1060,f977])).
fof(f1060,plain,(
  ( ! [X144,X143] : (sP6(X143,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X144)) )),
  inference(resolution,[],[f652,f748])).
fof(f6966,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1060,f977])).
fof(f6961,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1060,f863])).
fof(f6943,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1059,f977])).
fof(f6927,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1059,f977])).
fof(f6922,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1059,f863])).
fof(f6905,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1058,f977])).
fof(f1058,plain,(
  ( ! [X140,X139] : (sP7(X139,X140,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f743])).
fof(f6889,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1058,f977])).
fof(f6884,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1058,f863])).
fof(f6867,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1055,f977])).
fof(f1055,plain,(
  ( ! [X134,X133] : (sP8(X133,X134,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f739])).
fof(f6851,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1055,f977])).
fof(f6846,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1055,f863])).
fof(f6829,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1042,f977])).
fof(f1042,plain,(
  ( ! [X107,X108] : (sP4(X107,X108,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f722])).
fof(f6813,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1042,f977])).
fof(f6808,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1042,f863])).
fof(f6791,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1026,f977])).
fof(f1026,plain,(
  ( ! [X76,X75] : (sP2(X75,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X76)) )),
  inference(resolution,[],[f652,f706])).
fof(f6775,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1026,f977])).
fof(f6770,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1026,f863])).
fof(f6753,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1025,f977])).
fof(f1025,plain,(
  ( ! [X74,X73] : (sP2(X73,X74,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f705])).
fof(f6737,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1025,f977])).
fof(f6732,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1025,f863])).
fof(f6715,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1012,f977])).
fof(f1012,plain,(
  ( ! [X47,X48] : (sP0(X47,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X48)) )),
  inference(resolution,[],[f652,f689])).
fof(f6699,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1012,f977])).
fof(f6694,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1012,f863])).
fof(f6635,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f655,f977])).
fof(f6630,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f655,f863])).
fof(f6610,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f978])).
fof(f6515,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f655,f679])).
fof(f6514,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f678])).
fof(f6512,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f677])).
fof(f6510,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f655,f675])).
fof(f6508,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f673])).
fof(f6448,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f654,f977])).
fof(f6443,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f654,f863])).
fof(f6423,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f978])).
fof(f6328,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f654,f679])).
fof(f6327,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f678])).
fof(f6325,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f677])).
fof(f6323,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f654,f675])).
fof(f6321,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f673])).
fof(f6317,plain,(
  ( ! [X9] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X9))))))) )),
  inference(resolution,[],[f654,f952])).
fof(f6316,plain,(
  ( ! [X8] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X8))))))) )),
  inference(resolution,[],[f654,f952])).
fof(f6300,plain,(
  ( ! [X19] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1073,f977])).
fof(f1073,plain,(
  ( ! [X160] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)),X160)) )),
  inference(resolution,[],[f652,f800])).
fof(f6264,plain,(
  ( ! [X19] : (sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1072,f977])).
fof(f1072,plain,(
  ( ! [X159] : (sP11(X159,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f799])).
fof(f1176,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl18),
  inference(cnf_transformation,[],[f1176_D])).
fof(f1176_D,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f6174,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f652,f977])).
fof(f1166,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) ) | $spl6),
  inference(cnf_transformation,[],[f1166_D])).
fof(f1166_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f6148,plain,(
  ( ! [X169,X170] : (p(s(bool,X169)) | p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f978])).
fof(f6057,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f652,f679])).
fof(f6054,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f678])).
fof(f6053,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f677])).
fof(f6052,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f652,f675])).
fof(f6050,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f673])).
fof(f2010,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK14) | $spl55),
  inference(cnf_transformation,[],[f2010_D])).
fof(f2010_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK14) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f6030,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK14) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f1807,f977])).
fof(f1807,plain,(
  s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),sK15)),
  inference(resolution,[],[f654,f956])).
fof(f956,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f642])).
fof(f6025,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f1807,f863])).
fof(f5803,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f5562,f872])).
fof(f872,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X1))))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0,X1] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X1))))),
  inference(flattening,[],[f336])).
fof(f336,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X1))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X11] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X6),X11))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_NIL|split|0')).
fof(f5562,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK15)))) | $spl54),
  inference(backward_demodulation,[],[f2009,f654])).
fof(f5559,plain,(
  ( ! [X170] : (p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl114),
  inference(cnf_transformation,[],[f5559_D])).
fof(f5559_D,plain,(
  ( ! [X170] : (p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f5375,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl18),
  inference(resolution,[],[f1176,f673])).
fof(f5377,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X16) != s(bool,X17)) ) | $spl18),
  inference(resolution,[],[f1176,f675])).
fof(f5378,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl18),
  inference(resolution,[],[f1176,f677])).
fof(f5379,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl18),
  inference(resolution,[],[f1176,f678])).
fof(f5382,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X31) != s(bool,X32)) ) | $spl18),
  inference(resolution,[],[f1176,f679])).
fof(f5473,plain,(
  ( ! [X169,X170] : (p(s(bool,X169)) | p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl18),
  inference(resolution,[],[f1176,f978])).
fof(f3960,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f1807,f863])).
fof(f3965,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK14) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f1807,f977])).
fof(f5556,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) | $spl113),
  inference(cnf_transformation,[],[f5556_D])).
fof(f5556_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f2009,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK14) | $spl54),
  inference(cnf_transformation,[],[f2009_D])).
fof(f2009_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK14) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5372,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK14) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.NIL_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) | $spl18),
  inference(resolution,[],[f1176,f955])).
fof(f5545,plain,(
  $false | ($spl18 | $spl55)),
  inference(subsumption_resolution,[],[f5544,f872])).
fof(f5544,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)))) | ($spl18 | $spl55)),
  inference(subsumption_resolution,[],[f5370,f2010])).
fof(f5370,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)))) | $spl18),
  inference(resolution,[],[f1176,f854])).
fof(f854,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(flattening,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(ennf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))) => s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))) => s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X11,X18] : ((p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X11),s('type.list.list'(X6),X18)))) & p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X18),s('type.list.list'(X6),X11))))) => s('type.list.list'(X6),X11) = s('type.list.list'(X6),X18))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_ANTISYM')).
fof(f5543,plain,(
  $false | ($spl18 | $spl55)),
  inference(subsumption_resolution,[],[f5542,f872])).
fof(f5542,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)))) | ($spl18 | $spl55)),
  inference(subsumption_resolution,[],[f5369,f2010])).
fof(f5369,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK14)))) | $spl18),
  inference(resolution,[],[f1176,f854])).
fof(f5540,plain,(
  $false | ($spl18 | $spl55)),
  inference(subsumption_resolution,[],[f5363,f2010])).
fof(f5363,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') = s('type.list.list'(sK12),sK14) | $spl18),
  inference(resolution,[],[f1176,f863])).
fof(f5281,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f978])).
fof(f5186,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X33) != s(bool,X35)) ) | $spl19),
  inference(resolution,[],[f1175,f679])).
fof(f5185,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f678])).
fof(f5183,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f677])).
fof(f5181,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl19),
  inference(resolution,[],[f1175,f675])).
fof(f5179,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl19),
  inference(resolution,[],[f1175,f673])).
fof(f5137,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl108),
  inference(cnf_transformation,[],[f5137_D])).
fof(f5137_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f5068,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl17),
  inference(resolution,[],[f1174,f978])).
fof(f4973,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X33) != s(bool,X35)) ) | $spl17),
  inference(resolution,[],[f1174,f679])).
fof(f4972,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl17),
  inference(resolution,[],[f1174,f678])).
fof(f4970,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl17),
  inference(resolution,[],[f1174,f677])).
fof(f4968,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl17),
  inference(resolution,[],[f1174,f675])).
fof(f4966,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl17),
  inference(resolution,[],[f1174,f673])).
fof(f1174,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl17),
  inference(cnf_transformation,[],[f1174_D])).
fof(f1174_D,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f4905,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) | $spl1),
  inference(superposition,[],[f1161,f977])).
fof(f4900,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl1),
  inference(superposition,[],[f1161,f863])).
fof(f4880,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f978])).
fof(f4785,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1161,f679])).
fof(f4784,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f678])).
fof(f4782,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f677])).
fof(f4780,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X18) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f1161,f675])).
fof(f4778,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f673])).
fof(f4756,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1061,f977])).
fof(f4751,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X6,X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1061,f863])).
fof(f4740,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1061,f977])).
fof(f4735,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X6,X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1061,f863])).
fof(f4717,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1060,f977])).
fof(f4712,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1060,f863])).
fof(f4701,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1060,f977])).
fof(f4696,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1060,f863])).
fof(f4678,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1059,f977])).
fof(f4673,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1059,f863])).
fof(f4662,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1059,f977])).
fof(f4657,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1059,f863])).
fof(f4640,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1058,f977])).
fof(f4635,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1058,f863])).
fof(f4624,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1058,f977])).
fof(f4619,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1058,f863])).
fof(f4602,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1055,f977])).
fof(f4597,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1055,f863])).
fof(f4586,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1055,f977])).
fof(f4581,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1055,f863])).
fof(f4564,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1042,f977])).
fof(f4559,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1042,f863])).
fof(f4548,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1042,f977])).
fof(f4543,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1042,f863])).
fof(f4526,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1026,f977])).
fof(f4521,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1026,f863])).
fof(f4510,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1026,f977])).
fof(f4505,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1026,f863])).
fof(f4488,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1025,f977])).
fof(f4483,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1025,f863])).
fof(f4472,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1025,f977])).
fof(f4467,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1025,f863])).
fof(f4450,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1012,f977])).
fof(f4445,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1012,f863])).
fof(f4434,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1012,f977])).
fof(f4429,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1012,f863])).
fof(f4370,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f655,f977])).
fof(f4365,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f655,f863])).
fof(f4345,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f978])).
fof(f4250,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f655,f679])).
fof(f4249,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f678])).
fof(f4247,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f677])).
fof(f4245,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f655,f675])).
fof(f4243,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f673])).
fof(f4183,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f654,f977])).
fof(f4178,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f654,f863])).
fof(f4158,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f978])).
fof(f4063,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f654,f679])).
fof(f4062,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f678])).
fof(f4060,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f677])).
fof(f4058,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f654,f675])).
fof(f4056,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f673])).
fof(f4052,plain,(
  ( ! [X9] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X9))))))) )),
  inference(resolution,[],[f654,f952])).
fof(f4051,plain,(
  ( ! [X8] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X8))))))) )),
  inference(resolution,[],[f654,f952])).
fof(f4035,plain,(
  ( ! [X19] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1073,f977])).
fof(f4030,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X4) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1073,f863])).
fof(f3999,plain,(
  ( ! [X19] : (sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1072,f977])).
fof(f3994,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1072,f863])).
fof(f3957,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK13) | $spl107),
  inference(cnf_transformation,[],[f3957_D])).
fof(f3957_D,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK13) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f3947,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK13) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) | $spl3),
  inference(superposition,[],[f1162,f977])).
fof(f1162,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK14) | $spl3),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK14) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f3942,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK13) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl3),
  inference(superposition,[],[f1162,f863])).
fof(f3896,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f652,f977])).
fof(f3891,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f652,f863])).
fof(f3870,plain,(
  ( ! [X169,X170] : (p(s(bool,X169)) | p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f978])).
fof(f3779,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f652,f679])).
fof(f3776,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f678])).
fof(f3775,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f677])).
fof(f3774,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f652,f675])).
fof(f3772,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f673])).
fof(f1165,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) | $spl5),
  inference(cnf_transformation,[],[f1165_D])).
fof(f1165_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f1163,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) | $spl2),
  inference(cnf_transformation,[],[f1163_D])).
fof(f1163_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3769,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14)))),
  inference(resolution,[],[f652,f955])).
fof(f1161,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) | $spl1),
  inference(cnf_transformation,[],[f1161_D])).
fof(f1161_D,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f3767,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))),
  inference(resolution,[],[f652,f854])).
fof(f3766,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))),
  inference(resolution,[],[f652,f854])).
fof(f3551,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f3309,f653])).
fof(f3309,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) | $spl2),
  inference(backward_demodulation,[],[f1163,f655])).
fof(f3256,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12)))) | $spl1),
  inference(superposition,[],[f1161,f977])).
fof(f3251,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | $spl1),
  inference(superposition,[],[f1161,f863])).
fof(f3231,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f978])).
fof(f3136,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1161,f679])).
fof(f3135,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f678])).
fof(f3133,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f677])).
fof(f3131,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | s(bool,X18) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f1161,f675])).
fof(f3129,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14)))) ) | $spl1),
  inference(resolution,[],[f1161,f673])).
fof(f3105,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X26,X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1061,f977])).
fof(f3100,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X6,X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1061,f863])).
fof(f3089,plain,(
  ( ! [X26,X27] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X26,X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1061,f977])).
fof(f3084,plain,(
  ( ! [X6,X7] : (sP6('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X6,X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1061,f863])).
fof(f3064,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1060,f977])).
fof(f3059,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1060,f863])).
fof(f3048,plain,(
  ( ! [X26,X27] : (sP6(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1060,f977])).
fof(f3043,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1060,f863])).
fof(f3023,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1059,f977])).
fof(f3018,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1059,f863])).
fof(f3007,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1059,f977])).
fof(f3002,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1059,f863])).
fof(f2983,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1058,f977])).
fof(f2978,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1058,f863])).
fof(f2967,plain,(
  ( ! [X26,X27] : (sP7(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1058,f977])).
fof(f2962,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1058,f863])).
fof(f2943,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1055,f977])).
fof(f2938,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1055,f863])).
fof(f2927,plain,(
  ( ! [X26,X27] : (sP8(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1055,f977])).
fof(f2922,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1055,f863])).
fof(f2903,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1042,f977])).
fof(f2898,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1042,f863])).
fof(f2887,plain,(
  ( ! [X26,X27] : (sP4(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1042,f977])).
fof(f2882,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1042,f863])).
fof(f2863,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1026,f977])).
fof(f2858,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1026,f863])).
fof(f2847,plain,(
  ( ! [X26,X27] : (sP2(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1026,f977])).
fof(f2842,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1026,f863])).
fof(f2823,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1025,f977])).
fof(f2818,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1025,f863])).
fof(f2807,plain,(
  ( ! [X26,X27] : (sP2(X26,X27,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1025,f977])).
fof(f2802,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1025,f863])).
fof(f2794,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X4)) ) | $spl70),
  inference(cnf_transformation,[],[f2794_D])).
fof(f2794_D,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X4)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2780,plain,(
  ( ! [X19] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl38),
  inference(superposition,[],[f1689,f977])).
fof(f2775,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')),X4) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) | $spl38),
  inference(superposition,[],[f1689,f863])).
fof(f2741,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X27) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1012,f977])).
fof(f2736,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1012,f863])).
fof(f2725,plain,(
  ( ! [X26,X27] : (sP0(X26,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X27) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1012,f977])).
fof(f2720,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X7) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1012,f863])).
fof(f2712,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f2712_D])).
fof(f2712_D,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2698,plain,(
  ( ! [X19] : (sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) ) | $spl32),
  inference(superposition,[],[f1649,f977])).
fof(f2693,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) ) | $spl32),
  inference(superposition,[],[f1649,f863])).
fof(f2659,plain,(
  ( ! [X19] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1269,f977])).
fof(f2654,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')),X4) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1269,f863])).
fof(f1689,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X4)) ) | $spl38),
  inference(cnf_transformation,[],[f1689_D])).
fof(f1689_D,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X4)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2643,plain,(
  ( ! [X19] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1269,f977])).
fof(f2638,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X4) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1269,f863])).
fof(f2622,plain,(
  ( ! [X19] : (sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1268,f977])).
fof(f2617,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1268,f863])).
fof(f1649,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) | $spl32),
  inference(cnf_transformation,[],[f1649_D])).
fof(f1649_D,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2606,plain,(
  ( ! [X19] : (sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))) )),
  inference(superposition,[],[f1268,f977])).
fof(f2601,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1268,f863])).
fof(f2544,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))),
  inference(superposition,[],[f655,f977])).
fof(f2539,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f655,f863])).
fof(f2519,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f978])).
fof(f2424,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f655,f679])).
fof(f2423,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f678])).
fof(f2421,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f677])).
fof(f2419,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f655,f675])).
fof(f2417,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) )),
  inference(resolution,[],[f655,f673])).
fof(f2402,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2251,f1169])).
fof(f2251,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f1169,f945])).
fof(f945,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f638])).
fof(f638,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f378])).
fof(f378,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f50])).
fof(f50,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.NOT_CLAUSES|split|1')).
fof(f2401,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2250,f1169])).
fof(f2250,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f1169,f943])).
fof(f943,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f637])).
fof(f637,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f377])).
fof(f377,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f49])).
fof(f49,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.NOT_CLAUSES|split|2')).
fof(f2400,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2245,f1169])).
fof(f2245,plain,(
  ( ! [X269] : (p(s(bool,X269))) ) | $spl12),
  inference(resolution,[],[f1169,f890])).
fof(f890,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f612])).
fof(f612,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f611])).
fof(f611,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f466])).
fof(f466,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X19] : ((p(s(bool,X19)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.IMP_CLAUSES|split|1')).
fof(f2399,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2243,f1169])).
fof(f2243,plain,(
  ( ! [X267] : (p(s(bool,X267))) ) | $spl12),
  inference(resolution,[],[f1169,f887])).
fof(f887,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f610])).
fof(f610,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f609])).
fof(f609,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f465])).
fof(f465,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X19] : ((p(s(bool,X19)) => p(s(bool,X19))) <=> p(s(bool,'T_0')))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.IMP_CLAUSES|split|3')).
fof(f2398,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2242,f1169])).
fof(f2242,plain,(
  ( ! [X266] : (p(s(bool,X266))) ) | $spl12),
  inference(resolution,[],[f1169,f884])).
fof(f884,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f608])).
fof(f608,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f607])).
fof(f607,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X19] : ((p(s(bool,'T_0')) => p(s(bool,X19))) <=> p(s(bool,X19)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.IMP_CLAUSES|split|0')).
fof(f2397,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2240,f1169])).
fof(f2240,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f1169,f881])).
fof(f881,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f606])).
fof(f606,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f605])).
fof(f605,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f463])).
fof(f463,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X19] : ((p(s(bool,'F_0')) => p(s(bool,X19))) <=> p(s(bool,'T_0')))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.IMP_CLAUSES|split|2')).
fof(f2252,plain,(
  $false | $spl12),
  inference(resolution,[],[f1169,f947])).
fof(f947,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f18])).
fof(f18,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.TRUTH')).
fof(f2253,plain,(
  $false | $spl12),
  inference(resolution,[],[f1169,f948])).
fof(f2254,plain,(
  $false | $spl12),
  inference(resolution,[],[f1169,f982])).
fof(f982,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f865])).
fof(f865,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f595])).
fof(f595,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.REFL_CLAUSE')).
fof(f2262,plain,(
  $false | $spl12),
  inference(resolution,[],[f1169,f871])).
fof(f871,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X1))))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X1))))),
  inference(flattening,[],[f334])).
fof(f334,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X1))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X11] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X11),s('type.list.list'(X6),X11))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_REFL')).
fof(f2274,plain,(
  $false | $spl12),
  inference(resolution,[],[f1169,f963])).
fof(f963,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))))),
  inference(flattening,[],[f397])).
fof(f397,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ! [X24,X26] : p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(X6),X24),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X24),s('type.list.list'(X6),X26))))))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.rich_list.IS_PREFIX_APPEND3')).
fof(f2279,plain,(
  $false | $spl12),
  inference(resolution,[],[f1169,f872])).
fof(f2280,plain,(
  $false | $spl12),
  inference(resolution,[],[f1169,f652])).
fof(f2281,plain,(
  $false | $spl12),
  inference(resolution,[],[f1169,f653])).
fof(f2117,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f2046])).
fof(f2046,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50)) | ~p(s(bool,X50))) ) | $spl10),
  inference(superposition,[],[f1168,f930])).
fof(f930,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f636])).
fof(f636,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X19] : (s(bool,'T_0') = s(bool,X19) <=> p(s(bool,X19)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.EQ_CLAUSES|split|0')).
fof(f2118,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f2045])).
fof(f2045,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48)) | ~p(s(bool,X48))) ) | $spl10),
  inference(superposition,[],[f1168,f928])).
fof(f2000,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK14) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f1807,f977])).
fof(f1995,plain,(
  s('type.list.list'(sK12),'const.list.NIL_0') != s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f1807,f863])).
fof(f1936,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK15) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK15,sK12)),s('type.list.list'(sK12),sK39(sK15,sK12))))),
  inference(superposition,[],[f654,f977])).
fof(f1931,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f654,f863])).
fof(f1911,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f978])).
fof(f1816,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f654,f679])).
fof(f1815,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f678])).
fof(f1813,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f677])).
fof(f1811,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f654,f675])).
fof(f1809,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(resolution,[],[f654,f673])).
fof(f1805,plain,(
  ( ! [X9] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X9))))))) )),
  inference(resolution,[],[f654,f952])).
fof(f1804,plain,(
  ( ! [X8] : (p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14)))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),X8))))))) )),
  inference(resolution,[],[f654,f952])).
fof(f1785,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) | $spl42),
  inference(cnf_transformation,[],[f1785_D])).
fof(f1785_D,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1778,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f1706])).
fof(f1706,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46)) | p(s(bool,X46))) ) | $spl8),
  inference(superposition,[],[f1167,f926])).
fof(f926,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f359])).
fof(f359,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X19] : (s(bool,'F_0') = s(bool,X19) <=> ~p(s(bool,X19)))),
  file('thm.rich_list.prefixes_is_prefix_total','thm.bool.EQ_CLAUSES|split|2')).
fof(f1779,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f1705])).
fof(f1705,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | p(s(bool,X44))) ) | $spl8),
  inference(superposition,[],[f1167,f924])).
fof(f924,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f633])).
fof(f1780,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f1703])).
fof(f1703,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40)) | p(s(bool,X40))) ) | $spl8),
  inference(superposition,[],[f1167,f878])).
fof(f1678,plain,(
  ( ! [X19] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X19) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1073,f977])).
fof(f1673,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')),X4) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1073,f863])).
fof(f1662,plain,(
  ( ! [X19] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X19) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1073,f977])).
fof(f1657,plain,(
  ( ! [X4] : (sP11('const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13)),X4) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1073,f863])).
fof(f1638,plain,(
  ( ! [X19] : (sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))) )),
  inference(superposition,[],[f1072,f977])).
fof(f1633,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1072,f863])).
fof(f1622,plain,(
  ( ! [X19] : (sP11(X19,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK14,sK12)),s('type.list.list'(sK12),sK39(sK14,sK12))))) )),
  inference(superposition,[],[f1072,f977])).
fof(f1617,plain,(
  ( ! [X4] : (sP11(X4,'const.list.isPREFIX_2'(s('type.list.list'(sK12),'const.list.NIL_0'),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f1072,f863])).
fof(f1609,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1487,f1170])).
fof(f1487,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(resolution,[],[f1170,f946])).
fof(f946,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f638])).
fof(f1608,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1486,f1170])).
fof(f1486,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(resolution,[],[f1170,f944])).
fof(f944,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f637])).
fof(f1607,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1483,f1170])).
fof(f1483,plain,(
  ( ! [X270] : (~p(s(bool,X270))) ) | $spl14),
  inference(resolution,[],[f1170,f938])).
fof(f1606,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1476,f1170])).
fof(f1476,plain,(
  ( ! [X263] : (~p(s(bool,X263))) ) | $spl14),
  inference(resolution,[],[f1170,f876])).
fof(f1485,plain,(
  $false | $spl14),
  inference(resolution,[],[f1170,f942])).
fof(f1317,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f653,f977])).
fof(f1312,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f653,f863])).
fof(f1291,plain,(
  ( ! [X169,X170] : (p(s(bool,X169)) | p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f978])).
fof(f1200,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f653,f679])).
fof(f1197,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f678])).
fof(f1196,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f677])).
fof(f1195,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f653,f675])).
fof(f1193,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f653,f673])).
fof(f1361,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) | $spl27),
  inference(cnf_transformation,[],[f1361_D])).
fof(f1361_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1359,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) | $spl24),
  inference(cnf_transformation,[],[f1359_D])).
fof(f1359_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1190,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15)))),
  inference(resolution,[],[f653,f955])).
fof(f1357,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) | $spl23),
  inference(cnf_transformation,[],[f1357_D])).
fof(f1357_D,plain,(
  ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1188,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))))),
  inference(resolution,[],[f653,f854])).
fof(f1187,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK15))))),
  inference(resolution,[],[f653,f854])).
fof(f1121,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK38(sK13,sK12)),s('type.list.list'(sK12),sK39(sK13,sK12))))),
  inference(superposition,[],[f652,f977])).
fof(f1116,plain,(
  p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),'const.list.NIL_0')))) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),'const.list.NIL_0'))))),
  inference(superposition,[],[f652,f863])).
fof(f1095,plain,(
  ( ! [X169,X170] : (p(s(bool,X169)) | p(s(bool,X170)) | s(bool,X170) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f978])).
fof(f1004,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f652,f679])).
fof(f1001,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f678])).
fof(f1000,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f677])).
fof(f999,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f652,f675])).
fof(f997,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK13)))) )),
  inference(resolution,[],[f652,f673])).
fof(f994,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14)))),
  inference(resolution,[],[f652,f955])).
fof(f992,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))),
  inference(resolution,[],[f652,f854])).
fof(f991,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) | ~p(s(bool,'const.list.isPREFIX_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))),
  inference(resolution,[],[f652,f854])).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 772
% Passive clauses: 21333
% Generated clauses: 46262
% Final active clauses: 699
% Final passive clauses: 18766
% Input formulas: 261
% Initial clauses: 326
% 
% Selected by SInE selection: 97
% SInE iterations: 1
% 
% Duplicate literals: 808
% Trivial inequalities: 109
% Fw subsumption resolutions: 4189
% Fw demodulations: 5669
% Bw demodulations: 58
% 
% Simple tautologies: 1504
% Equational tautologies: 3
% Forward subsumptions: 12376
% 
% Binary resolution: 8772
% Forward superposition: 21124
% Backward superposition: 5245
% Equality resolution: 4
% 
% Split clauses: 1155
% Split components: 2355
% SAT solver clauses: 1195
% SAT solver unit clauses: 24
% SAT solver binary clauses: 776
% 
% Sat splits: 1155
% Sat splitting refutations: 35
% 
% Memory used [KB]: 16630
% Time elapsed: 1.183 s
% ------------------------------
% ------------------------------
% Success in time 12.96 s
