Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ring
% (108605)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 3265
% Generated clauses: 21998
% Final active clauses: 113
% Final passive clauses: 2924
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 688
% Trivial inequalities: 258
% Fw subsumption resolutions: 2524
% Fw demodulations: 4199
% Bw demodulations: 25
% 
% Simple tautologies: 243
% Equational tautologies: 71
% Forward subsumptions: 10132
% 
% Binary resolution: 349
% Factoring: 1
% Forward superposition: 6601
% Backward superposition: 6995
% Self superposition: 252
% Equality factoring: 156
% Equality resolution: 17
% 
% Split clauses: 290
% Split components: 580
% SAT solver clauses: 768
% SAT solver unit clauses: 22
% SAT solver binary clauses: 113
% 
% Sat splits: 290
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 766
% TWLsolver calls for satisfiability: 152
% 
% Memory used [KB]: 14839
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (108619)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 134
% Passive clauses: 1215
% Generated clauses: 16398
% Final active clauses: 104
% Final passive clauses: 888
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 489
% Trivial inequalities: 92
% Fw subsumption resolutions: 866
% Fw demodulations: 1919
% Bw demodulations: 218
% Condensations: 1158
% 
% Simple tautologies: 148
% Equational tautologies: 39
% Forward subsumptions: 8587
% 
% Binary resolution: 325
% Factoring: 1
% Forward superposition: 5471
% Backward superposition: 4244
% Self superposition: 189
% Equality factoring: 78
% Equality resolution: 10
% 
% Split clauses: 48
% Split components: 96
% SAT solver clauses: 1922
% SAT solver unit clauses: 131
% SAT solver binary clauses: 39
% 
% Sat splits: 48
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 1922
% TWLsolver calls for satisfiability: 265
% 
% Memory used [KB]: 23794
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f28542,plain,(
  $false),
  inference(subsumption_resolution,[],[f28541,f214])).
fof(f214,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|9',aHL_TRUTH)).
fof(f28541,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f28539,f3203])).
fof(f3203,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f3108,f295])).
fof(f295,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f215,f216])).
fof(f216,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ring.ring_fupdfupds_comp|split|9',aHL_BOOL_CASES)).
fof(f215,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|9',aHL_FALSITY)).
fof(f3108,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f3004,f1209])).
fof(f1209,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1173,f266])).
fof(f266,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7,X15] : ! [X10,X9,X13] : s(X6,'const.combin.o_3'(s(fun(X7,X6),X10),s(fun(X15,X7),X9),s(X15,X13))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,happ(s(fun(X15,X7),X9),s(X15,X13)))))),
  file('thm.ring.ring_fupdfupds_comp|split|9','thm.combin.o_THM')).
fof(f1173,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f217,f264])).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X37,X38,X39] : (p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X38))),s(X37,X39)))) <=> s(X37,X38) = s(X37,X39))),
  file('thm.ring.ring_fupdfupds_comp|split|9',arity_0_HL_EQ)).
fof(f217,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X26] : ! [X34] : s(X26,X34) = s(X26,happ(s(fun(X26,X26),'_i_0_0'),s(X26,X34)))),
  file('thm.ring.ring_fupdfupds_comp|split|9','_i0_ax')).
fof(f3004,plain,(
  ( ! [X76,X74,X75,X73] : (s(X75,happ(s(fun(X73,X75),X76),s(X73,X74))) = s(X75,'const.combin.o_3'(s(fun(X73,X75),X76),s(fun(X73,X73),'_i_0_0'),s(X73,X74)))) )),
  inference(superposition,[],[f266,f217])).
fof(f28539,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),'const.combin.o_2'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK4),s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),'const.combin.o_2'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK4),s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f28459,f8277])).
fof(f8277,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),'const.combin.o_2'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK4),s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))))),
  inference(unit_resulting_resolution,[],[f213,f263])).
fof(f263,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f213,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),'const.combin.o_2'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK4),s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),'const.combin.o_2'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK4),s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(sK1,sK1),fun(sK1,sK1)),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f195])).
fof(f195,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),'const.combin.o_2'(s(fun(fun(X1,X1),fun(X1,X1)),X4),s(fun(fun(X1,X1),fun(X1,X1)),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) != s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),'const.combin.o_2'(s(fun(fun(X1,X1),fun(X1,X1)),X4),s(fun(fun(X1,X1),fun(X1,X1)),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),'const.combin.o_2'(s(fun(fun(X1,X1),fun(X1,X1)),X4),s(fun(fun(X1,X1),fun(X1,X1)),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X1,X1),fun(X1,X1)),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),'const.combin.o_2'(s(fun(fun(X7,X7),fun(X7,X7)),X10),s(fun(fun(X7,X7),fun(X7,X7)),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),X9))),s(fun(X6,'type.ring.ring'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),'const.combin.o_2'(s(fun(fun(X7,X7),fun(X7,X7)),X10),s(fun(fun(X7,X7),fun(X7,X7)),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),X9))),s(fun(X6,'type.ring.ring'(X7)),X8)))))),
  file('thm.ring.ring_fupdfupds_comp|split|9','thm.ring.ring_fupdfupds_comp|split|9')).
fof(f28459,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.o_2'(s(fun(fun(X0,X0),fun(X0,X0)),X1),s(fun(fun(X0,X0),fun(X0,X0)),X2))))),s(fun(X3,'type.ring.ring'(X0)),X4))) = s(fun(X3,'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),X1))),s(fun(X3,'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),X2))),s(fun(X3,'type.ring.ring'(X0)),X4)))))) )),
  inference(superposition,[],[f269,f227])).
fof(f227,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.o_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(fun(X0,X0),fun(X0,X0)),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),X1)))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.o_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(fun(X0,X0),fun(X0,X0)),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),X1)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.o_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(fun(X0,X0),fun(X0,X0)),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X0,X0),fun(X0,X0)),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),'const.combin.o_2'(s(fun(fun(X7,X7),fun(X7,X7)),X10),s(fun(fun(X7,X7),fun(X7,X7)),X9))))) = s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),X10))),s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RN_fupd_1'(s(fun(fun(X7,X7),fun(X7,X7)),X9)))))),
  file('thm.ring.ring_fupdfupds_comp|split|9','thm.ring.ring_fupdfupds_comp|split|8')).
fof(f269,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7,X18,X15] : ! [X10,X9,X8] : s(fun(X18,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X18,X7),'const.combin.o_2'(s(fun(X15,X7),X9),s(fun(X18,X15),X8))))) = s(fun(X18,X6),'const.combin.o_2'(s(fun(X15,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X15,X7),X9))),s(fun(X18,X15),X8)))),
  file('thm.ring.ring_fupdfupds_comp|split|9','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 242
% Passive clauses: 2792
% Generated clauses: 32165
% Final active clauses: 104
% Final passive clauses: 977
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 966
% Trivial inequalities: 180
% Fw subsumption resolutions: 1448
% Fw demodulations: 4967
% Bw demodulations: 1418
% Condensations: 1858
% Global subsumptions: 6
% 
% Simple tautologies: 181
% Equational tautologies: 294
% Forward subsumptions: 19048
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 25
% 
% Unit resulting resolution: 1406
% Factoring: 1
% Forward superposition: 8252
% Backward superposition: 7044
% Self superposition: 271
% Equality factoring: 132
% Equality resolution: 10
% 
% Split clauses: 31
% Split components: 62
% SAT solver clauses: 3192
% SAT solver unit clauses: 877
% SAT solver binary clauses: 1082
% 
% Sat splits: 31
% Sat splitting refutations: 330
% 
% Memory used [KB]: 25202
% Time elapsed: 1.112 s
% ------------------------------
% ------------------------------
% Success in time 2.44 s
