Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.sum
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.sum
% SZS output start Proof for thm.sum
fof(f7467,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7457,f2559,f145,f242,f144,f240,f238,f5712,f5605,f5596,f146,f243,f241,f239])).
fof(f239,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) != s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) | $spl1),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) != s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f241,plain,(
  ~p(s(bool,sK2)) | $spl3),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  ~p(s(bool,sK2)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f243,plain,(
  s(sK1,sK3) != s(sK1,sK5) | $spl5),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  s(sK1,sK3) != s(sK1,sK5) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f146,plain,(
  s(sK1,sK3) != s(sK1,sK5) | ~p(s(bool,sK2)) | s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) != s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4)))))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  (s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) = s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) | (p(s(bool,sK2)) & s(sK1,sK3) = s(sK1,sK5))) & (s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) != s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) | ~p(s(bool,sK2)) | s(sK1,sK3) != s(sK1,sK5))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f102])).
fof(f102,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X5))) = s('type.sum.sum'(X1,X0),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))),s('type.sum.sum'(X1,X0),'const.sum.INR_1'(s(X0,X4))))) | (p(s(bool,X2)) & s(X1,X3) = s(X1,X5))) & (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X5))) != s('type.sum.sum'(X1,X0),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))),s('type.sum.sum'(X1,X0),'const.sum.INR_1'(s(X0,X4))))) | ~p(s(bool,X2)) | s(X1,X3) != s(X1,X5)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X5))) = s('type.sum.sum'(X1,X0),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))),s('type.sum.sum'(X1,X0),'const.sum.INR_1'(s(X0,X4))))) | (p(s(bool,X2)) & s(X1,X3) = s(X1,X5))) & (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X5))) != s('type.sum.sum'(X1,X0),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))),s('type.sum.sum'(X1,X0),'const.sum.INR_1'(s(X0,X4))))) | (~p(s(bool,X2)) | s(X1,X3) != s(X1,X5))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X5))) = s('type.sum.sum'(X1,X0),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))),s('type.sum.sum'(X1,X0),'const.sum.INR_1'(s(X0,X4))))) <~> (p(s(bool,X2)) & s(X1,X3) = s(X1,X5)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X5))) = s('type.sum.sum'(X1,X0),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))),s('type.sum.sum'(X1,X0),'const.sum.INR_1'(s(X0,X4))))) <=> (p(s(bool,X2)) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X5))) = s('type.sum.sum'(X1,X0),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))),s('type.sum.sum'(X1,X0),'const.sum.INR_1'(s(X0,X4))))) <=> (p(s(bool,X2)) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X11))) = s('type.sum.sum'(X7,X6),'const.bool.COND_3'(s(bool,X8),s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X9))),s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X10))))) <=> (p(s(bool,X8)) & s(X7,X9) = s(X7,X11)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X11))) = s('type.sum.sum'(X7,X6),'const.bool.COND_3'(s(bool,X8),s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X9))),s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X10))))) <=> (p(s(bool,X8)) & s(X7,X9) = s(X7,X11)))),
  file('thm.sum.cond_sum_expand|split|2','thm.sum.cond_sum_expand|split|2')).
fof(f5596,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5595,f243])).
fof(f5595,plain,(
  s(sK1,sK3) = s(sK1,sK5) | ($spl0 | $spl2)),
  inference(equality_resolution,[],[f2562])).
fof(f2562,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))) != s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,X0))) | s(sK1,sK5) = s(sK1,X0)) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f216,f2561])).
fof(f2561,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))) = s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f2560,f202])).
fof(f202,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X15] : ! [X13,X14] : s(X15,X13) = s(X15,'const.bool.COND_3'(s(bool,'T_0'),s(X15,X13),s(X15,X14)))),
  file('thm.sum.cond_sum_expand|split|2','thm.bool.COND_CLAUSES|split|0')).
fof(f2560,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) = s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f238,f254])).
fof(f254,plain,(
  s(bool,'T_0') = s(bool,sK2) | $spl2),
  inference(resolution,[],[f158,f240])).
fof(f158,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.sum.cond_sum_expand|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f216,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) != s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : ((s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) != s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) & (s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3)))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X19,X15] : ! [X10,X9] : (s('type.sum.sum'(X15,X19),'const.sum.INL_1'(s(X15,X9))) = s('type.sum.sum'(X15,X19),'const.sum.INL_1'(s(X15,X10))) <=> s(X15,X9) = s(X15,X10))),
  file('thm.sum.cond_sum_expand|split|2','thm.sum.INR_INL_|49||49||split|0')).
fof(f5605,plain,(
  s(bool,'F_0') = s(bool,sK2) | $spl62),
  inference(cnf_transformation,[],[f5605_D])).
fof(f5605_D,plain,(
  s(bool,'F_0') = s(bool,sK2) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5712,plain,(
  s(bool,'F_0') = s(bool,sK2) | $spl3),
  inference(resolution,[],[f241,f162])).
fof(f162,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.sum.cond_sum_expand|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f238,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) = s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) | $spl0),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) = s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f240,plain,(
  p(s(bool,sK2)) | $spl2),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  p(s(bool,sK2)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f144,plain,(
  p(s(bool,sK2)) | s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) = s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4)))))),
  inference(cnf_transformation,[],[f103])).
fof(f242,plain,(
  s(sK1,sK3) = s(sK1,sK5) | $spl4),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  s(sK1,sK3) = s(sK1,sK5) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f145,plain,(
  s(sK1,sK3) = s(sK1,sK5) | s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) = s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4)))))),
  inference(cnf_transformation,[],[f103])).
fof(f2559,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f2558])).
fof(f2558,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))) != s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f2557,f242])).
fof(f2557,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))) != s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2556,f202])).
fof(f2556,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) != s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f239,f254])).
fof(f7457,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f7433,f220])).
fof(f220,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ! [X2,X3] : ~s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X15,X19] : ! [X9,X10] : ~s('type.sum.sum'(X15,X19),'const.sum.INL_1'(s(X15,X9))) = s('type.sum.sum'(X15,X19),'const.sum.INR_1'(s(X19,X10)))),
  file('thm.sum.cond_sum_expand|split|2','thm.sum.sum_distinct')).
fof(f7433,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) = s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))) | ($spl0 | $spl62)),
  inference(superposition,[],[f7146,f201])).
fof(f201,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X15] : ! [X13,X14] : s(X15,X14) = s(X15,'const.bool.COND_3'(s(bool,'F_0'),s(X15,X13),s(X15,X14)))),
  file('thm.sum.cond_sum_expand|split|2','thm.bool.COND_CLAUSES|split|1')).
fof(f7146,plain,(
  s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK5))) = s('type.sum.sum'(sK1,sK0),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK1,sK0),'const.sum.INL_1'(s(sK1,sK3))),s('type.sum.sum'(sK1,sK0),'const.sum.INR_1'(s(sK0,sK4))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f238,f5605])).
% SZS output end Proof for thm.sum
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 75
% Passive clauses: 1360
% Generated clauses: 7591
% Final active clauses: 49
% Final passive clauses: 630
% Input formulas: 37
% Initial clauses: 91
% 
% Duplicate literals: 367
% Trivial inequalities: 153
% Fw subsumption resolutions: 953
% Fw demodulations: 936
% Bw demodulations: 3
% 
% Simple tautologies: 230
% Equational tautologies: 53
% Forward subsumptions: 2865
% 
% Binary resolution: 114
% Factoring: 1
% Forward superposition: 1941
% Backward superposition: 2206
% Self superposition: 111
% Equality factoring: 100
% Equality resolution: 26
% 
% Split clauses: 552
% Split components: 1172
% SAT solver clauses: 742
% SAT solver unit clauses: 13
% SAT solver binary clauses: 75
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 2
% 
% Sat splits: 552
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 726
% TWLsolver calls for satisfiability: 54
% 
% Memory used [KB]: 3326
% Time elapsed: 0.145 s
% ------------------------------
% ------------------------------
% Success in time 0.147 s
