Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.integer
% (116383)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 111
% Passive clauses: 1688
% Generated clauses: 24824
% Final active clauses: 102
% Final passive clauses: 1409
% Input formulas: 12
% Initial clauses: 18
% 
% Duplicate literals: 1269
% Trivial inequalities: 319
% Fw subsumption resolutions: 4086
% Fw demodulations: 3381
% Bw demodulations: 7
% 
% Simple tautologies: 377
% Equational tautologies: 71
% Forward subsumptions: 11094
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 158
% Factoring: 2
% Forward superposition: 6700
% Backward superposition: 8140
% Self superposition: 323
% Equality factoring: 217
% Equality resolution: 23
% 
% Split clauses: 1232
% Split components: 2489
% SAT solver clauses: 2148
% SAT solver unit clauses: 49
% SAT solver binary clauses: 221
% 
% Sat splits: 1232
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2131
% TWLsolver calls for satisfiability: 122
% 
% Memory used [KB]: 10234
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (116397)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 583
% Generated clauses: 9927
% Final active clauses: 78
% Final passive clauses: 404
% Input formulas: 12
% Initial clauses: 18
% 
% Duplicate literals: 929
% Trivial inequalities: 79
% Fw subsumption resolutions: 524
% Fw demodulations: 1032
% Bw demodulations: 86
% Condensations: 1378
% 
% Simple tautologies: 108
% Equational tautologies: 53
% Forward subsumptions: 4317
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 145
% Factoring: 2
% Forward superposition: 2085
% Backward superposition: 2474
% Self superposition: 205
% Equality factoring: 146
% Equality resolution: 14
% 
% Split clauses: 142
% Split components: 289
% SAT solver clauses: 959
% SAT solver unit clauses: 25
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 142
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 958
% TWLsolver calls for satisfiability: 133
% 
% Memory used [KB]: 8315
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (116425)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 446
% Passive clauses: 6228
% Generated clauses: 142675
% Final active clauses: 174
% Final passive clauses: 905
% Input formulas: 12
% Initial clauses: 18
% Discarded non-redundant clauses: 735
% 
% Duplicate literals: 7612
% Trivial inequalities: 924
% Fw subsumption resolutions: 13224
% Fw demodulations: 20353
% Bw demodulations: 3513
% Condensations: 14758
% Global subsumptions: 1319
% 
% Simple tautologies: 1069
% Equational tautologies: 2649
% Forward subsumptions: 72736
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 106
% 
% Unit resulting resolution: 8622
% Factoring: 1
% Forward superposition: 37447
% Backward superposition: 26287
% Self superposition: 747
% Equality factoring: 400
% Equality resolution: 106
% 
% Split clauses: 468
% Split components: 968
% SAT solver clauses: 9802
% SAT solver unit clauses: 72
% SAT solver binary clauses: 1802
% 
% Sat splits: 468
% Sat splitting refutations: 1317
% 
% Memory used [KB]: 78804
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f13878,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f61,f62,f63,f64,f65,f66,f129,f130,f138,f140,f174,f175,f137,f180,f181,f221,f222,f223,f224,f225,f226,f259,f306,f308,f303,f304,f309,f425,f426,f432,f792,f992,f1123,f1133,f1132,f1131,f1238,f1241,f1250,f1253,f1268,f1306,f135,f1334,f136,f1360,f1586,f1587,f1694,f2046,f994,f2882,f2884,f2904,f2905,f3005,f3007,f3062,f3063,f3070,f3072,f3696,f3698,f3717,f3719,f3699,f3907,f3908,f3917,f3919,f4104,f4105,f4313,f4314,f4323,f4325,f4382,f4383,f4392,f4394,f4451,f4452,f4461,f4463,f4520,f4521,f4530,f4532,f4600,f4602,f4673,f4674,f4683,f4685,f4868,f4870,f4943,f4944,f4953,f4955,f5009,f5011,f5023,f5073,f5074,f5082,f5083,f5175,f5177,f5233,f5235,f5234,f5303,f5305,f5304,f5317,f5318,f5372,f5374,f5386,f5387,f5477,f5479,f5533,f5535,f5543,f5544,f5639,f5641,f5695,f5697,f5705,f5706,f5761,f5762,f5769,f5863,f5865,f5919,f5921,f5928,f5929,f6134,f6138,f6136,f6207,f6209,f6217,f6218,f6283,f6284,f5373,f6294,f6295,f6357,f6359,f6358,f6433,f6434,f6446,f6447,f6818,f6932,f6933,f6935,f6936,f6937,f6938,f6939,f6940,f6941,f6942,f6943,f6944,f6945,f6946,f6947,f6948,f6957,f6961,f6966,f6967,f6968,f6958,f6971,f6972,f6973,f6974,f6975,f6882,f6979,f6980,f6981,f6982,f6983,f6984,f7445,f7564,f7565,f7566,f7567,f7568,f7569,f7570,f7571,f7572,f7573,f7575,f7576,f7578,f7579,f7581,f7582,f7584,f7585,f7586,f7587,f7588,f7589,f7590,f7591,f7592,f7593,f7594,f7595,f7604,f7608,f7613,f7614,f7615,f7605,f7618,f7619,f7621,f7623,f7627,f7628,f7629,f7630,f7631,f7632,f8376,f8377,f5024,f9231,f9221,f9216,f9440,f9445,f9450,f9451,f9452,f9457,f9458,f9462,f9463,f9466,f9470,f9471,f9472,f9473,f9474,f9475,f9491,f9493,f9495,f9497,f9499,f9501,f9503,f9505,f9507,f9509,f9511,f9513,f9515,f9535,f4869,f9541,f9547,f5696,f9552,f9556,f9567,f9569,f5920,f9571,f9595,f4684,f9597,f4954,f9599,f4324,f9601,f4393,f9603,f4462,f9605,f4531,f9607,f3071,f9609,f3918,f9611,f5640,f9613,f5864,f9615,f4601,f9617,f5176,f9619,f5478,f9634,f9635,f9642,f9185,f9441,f9682,f9264,f10425,f6934,f10426,f6931,f10427,f10428,f10430,f10479,f10480,f10482,f10539,f7736,f7735,f10540,f10543,f7723,f7722,f7622,f7633,f10544,f3818,f3817,f10545,f10546,f3805,f3804,f6208,f3497,f3496,f3462,f3461,f3400,f3399,f3365,f3364,f3319,f3318,f3251,f3250,f3187,f3186,f3177,f3176,f3154,f3153,f3144,f3143,f3123,f3122,f3095,f3094,f5010,f2724,f2956,f2714,f2638,f10548,f2931,f2628,f10549,f2909,f2552,f2551,f2889,f2493,f5534,f2492,f5770,f10541,f10550,f10481,f2210,f2047,f995,f10552,f10551,f10553,f307,f10542,f10566,f10624,f10693,f10694,f77,f76,f10695,f10696,f78,f10697,f10698,f10699,f10700,f10709,f10701,f10702,f72,f10703,f10704,f74,f10711,f10705,f10706,f794,f10712,f10713,f70,f10751,f10752,f69,f10879,f10880,f3006,f75,f10881,f10882,f11086,f11087,f68,f11136,f11137,f11138,f11139,f11545,f11546,f11548,f11549,f11550,f11551,f11552,f11553,f11554,f11555,f11556,f11557,f11558,f11559,f11560,f11561,f11570,f11574,f11575,f11576,f11571,f11579,f10710,f11584,f11585,f11586,f11587,f12257,f12258,f139,f10547,f12677,f12939,f12941,f12940,f13008,f13009,f2883,f13013,f13080,f7514,f13093,f13095,f13094,f793,f13528,f13529,f2885,f993,f13846,f13847,f13848])).
fof(f13848,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f13773,f10757])).
fof(f10757,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl2),
  inference(backward_demodulation,[],[f69,f32])).
fof(f32,plain,(
  s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) & s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f23])).
fof(f23,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))))))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) => s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X8))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) => s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X8))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X8))))))),
  file('thm.integer.INT_DIVIDES_LADD','thm.integer.INT_DIVIDES_LADD')).
fof(f13773,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl2),
  inference(superposition,[],[f9973,f10837])).
fof(f10837,plain,(
  ( ! [X7] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',sK0))),s('type.integer.int',X7)))),s('type.integer.int',sK0)))) ) | $spl2),
  inference(subsumption_resolution,[],[f10778,f48])).
fof(f48,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f42])).
fof(f42,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11] : ! [X12] : (s(X11,X12) = s(X11,X12) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_DIVIDES_LADD','thm.bool.REFL_CLAUSE')).
fof(f10778,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',sK0))),s('type.integer.int',X7)))),s('type.integer.int',sK0)))) ) | $spl2),
  inference(superposition,[],[f1523,f69])).
fof(f1523,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int',X6)))) | s('type.integer.int',X6) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X6,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X4)))),s('type.integer.int',X5)))) )),
  inference(forward_demodulation,[],[f1472,f41])).
fof(f41,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X12,X13] : s('type.integer.int',X13) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int',X12)))),
  file('thm.integer.INT_DIVIDES_LADD','thm.integer.INT_ADD_SUB')).
fof(f1472,plain,(
  ( ! [X6,X4,X5] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X6,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X4)))),s('type.integer.int',X5))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X4))),s('type.integer.int',X6))))) )),
  inference(superposition,[],[f35,f41])).
fof(f35,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X1,X0)),s('type.integer.int',X0))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X1,X0)),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f28])).
fof(f28,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ? [X3] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(rectify,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ? [X2] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ? [X2] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> ? [X15] : s('type.integer.int',X7) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',X6))))),
  file('thm.integer.INT_DIVIDES_LADD','thm.integer.INT_DIVIDES')).
fof(f9973,plain,(
  ( ! [X16] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',sK0)))))))) )),
  inference(superposition,[],[f1460,f341])).
fof(f341,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X2))),s('type.integer.int',sK0)))) )),
  inference(superposition,[],[f34,f49])).
fof(f49,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)))),
  inference(resolution,[],[f31,f35])).
fof(f31,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))),
  inference(cnf_transformation,[],[f26])).
fof(f34,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))) )),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13,X14] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int',X14))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X14)))))),
  file('thm.integer.INT_DIVIDES_LADD','thm.integer.INT_RDISTRIB')).
fof(f1460,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',X14)))))) )),
  inference(subsumption_resolution,[],[f1455,f45])).
fof(f45,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_DIVIDES_LADD',aHL_FALSITY)).
fof(f1455,plain,(
  ( ! [X14,X15] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',X14)))))) )),
  inference(superposition,[],[f1348,f44])).
fof(f44,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_DIVIDES_LADD',aHL_BOOL_CASES)).
fof(f1348,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))) )),
  inference(equality_resolution,[],[f36])).
fof(f36,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f29])).
fof(f13847,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f13772,f10757])).
fof(f13772,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl2),
  inference(superposition,[],[f9973,f10830])).
fof(f10830,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0))) | $spl2),
  inference(subsumption_resolution,[],[f10770,f48])).
fof(f10770,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0))) | $spl2),
  inference(superposition,[],[f35,f69])).
fof(f13846,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f13769,f10757])).
fof(f13769,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl2),
  inference(superposition,[],[f9973,f10838])).
fof(f10838,plain,(
  ( ! [X8] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',sK2))),s('type.integer.int',X8)),sK0)),s('type.integer.int',sK0)))) ) | $spl2),
  inference(subsumption_resolution,[],[f10779,f48])).
fof(f10779,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',sK2))),s('type.integer.int',X8)),sK0)),s('type.integer.int',sK0)))) ) | $spl2),
  inference(superposition,[],[f1531,f69])).
fof(f1531,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X5)))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X4)),X6)),s('type.integer.int',X6)))) )),
  inference(forward_demodulation,[],[f1482,f41])).
fof(f1482,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X5)))) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X4))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X4)),X6)),s('type.integer.int',X6)))) )),
  inference(superposition,[],[f35,f41])).
fof(f993,plain,(
  ( ! [X42] : (s(bool,'F_0') = s(bool,X42)) ) | $spl28),
  inference(cnf_transformation,[],[f993_D])).
fof(f993_D,plain,(
  ( ! [X42] : (s(bool,'F_0') = s(bool,X42)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2885,plain,(
  ( ! [X78,X76,X74,X77] : (s(bool,X74) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X76),s('type.integer.int',X77))) | s(bool,'F_0') = s(bool,X74) | s('type.integer.int',X77) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X78),s('type.integer.int',X76)))) ) | $spl38),
  inference(cnf_transformation,[],[f2885_D])).
fof(f2885_D,plain,(
  ( ! [X78,X76,X74,X77] : (s(bool,X74) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X76),s('type.integer.int',X77))) | s(bool,'F_0') = s(bool,X74) | s('type.integer.int',X77) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X78),s('type.integer.int',X76)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f13529,plain,(
  ( ! [X66,X64,X67,X65,X63] : (s(bool,X63) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int',X66))) | s('type.integer.int',X66) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,X64) | s(bool,'F_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f13505,f1269])).
fof(f1269,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(resolution,[],[f45,f37])).
fof(f37,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.integer.INT_DIVIDES_LADD','thm.bool.IMP_ANTISYM_AX')).
fof(f13505,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X64)) | s(bool,X63) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int',X66))) | s('type.integer.int',X66) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,X64) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f1311,f1623])).
fof(f1623,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,X21) | s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,X21)) )),
  inference(resolution,[],[f1369,f1269])).
fof(f1369,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(resolution,[],[f1269,f40])).
fof(f40,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f25])).
fof(f1311,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) )),
  inference(resolution,[],[f36,f40])).
fof(f13528,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X58) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X60),s('type.integer.int',X61))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f13504,f1269])).
fof(f13504,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X59)) | s(bool,X58) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X60),s('type.integer.int',X61))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f1311,f1623])).
fof(f793,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,X39)) ) | $spl24),
  inference(cnf_transformation,[],[f793_D])).
fof(f793_D,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,X39)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f13094,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | $spl282),
  inference(cnf_transformation,[],[f13094_D])).
fof(f13094_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f13095,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X4))) ) | $spl16),
  inference(subsumption_resolution,[],[f13045,f106])).
fof(f106,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(forward_demodulation,[],[f105,f60])).
fof(f60,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))),
  inference(subsumption_resolution,[],[f53,f45])).
fof(f53,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))),
  inference(superposition,[],[f31,f44])).
fof(f105,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | ~p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f93,f48])).
fof(f93,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | ~p(s(bool,X7))) )),
  inference(superposition,[],[f40,f60])).
fof(f13045,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | $spl16),
  inference(superposition,[],[f834,f37])).
fof(f834,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'T_0') = s(bool,X16) | s(bool,'F_0') = s(bool,X15)) ) | $spl16),
  inference(resolution,[],[f820,f218])).
fof(f218,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10)) ) | $spl16),
  inference(trivial_inequality_removal,[],[f215])).
fof(f215,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,X10)) ) | $spl16),
  inference(superposition,[],[f137,f44])).
fof(f820,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl16),
  inference(forward_demodulation,[],[f819,f60])).
fof(f819,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | ~p(s(bool,X6))) ) | $spl16),
  inference(subsumption_resolution,[],[f799,f48])).
fof(f799,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | ~p(s(bool,'T_0')) | ~p(s(bool,X6))) ) | $spl16),
  inference(superposition,[],[f565,f40])).
fof(f565,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,X20) | s(bool,X20) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl16),
  inference(subsumption_resolution,[],[f558,f48])).
fof(f558,plain,(
  ( ! [X20] : (~p(s(bool,'T_0')) | s(bool,X20) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') != s(bool,X20)) ) | $spl16),
  inference(superposition,[],[f206,f60])).
fof(f206,plain,(
  ( ! [X4,X3] : (~p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') != s(bool,X3)) ) | $spl16),
  inference(resolution,[],[f137,f40])).
fof(f13093,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X1))) ) | $spl16),
  inference(subsumption_resolution,[],[f13044,f106])).
fof(f13044,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl16),
  inference(superposition,[],[f834,f37])).
fof(f7514,plain,(
  ( ! [X149,X147,X148] : (s(bool,X149) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X147),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X148),s('type.integer.int',X147))))) | s(bool,'T_0') = s(bool,X149)) ) | $spl164),
  inference(cnf_transformation,[],[f7514_D])).
fof(f7514_D,plain,(
  ( ! [X149,X147,X148] : (s(bool,X149) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X147),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X148),s('type.integer.int',X147))))) | s(bool,'T_0') = s(bool,X149)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f13080,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X43) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X42),s('type.integer.int',X41))))) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X40)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f13028])).
fof(f13028,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X43) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X42),s('type.integer.int',X41))))) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X40) | s(bool,'F_0') = s(bool,X40)) ) | $spl16),
  inference(superposition,[],[f834,f1625])).
fof(f1625,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X27),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',X27))))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(resolution,[],[f1369,f1348])).
fof(f13013,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,X74) | s(bool,X74) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X71))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl16),
  inference(subsumption_resolution,[],[f12990,f44])).
fof(f12990,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') != s(bool,X73) | s(bool,'T_0') = s(bool,X74) | s(bool,X74) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X71))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl16),
  inference(superposition,[],[f833,f1625])).
fof(f833,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,X13) | s(bool,'T_0') = s(bool,X14) | s(bool,X13) != s(bool,X14)) ) | $spl16),
  inference(resolution,[],[f820,f137])).
fof(f2883,plain,(
  ( ! [X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X69)) ) | $spl36),
  inference(cnf_transformation,[],[f2883_D])).
fof(f2883_D,plain,(
  ( ! [X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X69)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f13009,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,X39) | s(bool,X37) != s(bool,X39) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X37)) ) | $spl16),
  inference(subsumption_resolution,[],[f12975,f44])).
fof(f12975,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') != s(bool,X38) | s(bool,'T_0') = s(bool,X39) | s(bool,X37) != s(bool,X39) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X37)) ) | $spl16),
  inference(superposition,[],[f833,f1623])).
fof(f13008,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,X34) != s(bool,X36) | s(bool,'F_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X35)) ) | $spl16),
  inference(subsumption_resolution,[],[f12974,f44])).
fof(f12974,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,'T_0') = s(bool,X36) | s(bool,X34) != s(bool,X36) | s(bool,'F_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X35)) ) | $spl16),
  inference(superposition,[],[f833,f1623])).
fof(f12940,plain,(
  ( ! [X41,X42] : (s('type.integer.int',X42) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X42,X41)),s('type.integer.int',X41))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int',X42)))) ) | $spl280),
  inference(cnf_transformation,[],[f12940_D])).
fof(f12940_D,plain,(
  ( ! [X41,X42] : (s('type.integer.int',X42) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X42,X41)),s('type.integer.int',X41))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int',X42)))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f12941,plain,(
  ( ! [X45,X46,X44] : (s('type.integer.int',X45) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X45,X44)),s('type.integer.int',X44))) | s(bool,'F_0') = s(bool,X46) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X44),s('type.integer.int',X45)))) )),
  inference(subsumption_resolution,[],[f12837,f1269])).
fof(f12837,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,X46)) | s('type.integer.int',X45) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X45,X44)),s('type.integer.int',X44))) | s(bool,'F_0') = s(bool,X46) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X44),s('type.integer.int',X45)))) )),
  inference(superposition,[],[f35,f1623])).
fof(f12939,plain,(
  ( ! [X43,X41,X42] : (s('type.integer.int',X42) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X42,X41)),s('type.integer.int',X41))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int',X42))) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f12836,f1269])).
fof(f12836,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,X43)) | s('type.integer.int',X42) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X42,X41)),s('type.integer.int',X41))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int',X42))) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f35,f1623])).
fof(f12677,plain,(
  $false | ($spl2 | $spl14 | $spl34)),
  inference(subsumption_resolution,[],[f12595,f12334])).
fof(f12334,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | ($spl14 | $spl34)),
  inference(subsumption_resolution,[],[f1587,f10730])).
fof(f10730,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X17)) ) | $spl14),
  inference(resolution,[],[f136,f1369])).
fof(f12595,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ($spl2 | $spl14 | $spl34)),
  inference(superposition,[],[f10757,f12334])).
fof(f10547,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,X14) = s(bool,X15)) ) | $spl242),
  inference(cnf_transformation,[],[f10547_D])).
fof(f10547_D,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,X14) = s(bool,X15)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f139,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f12258,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,X81)) | s(bool,'T_0') != s(bool,X82) | s(bool,X82) = s(bool,X83)) ) | ($spl14 | $spl16 | $spl26)),
  inference(subsumption_resolution,[],[f12177,f1620])).
fof(f1620,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') = s(bool,X12) | s(bool,X12) = s(bool,X13)) ) | $spl16),
  inference(resolution,[],[f1369,f137])).
fof(f12177,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,X81)) | s(bool,'T_0') != s(bool,X82) | s(bool,X82) = s(bool,X83) | s(bool,'F_0') != s(bool,X83)) ) | ($spl14 | $spl16 | $spl26)),
  inference(superposition,[],[f3843,f12075])).
fof(f12075,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,X40)) ) | ($spl14 | $spl26)),
  inference(subsumption_resolution,[],[f794,f10730])).
fof(f3843,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | s(bool,'T_0') != s(bool,X19) | s(bool,X18) = s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl16),
  inference(subsumption_resolution,[],[f3797,f48])).
fof(f3797,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,'T_0') != s(bool,X19) | s(bool,X18) = s(bool,X19) | ~p(s(bool,'T_0')) | ~p(s(bool,X17))) ) | $spl16),
  inference(superposition,[],[f544,f40])).
fof(f544,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | s(bool,'T_0') != s(bool,X12) | s(bool,X12) = s(bool,X13)) ) | $spl16),
  inference(resolution,[],[f206,f137])).
fof(f12257,plain,(
  ( ! [X52,X51] : (~p(s(bool,X51)) | s(bool,'T_0') = s(bool,X52)) ) | ($spl14 | $spl16 | $spl26)),
  inference(subsumption_resolution,[],[f12167,f44])).
fof(f12167,plain,(
  ( ! [X52,X51] : (~p(s(bool,X51)) | s(bool,'F_0') != s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | ($spl14 | $spl16 | $spl26)),
  inference(superposition,[],[f820,f12075])).
fof(f11587,plain,(
  ( ! [X101,X102] : (s(bool,'T_0') != s(bool,X101) | s(bool,X102) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X101) = s(bool,X102)) ) | $spl278),
  inference(cnf_transformation,[],[f11587_D])).
fof(f11587_D,plain,(
  ( ! [X101,X102] : (s(bool,'T_0') != s(bool,X101) | s(bool,X102) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X101) = s(bool,X102)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f11586,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') != s(bool,X101) | s(bool,X101) = s(bool,X102) | s(bool,X102) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X100)) ) | ($spl14 | $spl16)),
  inference(subsumption_resolution,[],[f11452,f1269])).
fof(f11452,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,X100)) | s(bool,'T_0') != s(bool,X101) | s(bool,X101) = s(bool,X102) | s(bool,X102) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X100)) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3843,f10730])).
fof(f11585,plain,(
  ( ! [X99,X98] : (s(bool,'F_0') = s(bool,X98) | s(bool,X99) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X98) = s(bool,X99)) ) | $spl276),
  inference(cnf_transformation,[],[f11585_D])).
fof(f11585_D,plain,(
  ( ! [X99,X98] : (s(bool,'F_0') = s(bool,X98) | s(bool,X99) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X98) = s(bool,X99)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f11584,plain,(
  ( ! [X99,X97,X98] : (s(bool,'F_0') = s(bool,X98) | s(bool,X98) = s(bool,X99) | s(bool,X99) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X97)) ) | ($spl14 | $spl16)),
  inference(subsumption_resolution,[],[f11451,f1269])).
fof(f11451,plain,(
  ( ! [X99,X97,X98] : (~p(s(bool,X97)) | s(bool,'F_0') = s(bool,X98) | s(bool,X98) = s(bool,X99) | s(bool,X99) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X97)) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f2350,f10730])).
fof(f2350,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') = s(bool,X11) | s(bool,X10) = s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl16),
  inference(subsumption_resolution,[],[f2324,f48])).
fof(f2324,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,X10) = s(bool,X11) | ~p(s(bool,'T_0')) | ~p(s(bool,X9))) ) | $spl16),
  inference(superposition,[],[f1620,f40])).
fof(f10710,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X6)) ) | $spl252),
  inference(cnf_transformation,[],[f10710_D])).
fof(f10710_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f11579,plain,(
  ( ! [X61,X62] : (s(bool,X62) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X61)) ) | ($spl14 | $spl16)),
  inference(subsumption_resolution,[],[f11436,f1269])).
fof(f11436,plain,(
  ( ! [X61,X62] : (~p(s(bool,X61)) | s(bool,X62) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X61)) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f820,f10730])).
fof(f11571,plain,(
  ( ! [X43] : (s(bool,X43) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') != s(bool,X43)) ) | $spl272),
  inference(cnf_transformation,[],[f11571_D])).
fof(f11571_D,plain,(
  ( ! [X43] : (s(bool,X43) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') != s(bool,X43)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f11576,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X56) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X55)) ) | ($spl14 | $spl16)),
  inference(subsumption_resolution,[],[f11433,f44])).
fof(f11433,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') != s(bool,X55) | s(bool,'T_0') != s(bool,X56) | s(bool,X56) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X55)) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f544,f10730])).
fof(f11575,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | $spl274),
  inference(cnf_transformation,[],[f11575_D])).
fof(f11575_D,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f11574,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X52)) ) | $spl14),
  inference(subsumption_resolution,[],[f11432,f1269])).
fof(f11432,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,X52)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X52)) ) | $spl14),
  inference(superposition,[],[f521,f10730])).
fof(f521,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X9)) | s(bool,'T_0') = s(bool,X7) | s(bool,X8) = s(bool,X9) | s(bool,X7) = s(bool,X8)) )),
  inference(resolution,[],[f110,f40])).
fof(f110,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(resolution,[],[f106,f37])).
fof(f11570,plain,(
  ( ! [X43,X42] : (s(bool,X43) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') != s(bool,X43) | s(bool,'F_0') = s(bool,X42)) ) | ($spl14 | $spl16)),
  inference(subsumption_resolution,[],[f11427,f1269])).
fof(f11427,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | s(bool,X43) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') != s(bool,X43) | s(bool,'F_0') = s(bool,X42)) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f206,f10730])).
fof(f11561,plain,(
  ( ! [X19] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',sK2))),s('type.integer.int',X19)),sK0)),s('type.integer.int',sK0)))) ) | $spl270),
  inference(cnf_transformation,[],[f11561_D])).
fof(f11561_D,plain,(
  ( ! [X19] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',sK2))),s('type.integer.int',X19)),sK0)),s('type.integer.int',sK0)))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f11560,plain,(
  ( ! [X19,X18] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',sK2))),s('type.integer.int',X19)),sK0)),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X18)) ) | $spl14),
  inference(subsumption_resolution,[],[f11413,f1269])).
fof(f11413,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',sK2))),s('type.integer.int',X19)),sK0)),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X18)) ) | $spl14),
  inference(superposition,[],[f1531,f10730])).
fof(f11559,plain,(
  ( ! [X17] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK0))),s('type.integer.int',X17)))),s('type.integer.int',sK0)))) ) | $spl268),
  inference(cnf_transformation,[],[f11559_D])).
fof(f11559_D,plain,(
  ( ! [X17] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK0))),s('type.integer.int',X17)))),s('type.integer.int',sK0)))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f11558,plain,(
  ( ! [X17,X16] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK0))),s('type.integer.int',X17)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X16)) ) | $spl14),
  inference(subsumption_resolution,[],[f11412,f1269])).
fof(f11412,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK0))),s('type.integer.int',X17)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X16)) ) | $spl14),
  inference(superposition,[],[f1523,f10730])).
fof(f11557,plain,(
  ( ! [X15] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',X15))),s('type.integer.int',sK0)))) ) | $spl266),
  inference(cnf_transformation,[],[f11557_D])).
fof(f11557_D,plain,(
  ( ! [X15] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',X15))),s('type.integer.int',sK0)))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f11556,plain,(
  ( ! [X14,X15] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',X15))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X14)) ) | $spl14),
  inference(subsumption_resolution,[],[f11411,f1269])).
fof(f11411,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',X15))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X14)) ) | $spl14),
  inference(superposition,[],[f1502,f10730])).
fof(f1502,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X15))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X14,X15)),s('type.integer.int',X16))),s('type.integer.int',X15)))) )),
  inference(superposition,[],[f33,f35])).
fof(f33,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))) )),
  inference(cnf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0,X1,X2] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12,X13,X14] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X12),s('type.integer.int',X13))),s('type.integer.int',X14))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X14)))))),
  file('thm.integer.INT_DIVIDES_LADD','thm.integer.INT_SUB_RDISTRIB')).
fof(f11555,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK0))),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3(sK2,sK0)))),s('type.integer.int',sK0)))) ) | $spl264),
  inference(cnf_transformation,[],[f11555_D])).
fof(f11555_D,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK0))),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3(sK2,sK0)))),s('type.integer.int',sK0)))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f11554,plain,(
  ( ! [X12,X13] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK0))),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3(sK2,sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X12)) ) | $spl14),
  inference(subsumption_resolution,[],[f11410,f1269])).
fof(f11410,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK0))),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3(sK2,sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X12)) ) | $spl14),
  inference(superposition,[],[f1501,f10730])).
fof(f1501,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int',X11)))) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X12))),s('type.integer.int',X11))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3(X11,X12)))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f33,f35])).
fof(f11553,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',X11))),s('type.integer.int',sK0)))) ) | $spl262),
  inference(cnf_transformation,[],[f11553_D])).
fof(f11553_D,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',X11))),s('type.integer.int',sK0)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f11552,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',X11))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X10)) ) | $spl14),
  inference(subsumption_resolution,[],[f11409,f1269])).
fof(f11409,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',X11))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X10)) ) | $spl14),
  inference(superposition,[],[f1500,f10730])).
fof(f1500,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int',X8)))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',X9))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X8,X9)),s('type.integer.int',X10))),s('type.integer.int',X9)))) )),
  inference(superposition,[],[f34,f35])).
fof(f11551,plain,(
  ( ! [X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK0))),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK3(sK2,sK0)))),s('type.integer.int',sK0)))) ) | $spl260),
  inference(cnf_transformation,[],[f11551_D])).
fof(f11551_D,plain,(
  ( ! [X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK0))),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK3(sK2,sK0)))),s('type.integer.int',sK0)))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f11550,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK0))),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK3(sK2,sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X8)) ) | $spl14),
  inference(subsumption_resolution,[],[f11408,f1269])).
fof(f11408,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK0))),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK3(sK2,sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X8)) ) | $spl14),
  inference(superposition,[],[f1499,f10730])).
fof(f1499,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X5)))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int',X5))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',sK3(X5,X6)))),s('type.integer.int',X6)))) )),
  inference(superposition,[],[f34,f35])).
fof(f11549,plain,(
  ( ! [X7] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int',sK2))))) ) | $spl258),
  inference(cnf_transformation,[],[f11549_D])).
fof(f11549_D,plain,(
  ( ! [X7] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int',sK2))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f11548,plain,(
  ( ! [X6,X7] : (s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int',sK2)))) | s(bool,'F_0') = s(bool,X6)) ) | ($spl2 | $spl14)),
  inference(forward_demodulation,[],[f11547,f10830])).
fof(f11547,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int',sK2)))) | s(bool,'F_0') = s(bool,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f11407,f1269])).
fof(f11407,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int',sK2)))) | s(bool,'F_0') = s(bool,X6)) ) | $spl14),
  inference(superposition,[],[f1485,f10730])).
fof(f1485,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X13),s('type.integer.int',X12)))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X12,X13)),s('type.integer.int',X13))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X12,X13)),s('type.integer.int',X13)),X14)),s('type.integer.int',X14))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X14),s('type.integer.int',X12))))) )),
  inference(superposition,[],[f35,f35])).
fof(f11546,plain,(
  ( ! [X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)))),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK2),s('type.integer.int',X5))))) ) | $spl256),
  inference(cnf_transformation,[],[f11546_D])).
fof(f11546_D,plain,(
  ( ! [X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)))),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK2),s('type.integer.int',X5))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f11545,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)))),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK2),s('type.integer.int',X5)))) | s(bool,'F_0') = s(bool,X4)) ) | ($spl2 | $spl14)),
  inference(forward_demodulation,[],[f11544,f10830])).
fof(f11544,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK2),s('type.integer.int',X5)))) | s(bool,'F_0') = s(bool,X4)) ) | $spl14),
  inference(subsumption_resolution,[],[f11406,f1269])).
fof(f11406,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK2,sK0)),s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK2),s('type.integer.int',X5)))) | s(bool,'F_0') = s(bool,X4)) ) | $spl14),
  inference(superposition,[],[f1480,f10730])).
fof(f1480,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X13),s('type.integer.int',X12)))) | s('type.integer.int',X14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X14,'const.integer.int_mul_2'(s('type.integer.int',sK3(X12,X13)),s('type.integer.int',X13)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X12,X13)),s('type.integer.int',X13))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int',X14))))) )),
  inference(superposition,[],[f35,f35])).
fof(f11139,plain,(
  $false | ($spl2 | $spl12 | $spl16)),
  inference(trivial_inequality_removal,[],[f11132])).
fof(f11132,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl12 | $spl16)),
  inference(superposition,[],[f11094,f60])).
fof(f11094,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl2 | $spl12 | $spl16)),
  inference(subsumption_resolution,[],[f11093,f137])).
fof(f11093,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl2 | $spl12)),
  inference(forward_demodulation,[],[f78,f69])).
fof(f11138,plain,(
  $false | ($spl2 | $spl12 | $spl16)),
  inference(trivial_inequality_removal,[],[f11133])).
fof(f11133,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl12 | $spl16)),
  inference(superposition,[],[f11094,f69])).
fof(f11137,plain,(
  $false | ($spl2 | $spl12 | $spl16)),
  inference(trivial_inequality_removal,[],[f11135])).
fof(f11135,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl12 | $spl16)),
  inference(superposition,[],[f11094,f1460])).
fof(f11136,plain,(
  $false | ($spl2 | $spl12 | $spl16)),
  inference(equality_resolution,[],[f11094])).
fof(f68,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl0),
  inference(cnf_transformation,[],[f68_D])).
fof(f68_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f11087,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f11078,f1269])).
fof(f11078,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl9),
  inference(superposition,[],[f75,f1623])).
fof(f11086,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X10)) ) | $spl9),
  inference(subsumption_resolution,[],[f11077,f1269])).
fof(f11077,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X10)) ) | $spl9),
  inference(superposition,[],[f75,f1623])).
fof(f10882,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f10873,f44])).
fof(f10873,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl2),
  inference(superposition,[],[f10757,f1623])).
fof(f10881,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f10872,f44])).
fof(f10872,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X10)) ) | $spl2),
  inference(superposition,[],[f10757,f1623])).
fof(f75,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl9),
  inference(cnf_transformation,[],[f75_D])).
fof(f75_D,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f3006,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25)) ) | $spl40),
  inference(cnf_transformation,[],[f3006_D])).
fof(f3006_D,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f10880,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f10863,f137])).
fof(f10863,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl2),
  inference(superposition,[],[f10757,f40])).
fof(f10879,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f10862,f137])).
fof(f10862,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ~p(s(bool,X2))) ) | $spl2),
  inference(superposition,[],[f10757,f40])).
fof(f69,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | $spl2),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f10752,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | $spl14),
  inference(subsumption_resolution,[],[f10742,f106])).
fof(f10742,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | $spl14),
  inference(superposition,[],[f136,f520])).
fof(f520,plain,(
  ( ! [X6,X5] : (s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) )),
  inference(resolution,[],[f110,f106])).
fof(f10751,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f10741,f106])).
fof(f10741,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X6)) ) | $spl14),
  inference(superposition,[],[f136,f520])).
fof(f70,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | $spl3),
  inference(cnf_transformation,[],[f70_D])).
fof(f70_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f10713,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(forward_demodulation,[],[f10708,f60])).
fof(f10708,plain,(
  s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(superposition,[],[f32,f1644])).
fof(f1644,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1635,f48])).
fof(f1635,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | s(bool,X22) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1369,f60])).
fof(f10712,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(forward_demodulation,[],[f10707,f1460])).
fof(f10707,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X12))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) )),
  inference(superposition,[],[f32,f1625])).
fof(f794,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X40)) ) | $spl26),
  inference(cnf_transformation,[],[f794_D])).
fof(f794_D,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X40)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f10706,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) )),
  inference(superposition,[],[f32,f1623])).
fof(f10705,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f32,f1623])).
fof(f10711,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl254),
  inference(cnf_transformation,[],[f10711_D])).
fof(f10711_D,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X8)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f74,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | $spl7),
  inference(cnf_transformation,[],[f74_D])).
fof(f74_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f10704,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X9) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f32,f1623])).
fof(f10703,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X8) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f32,f1623])).
fof(f72,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl4),
  inference(cnf_transformation,[],[f72_D])).
fof(f72_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f10702,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) )),
  inference(superposition,[],[f32,f520])).
fof(f10701,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f32,f520])).
fof(f10709,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl250),
  inference(cnf_transformation,[],[f10709_D])).
fof(f10709_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f10700,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X5) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f32,f520])).
fof(f10699,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X4) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f32,f520])).
fof(f10698,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(superposition,[],[f32,f44])).
fof(f10697,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(superposition,[],[f32,f44])).
fof(f78,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ~p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ~p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f10696,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ~p(s(bool,X3)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) )),
  inference(superposition,[],[f32,f40])).
fof(f10695,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f32,f40])).
fof(f76,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl8),
  inference(cnf_transformation,[],[f76_D])).
fof(f76_D,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f77,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,X0))) ) | $spl10),
  inference(cnf_transformation,[],[f77_D])).
fof(f77_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,X0))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f10694,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f32,f37])).
fof(f10693,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,X0)) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) )),
  inference(superposition,[],[f32,f37])).
fof(f10624,plain,(
  $false | ($spl7 | $spl15 | $spl16)),
  inference(resolution,[],[f261,f45])).
fof(f261,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl7 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f260,f218])).
fof(f260,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl7 | $spl15)),
  inference(subsumption_resolution,[],[f254,f135])).
fof(f254,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | p(s(bool,X0))) ) | $spl7),
  inference(superposition,[],[f74,f37])).
fof(f10566,plain,(
  $false | $spl20),
  inference(resolution,[],[f307,f45])).
fof(f10542,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl237),
  inference(cnf_transformation,[],[f10542_D])).
fof(f10542_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f307,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl20),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f10553,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl16),
  inference(subsumption_resolution,[],[f216,f137])).
fof(f216,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X11)) ) | $spl16),
  inference(superposition,[],[f137,f44])).
fof(f10551,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl248),
  inference(cnf_transformation,[],[f10551_D])).
fof(f10551_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f10552,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X9) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl16),
  inference(subsumption_resolution,[],[f807,f565])).
fof(f807,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X9) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X9)) ) | $spl16),
  inference(superposition,[],[f565,f44])).
fof(f995,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl16),
  inference(forward_demodulation,[],[f938,f60])).
fof(f938,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl16),
  inference(equality_factoring,[],[f592])).
fof(f592,plain,(
  ( ! [X20] : (s(bool,X20) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X20)) ) | $spl16),
  inference(subsumption_resolution,[],[f585,f48])).
fof(f585,plain,(
  ( ! [X20] : (~p(s(bool,'T_0')) | s(bool,X20) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X20)) ) | $spl16),
  inference(superposition,[],[f230,f60])).
fof(f230,plain,(
  ( ! [X4,X3] : (~p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X3)) ) | $spl16),
  inference(resolution,[],[f218,f40])).
fof(f2047,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) )),
  inference(forward_demodulation,[],[f1994,f60])).
fof(f1994,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) )),
  inference(equality_factoring,[],[f1644])).
fof(f2210,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X9)) ) | $spl16),
  inference(forward_demodulation,[],[f2205,f60])).
fof(f2205,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X9)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2182])).
fof(f2182,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl16),
  inference(superposition,[],[f858,f44])).
fof(f858,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X22)) ) | $spl16),
  inference(subsumption_resolution,[],[f850,f48])).
fof(f850,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X22)) ) | $spl16),
  inference(superposition,[],[f820,f60])).
fof(f10481,plain,(
  ( ! [X78,X75] : (s(bool,X75) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) | $spl232),
  inference(cnf_transformation,[],[f10481_D])).
fof(f10481_D,plain,(
  ( ! [X78,X75] : (s(bool,X75) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f10550,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,X15)) ) | $spl16),
  inference(subsumption_resolution,[],[f2332,f1620])).
fof(f2332,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) ) | $spl16),
  inference(superposition,[],[f1620,f44])).
fof(f10541,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl236),
  inference(cnf_transformation,[],[f10541_D])).
fof(f10541_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f5770,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') = s(bool,X83) | s(bool,X83) != s(bool,X84) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl110),
  inference(cnf_transformation,[],[f5770_D])).
fof(f5770_D,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') = s(bool,X83) | s(bool,X83) != s(bool,X84) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2492,plain,(
  ( ! [X99,X98] : (s(bool,X98) != s(bool,X99) | s(bool,X99) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X98)) ) | $spl16),
  inference(superposition,[],[f565,f1623])).
fof(f5534,plain,(
  ( ! [X52,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X49) != s(bool,X52) | s(bool,X51) = s(bool,X52)) ) | $spl100),
  inference(cnf_transformation,[],[f5534_D])).
fof(f5534_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X49) != s(bool,X52) | s(bool,X51) = s(bool,X52)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f2493,plain,(
  ( ! [X101,X102,X100] : (s(bool,X100) != s(bool,X101) | s(bool,'F_0') = s(bool,X102) | s(bool,X101) = s(bool,X102) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X100)) ) | $spl16),
  inference(superposition,[],[f1620,f1623])).
fof(f2889,plain,(
  ( ! [X118,X117] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X117) != s(bool,X118) | s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,X117)) ) | $spl16),
  inference(forward_demodulation,[],[f2501,f60])).
fof(f2501,plain,(
  ( ! [X118,X117] : (s(bool,X117) != s(bool,X118) | s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X117)) ) | $spl16),
  inference(superposition,[],[f858,f1623])).
fof(f2551,plain,(
  ( ! [X99,X98] : (s(bool,X98) != s(bool,X99) | s(bool,X99) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X98) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl16),
  inference(superposition,[],[f565,f1623])).
fof(f2552,plain,(
  ( ! [X101,X102,X100] : (s(bool,X100) != s(bool,X101) | s(bool,'F_0') = s(bool,X102) | s(bool,X101) = s(bool,X102) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl16),
  inference(superposition,[],[f1620,f1623])).
fof(f2909,plain,(
  ( ! [X118,X117] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X117) != s(bool,X118) | s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,X117)) ) | $spl16),
  inference(forward_demodulation,[],[f2560,f60])).
fof(f2560,plain,(
  ( ! [X118,X117] : (s(bool,X117) != s(bool,X118) | s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,X117) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl16),
  inference(superposition,[],[f858,f1623])).
fof(f10549,plain,(
  ( ! [X57,X58] : (s(bool,X57) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X58) | s(bool,X57) = s(bool,X58)) ) | $spl246),
  inference(cnf_transformation,[],[f10549_D])).
fof(f10549_D,plain,(
  ( ! [X57,X58] : (s(bool,X57) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X58) | s(bool,X57) = s(bool,X58)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f2628,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X57) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58)) ) | $spl16),
  inference(superposition,[],[f565,f1623])).
fof(f2931,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70)) ) | $spl16),
  inference(forward_demodulation,[],[f2930,f60])).
fof(f2930,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70)) ) | $spl16),
  inference(subsumption_resolution,[],[f2633,f1620])).
fof(f2633,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70)) ) | $spl16),
  inference(superposition,[],[f858,f1623])).
fof(f10548,plain,(
  ( ! [X87,X85,X86] : (s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,X87) | s(bool,X85) = s(bool,X87) | s(bool,X85) = s(bool,X86)) ) | $spl244),
  inference(cnf_transformation,[],[f10548_D])).
fof(f10548_D,plain,(
  ( ! [X87,X85,X86] : (s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,X87) | s(bool,X85) = s(bool,X87) | s(bool,X85) = s(bool,X86)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f2638,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X86) | s(bool,X85) = s(bool,X86) | s(bool,X85) = s(bool,X87) | s(bool,'F_0') = s(bool,X87)) ) | $spl16),
  inference(superposition,[],[f1620,f1623])).
fof(f2714,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X57) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58)) ) | $spl16),
  inference(superposition,[],[f565,f1623])).
fof(f2956,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70)) ) | $spl16),
  inference(forward_demodulation,[],[f2955,f60])).
fof(f2955,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70)) ) | $spl16),
  inference(subsumption_resolution,[],[f2719,f1620])).
fof(f2719,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70)) ) | $spl16),
  inference(superposition,[],[f858,f1623])).
fof(f2724,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X86) | s(bool,X85) = s(bool,X86) | s(bool,X85) = s(bool,X87) | s(bool,'F_0') = s(bool,X87)) ) | $spl16),
  inference(superposition,[],[f1620,f1623])).
fof(f5010,plain,(
  ( ! [X30,X33,X32] : (s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33)) ) | $spl80),
  inference(cnf_transformation,[],[f5010_D])).
fof(f5010_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3094,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X51) | s(bool,X50) = s(bool,X52) | s(bool,'F_0') = s(bool,X52)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3095,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3122,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X50) = s(bool,X52) | s(bool,'F_0') = s(bool,X52)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3123,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X53) = s(bool,X55) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3143,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) = s(bool,X38) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X37)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3144,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f520,f1623])).
fof(f3153,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,X27) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3154,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3176,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) = s(bool,X38) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X37)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3177,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f520,f1623])).
fof(f3186,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,X27) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3187,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f520,f1623])).
fof(f3250,plain,(
  ( ! [X136,X135,X137] : (s(bool,X135) = s(bool,X136) | s(bool,X136) = s(bool,X137) | s(bool,'F_0') = s(bool,X137) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X135)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3251,plain,(
  ( ! [X140,X138,X139] : (s(bool,X138) = s(bool,X139) | s(bool,X139) = s(bool,X140) | s(bool,'F_0') = s(bool,X140) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3318,plain,(
  ( ! [X136,X135,X137] : (s(bool,X135) = s(bool,X136) | s(bool,X136) = s(bool,X137) | s(bool,'F_0') = s(bool,X137) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f1623,f520])).
fof(f3319,plain,(
  ( ! [X140,X138,X139] : (s(bool,X138) = s(bool,X139) | s(bool,X139) = s(bool,X140) | s(bool,'F_0') = s(bool,X140) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f1623,f520])).
fof(f3364,plain,(
  ( ! [X103,X105,X104] : (s(bool,X103) = s(bool,X104) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X104) | s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3365,plain,(
  ( ! [X107,X108,X106] : (s(bool,X106) = s(bool,X107) | s(bool,'F_0') = s(bool,X107) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3399,plain,(
  ( ! [X92,X93,X91] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X91) = s(bool,X92) | s(bool,'F_0') = s(bool,X92) | s(bool,X91) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3400,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X94) = s(bool,X95) | s(bool,'F_0') = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3461,plain,(
  ( ! [X103,X105,X104] : (s(bool,X103) = s(bool,X104) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X104) | s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3462,plain,(
  ( ! [X107,X108,X106] : (s(bool,X106) = s(bool,X107) | s(bool,'F_0') = s(bool,X107) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3496,plain,(
  ( ! [X92,X93,X91] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X91) = s(bool,X92) | s(bool,'F_0') = s(bool,X92) | s(bool,X91) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f1623,f520])).
fof(f3497,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X94) = s(bool,X95) | s(bool,'F_0') = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f1623,f520])).
fof(f6208,plain,(
  ( ! [X52,X51,X49] : (s(bool,'T_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X49) != s(bool,X52) | s(bool,X51) = s(bool,X52)) ) | $spl122),
  inference(cnf_transformation,[],[f6208_D])).
fof(f6208_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,'T_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X49) != s(bool,X52) | s(bool,X51) = s(bool,X52)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f3804,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) != s(bool,X38) | s(bool,'T_0') != s(bool,X39) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X37)) ) | $spl16),
  inference(superposition,[],[f544,f1623])).
fof(f3805,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'T_0') != s(bool,X42) | s(bool,X41) = s(bool,X42) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl16),
  inference(superposition,[],[f544,f1623])).
fof(f10546,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X15) | s(bool,X14) = s(bool,X15)) ) | $spl16),
  inference(subsumption_resolution,[],[f3811,f544])).
fof(f3811,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X15) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) ) | $spl16),
  inference(superposition,[],[f544,f44])).
fof(f10545,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X34) | s(bool,'F_0') = s(bool,X35) | s(bool,X33) = s(bool,X35) | s(bool,X33) = s(bool,X34)) ) | $spl240),
  inference(cnf_transformation,[],[f10545_D])).
fof(f10545_D,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X34) | s(bool,'F_0') = s(bool,X35) | s(bool,X33) = s(bool,X35) | s(bool,X33) = s(bool,X34)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f3817,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X31) | s(bool,X30) = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl16),
  inference(superposition,[],[f544,f1623])).
fof(f3818,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X34) | s(bool,X33) = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | $spl16),
  inference(superposition,[],[f544,f1623])).
fof(f10544,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20)))))) ) | $spl238),
  inference(cnf_transformation,[],[f10544_D])).
fof(f10544_D,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20)))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f7633,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) )),
  inference(forward_demodulation,[],[f7434,f1460])).
fof(f7434,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1))))) | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))) )),
  inference(equality_factoring,[],[f1625])).
fof(f7622,plain,(
  ( ! [X218,X215,X217,X214] : (s(bool,'F_0') = s(bool,X217) | s(bool,X217) != s(bool,X218) | s(bool,X218) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X214),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X215),s('type.integer.int',X214)))))) ) | $spl200),
  inference(cnf_transformation,[],[f7622_D])).
fof(f7622_D,plain,(
  ( ! [X218,X215,X217,X214] : (s(bool,'F_0') = s(bool,X217) | s(bool,X217) != s(bool,X218) | s(bool,X218) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X214),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X215),s('type.integer.int',X214)))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f7722,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X50) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int',X51))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X49)) ) | $spl16),
  inference(superposition,[],[f1438,f1623])).
fof(f1438,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X8) | s(bool,X8) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',X9)))))) ) | $spl16),
  inference(resolution,[],[f1348,f206])).
fof(f7723,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | s(bool,X54) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X56),s('type.integer.int',X55))))) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl16),
  inference(superposition,[],[f1438,f1623])).
fof(f10543,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X19) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f7729,f1438])).
fof(f7729,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X19) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))))) | s(bool,'T_0') = s(bool,X19)) ) | $spl16),
  inference(superposition,[],[f1438,f44])).
fof(f10540,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X45) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X46),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))))) | s(bool,'F_0') = s(bool,X48) | s(bool,X45) = s(bool,X48)) ) | $spl234),
  inference(cnf_transformation,[],[f10540_D])).
fof(f10540_D,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X45) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X46),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))))) | s(bool,'F_0') = s(bool,X48) | s(bool,X45) = s(bool,X48)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f7735,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X41) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X42),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42))))) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | $spl16),
  inference(superposition,[],[f1438,f1623])).
fof(f7736,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X45) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X46),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))))) | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | $spl16),
  inference(superposition,[],[f1438,f1623])).
fof(f10539,plain,(
  $false | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f10538,f47])).
fof(f47,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_LADD',aHL_TRUTH)).
fof(f10538,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f75,f72])).
fof(f10482,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) != s(bool,X73) | s(bool,X71) = s(bool,X74) | s(bool,'F_0') = s(bool,X74)) ) | $spl16),
  inference(subsumption_resolution,[],[f6271,f1369])).
fof(f6271,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X71)) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X73) | s(bool,X72) != s(bool,X73) | s(bool,X71) = s(bool,X74) | s(bool,'F_0') = s(bool,X74)) ) | $spl16),
  inference(superposition,[],[f830,f1623])).
fof(f830,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X5) | s(bool,X4) = s(bool,X6) | s(bool,X4) != s(bool,X5)) ) | $spl16),
  inference(resolution,[],[f820,f37])).
fof(f10480,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X77) | s(bool,X76) != s(bool,X77) | s(bool,X75) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) | $spl16),
  inference(subsumption_resolution,[],[f6272,f1369])).
fof(f6272,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X75)) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X77) | s(bool,X76) != s(bool,X77) | s(bool,X75) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) | $spl16),
  inference(superposition,[],[f830,f1623])).
fof(f10479,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') = s(bool,X80) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | $spl16),
  inference(subsumption_resolution,[],[f6273,f106])).
fof(f6273,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X79)) | s(bool,'T_0') = s(bool,X80) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | $spl16),
  inference(superposition,[],[f830,f44])).
fof(f10430,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) | $spl230),
  inference(cnf_transformation,[],[f10430_D])).
fof(f10430_D,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f10428,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X1)) ) | $spl4),
  inference(subsumption_resolution,[],[f6736,f1269])).
fof(f6736,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X1)) ) | $spl4),
  inference(superposition,[],[f35,f1645])).
fof(f1645,plain,(
  ( ! [X26] : (s(bool,X26) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X26)) ) | $spl4),
  inference(subsumption_resolution,[],[f1638,f48])).
fof(f1638,plain,(
  ( ! [X26] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X26)) ) | $spl4),
  inference(superposition,[],[f1369,f72])).
fof(f10427,plain,(
  ( ! [X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X5))))) ) | $spl228),
  inference(cnf_transformation,[],[f10427_D])).
fof(f10427_D,plain,(
  ( ! [X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X5))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f6931,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X5)))) | s(bool,'F_0') = s(bool,X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f6738,f1269])).
fof(f6738,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X5)))) | s(bool,'F_0') = s(bool,X4)) ) | $spl4),
  inference(superposition,[],[f1480,f1645])).
fof(f10426,plain,(
  ( ! [X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl226),
  inference(cnf_transformation,[],[f10426_D])).
fof(f10426_D,plain,(
  ( ! [X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f6934,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s(bool,'F_0') = s(bool,X6)) ) | $spl4),
  inference(subsumption_resolution,[],[f6739,f1269])).
fof(f6739,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s(bool,'F_0') = s(bool,X6)) ) | $spl4),
  inference(superposition,[],[f1485,f1645])).
fof(f10425,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f10416,f41])).
fof(f10416,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | ($spl6 | $spl8)),
  inference(superposition,[],[f10201,f391])).
fof(f391,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) | $spl8),
  inference(resolution,[],[f76,f35])).
fof(f10201,plain,(
  ( ! [X2] : (s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | $spl6),
  inference(superposition,[],[f290,f342])).
fof(f342,plain,(
  ( ! [X3] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X3),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0)))) )),
  inference(superposition,[],[f33,f49])).
fof(f290,plain,(
  ( ! [X0] : (s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK0)))) ) | $spl6),
  inference(subsumption_resolution,[],[f273,f45])).
fof(f273,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK0)))) ) | $spl6),
  inference(superposition,[],[f36,f73])).
fof(f73,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | $spl6),
  inference(cnf_transformation,[],[f73_D])).
fof(f73_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f9264,plain,(
  ( ! [X55,X53] : (s(bool,X53) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X53) = s(bool,X55) | s(bool,X55) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | $spl208),
  inference(cnf_transformation,[],[f9264_D])).
fof(f9264_D,plain,(
  ( ! [X55,X53] : (s(bool,X53) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X53) = s(bool,X55) | s(bool,X55) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f9682,plain,(
  ( ! [X103,X105,X104] : (s(bool,X103) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X104) | s(bool,X105) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X103) = s(bool,X105)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9074,f44])).
fof(f9074,plain,(
  ( ! [X103,X105,X104] : (s(bool,'F_0') != s(bool,X104) | s(bool,X103) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X104) | s(bool,X105) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X103) = s(bool,X105)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f620,f2782])).
fof(f2782,plain,(
  ( ! [X54,X53] : (s(bool,X53) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X54)) ) | ($spl6 | $spl22)),
  inference(trivial_inequality_removal,[],[f2712])).
fof(f2712,plain,(
  ( ! [X54,X53] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X53) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X54)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f478,f1623])).
fof(f478,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21) | s(bool,X21) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f473,f45])).
fof(f473,plain,(
  ( ! [X21] : (p(s(bool,'F_0')) | s(bool,X21) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') != s(bool,X21)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f313,f73])).
fof(f313,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') != s(bool,X5)) ) | $spl22),
  inference(resolution,[],[f309,f37])).
fof(f620,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f478,f44])).
fof(f9441,plain,(
  ( ! [X37,X35] : (s(bool,X35) = s(bool,X37) | s(bool,X35) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') != s(bool,X37)) ) | $spl210),
  inference(cnf_transformation,[],[f9441_D])).
fof(f9441_D,plain,(
  ( ! [X37,X35] : (s(bool,X35) = s(bool,X37) | s(bool,X35) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') != s(bool,X37)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f9185,plain,(
  ( ! [X187,X185,X186] : (s(bool,'T_0') != s(bool,X185) | s(bool,'F_0') = s(bool,X186) | s(bool,X187) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X185) = s(bool,X187)) ) | ($spl6 | $spl16 | $spl22)),
  inference(duplicate_literal_removal,[],[f9028])).
fof(f9028,plain,(
  ( ! [X187,X185,X186] : (s(bool,'T_0') != s(bool,X185) | s(bool,'F_0') = s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,X187) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X185) = s(bool,X187)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f1620,f2782])).
fof(f9642,plain,(
  ( ! [X116,X117,X115] : (s(bool,'T_0') != s(bool,X115) | s(bool,'T_0') != s(bool,X116) | s(bool,X117) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X115) = s(bool,X117)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f9007,f312])).
fof(f312,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') != s(bool,X4)) ) | ($spl16 | $spl22)),
  inference(resolution,[],[f309,f137])).
fof(f9007,plain,(
  ( ! [X116,X117,X115] : (s(bool,'T_0') != s(bool,X115) | s(bool,'T_0') != s(bool,X116) | s(bool,'F_0') = s(bool,X116) | s(bool,X117) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X115) = s(bool,X117)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f544,f2782])).
fof(f9635,plain,(
  ( ! [X81] : (s(bool,'F_0') != s(bool,X81)) ) | $spl224),
  inference(cnf_transformation,[],[f9635_D])).
fof(f9635_D,plain,(
  ( ! [X81] : (s(bool,'F_0') != s(bool,X81)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f9634,plain,(
  ( ! [X80,X81,X82] : (s(bool,'T_0') != s(bool,X80) | s(bool,'F_0') != s(bool,X81) | s(bool,X82) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X80) = s(bool,X82)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f8996,f312])).
fof(f8996,plain,(
  ( ! [X80,X81,X82] : (s(bool,'T_0') != s(bool,X80) | s(bool,'F_0') != s(bool,X81) | s(bool,'T_0') = s(bool,X81) | s(bool,X82) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X80) = s(bool,X82)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f368,f2782])).
fof(f368,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f312,f44])).
fof(f5478,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.integer.int',X69) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X69,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))),s('type.integer.int',X69)))) ) | $spl98),
  inference(cnf_transformation,[],[f5478_D])).
fof(f5478_D,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.integer.int',X69) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X69,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))),s('type.integer.int',X69)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f9619,plain,(
  ( ! [X356,X358,X360,X357,X359] : (s(bool,'F_0') = s(bool,X360) | s('type.integer.int',X359) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X359,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X356),s('type.integer.int',X358))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X357),s('type.integer.int',X358)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))),s('type.integer.int',X359)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9618,f1269])).
fof(f9618,plain,(
  ( ! [X356,X358,X360,X357,X359] : (s(bool,'F_0') = s(bool,X360) | ~p(s(bool,X360)) | s('type.integer.int',X359) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X359,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X356),s('type.integer.int',X358))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X357),s('type.integer.int',X358)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))),s('type.integer.int',X359)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8968,f73])).
fof(f8968,plain,(
  ( ! [X356,X358,X360,X357,X359] : (~p(s(bool,X360)) | s('type.integer.int',X359) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X359,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X356),s('type.integer.int',X358))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X357),s('type.integer.int',X358)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))))) | s(bool,X360) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))),s('type.integer.int',X359)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1524,f2782])).
fof(f1524,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))),s('type.integer.int',X8))),s('type.integer.int',X10)))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X10,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X8))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))),s('type.integer.int',X8)))))) )),
  inference(forward_demodulation,[],[f1473,f34])).
fof(f1473,plain,(
  ( ! [X10,X8,X7,X9] : (s('type.integer.int',X10) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X10,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X8))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))),s('type.integer.int',X8))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X8))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8))))),s('type.integer.int',X10))))) )),
  inference(superposition,[],[f35,f34])).
fof(f5176,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.integer.int',X69) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X69,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))),s('type.integer.int',X69)))) ) | $spl86),
  inference(cnf_transformation,[],[f5176_D])).
fof(f5176_D,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.integer.int',X69) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X69,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))),s('type.integer.int',X69)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f9617,plain,(
  ( ! [X352,X354,X351,X353,X355] : (s(bool,'F_0') = s(bool,X355) | s('type.integer.int',X354) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X354,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X351),s('type.integer.int',X353))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X352),s('type.integer.int',X353)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))),s('type.integer.int',X354)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9616,f1269])).
fof(f9616,plain,(
  ( ! [X352,X354,X351,X353,X355] : (s(bool,'F_0') = s(bool,X355) | ~p(s(bool,X355)) | s('type.integer.int',X354) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X354,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X351),s('type.integer.int',X353))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X352),s('type.integer.int',X353)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))),s('type.integer.int',X354)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8967,f73])).
fof(f8967,plain,(
  ( ! [X352,X354,X351,X353,X355] : (~p(s(bool,X355)) | s('type.integer.int',X354) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X354,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X351),s('type.integer.int',X353))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X352),s('type.integer.int',X353)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))))) | s(bool,X355) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))),s('type.integer.int',X354)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1522,f2782])).
fof(f1522,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int',X1))),s('type.integer.int',X3)))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X3,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int',X1)))))) )),
  inference(forward_demodulation,[],[f1471,f33])).
fof(f1471,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int',X3) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X3,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int',X1))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1))))),s('type.integer.int',X3))))) )),
  inference(superposition,[],[f35,f33])).
fof(f4601,plain,(
  ( ! [X54,X52,X53] : (s('type.integer.int',X54) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X54,'const.integer.int_mul_2'(s('type.integer.int',sK3(X52,X53)),s('type.integer.int',X53)))),s('type.integer.int',X52))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X52,X53)),s('type.integer.int',X53))),s('type.integer.int',X54))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X53),s('type.integer.int',X52))))) ) | $spl68),
  inference(cnf_transformation,[],[f4601_D])).
fof(f4601_D,plain,(
  ( ! [X54,X52,X53] : (s('type.integer.int',X54) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X54,'const.integer.int_mul_2'(s('type.integer.int',sK3(X52,X53)),s('type.integer.int',X53)))),s('type.integer.int',X52))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X52,X53)),s('type.integer.int',X53))),s('type.integer.int',X54))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X53),s('type.integer.int',X52))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9615,plain,(
  ( ! [X350,X349,X347,X348] : (s(bool,'F_0') = s(bool,X350) | s('type.integer.int',X349) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X349,'const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348)))),s('type.integer.int',X347))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X348),s('type.integer.int',X347)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348))),s('type.integer.int',X349)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9614,f1269])).
fof(f9614,plain,(
  ( ! [X350,X349,X347,X348] : (s(bool,'F_0') = s(bool,X350) | ~p(s(bool,X350)) | s('type.integer.int',X349) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X349,'const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348)))),s('type.integer.int',X347))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X348),s('type.integer.int',X347)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348))),s('type.integer.int',X349)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8966,f73])).
fof(f8966,plain,(
  ( ! [X350,X349,X347,X348] : (~p(s(bool,X350)) | s('type.integer.int',X349) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X349,'const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348)))),s('type.integer.int',X347))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X348),s('type.integer.int',X347)))) | s(bool,X350) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348))),s('type.integer.int',X349)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1475,f2782])).
fof(f1475,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X12,X13)),s('type.integer.int',X13))),s('type.integer.int',X14)))) | s('type.integer.int',X14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X14,'const.integer.int_mul_2'(s('type.integer.int',sK3(X12,X13)),s('type.integer.int',X13)))),s('type.integer.int',X12))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X13),s('type.integer.int',X12))))) )),
  inference(superposition,[],[f35,f35])).
fof(f5864,plain,(
  ( ! [X68,X66,X67,X65] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))),X65)),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68)))))) ) | $spl112),
  inference(cnf_transformation,[],[f5864_D])).
fof(f5864_D,plain,(
  ( ! [X68,X66,X67,X65] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))),X65)),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f9613,plain,(
  ( ! [X344,X346,X343,X345,X342] : (s(bool,'F_0') = s(bool,X346) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X343),s('type.integer.int',X345))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X344),s('type.integer.int',X345)))),X342)),s('type.integer.int',X342))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X342),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345)))))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9612,f1269])).
fof(f9612,plain,(
  ( ! [X344,X346,X343,X345,X342] : (s(bool,'F_0') = s(bool,X346) | ~p(s(bool,X346)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X343),s('type.integer.int',X345))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X344),s('type.integer.int',X345)))),X342)),s('type.integer.int',X342))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X342),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345)))))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8965,f73])).
fof(f8965,plain,(
  ( ! [X344,X346,X343,X345,X342] : (~p(s(bool,X346)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X343),s('type.integer.int',X345))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X344),s('type.integer.int',X345)))),X342)),s('type.integer.int',X342))) | s(bool,X346) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X342),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345)))))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1532,f2782])).
fof(f1532,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))),s('type.integer.int',X8)))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))),s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X8))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8)))),X10)),s('type.integer.int',X10)))) )),
  inference(forward_demodulation,[],[f1483,f34])).
fof(f1483,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))),s('type.integer.int',X8)))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X8))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X8))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8)))),X10)),s('type.integer.int',X10)))) )),
  inference(superposition,[],[f35,f34])).
fof(f5640,plain,(
  ( ! [X68,X66,X67,X65] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))),X65)),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68)))))) ) | $spl104),
  inference(cnf_transformation,[],[f5640_D])).
fof(f5640_D,plain,(
  ( ! [X68,X66,X67,X65] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))),X65)),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f9611,plain,(
  ( ! [X341,X337,X339,X340,X338] : (s(bool,'F_0') = s(bool,X341) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X338),s('type.integer.int',X340))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X339),s('type.integer.int',X340)))),X337)),s('type.integer.int',X337))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340)))))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9610,f1269])).
fof(f9610,plain,(
  ( ! [X341,X337,X339,X340,X338] : (s(bool,'F_0') = s(bool,X341) | ~p(s(bool,X341)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X338),s('type.integer.int',X340))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X339),s('type.integer.int',X340)))),X337)),s('type.integer.int',X337))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340)))))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8964,f73])).
fof(f8964,plain,(
  ( ! [X341,X337,X339,X340,X338] : (~p(s(bool,X341)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X338),s('type.integer.int',X340))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X339),s('type.integer.int',X340)))),X337)),s('type.integer.int',X337))) | s(bool,X341) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340)))))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1530,f2782])).
fof(f1530,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int',X1)))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1)))),X3)),s('type.integer.int',X3)))) )),
  inference(forward_demodulation,[],[f1481,f33])).
fof(f1481,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X2))),s('type.integer.int',X1)))))) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1)))),X3)),s('type.integer.int',X3)))) )),
  inference(superposition,[],[f35,f33])).
fof(f3918,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int',X60) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X60))),s('type.integer.int',X62)),X59)),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) | $spl50),
  inference(cnf_transformation,[],[f3918_D])).
fof(f3918_D,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int',X60) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X60))),s('type.integer.int',X62)),X59)),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f9609,plain,(
  ( ! [X325,X323,X324,X322] : (s(bool,'F_0') = s(bool,X324) | s('type.integer.int',X323) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X325),s('type.integer.int',X323))),s('type.integer.int',X325)),X322)),s('type.integer.int',X322))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X322),s('type.integer.int',X323)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9608,f1269])).
fof(f9608,plain,(
  ( ! [X325,X323,X324,X322] : (s(bool,'F_0') = s(bool,X324) | ~p(s(bool,X324)) | s('type.integer.int',X323) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X325),s('type.integer.int',X323))),s('type.integer.int',X325)),X322)),s('type.integer.int',X322))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X322),s('type.integer.int',X323)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8960,f73])).
fof(f8960,plain,(
  ( ! [X325,X323,X324,X322] : (~p(s(bool,X324)) | s('type.integer.int',X323) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X325),s('type.integer.int',X323))),s('type.integer.int',X325)),X322)),s('type.integer.int',X322))) | s(bool,X324) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X322),s('type.integer.int',X323)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1531,f2782])).
fof(f3071,plain,(
  ( ! [X43,X46,X44] : (s('type.integer.int',X44) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X44,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X46),s('type.integer.int',X43))),s('type.integer.int',X46)))),s('type.integer.int',X43))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl44),
  inference(cnf_transformation,[],[f3071_D])).
fof(f3071_D,plain,(
  ( ! [X43,X46,X44] : (s('type.integer.int',X44) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X44,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X46),s('type.integer.int',X43))),s('type.integer.int',X46)))),s('type.integer.int',X43))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f9607,plain,(
  ( ! [X321,X318,X320,X319] : (s(bool,'F_0') = s(bool,X320) | s('type.integer.int',X319) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X319,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X321),s('type.integer.int',X318))),s('type.integer.int',X321)))),s('type.integer.int',X318))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X318),s('type.integer.int',X319)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9606,f1269])).
fof(f9606,plain,(
  ( ! [X321,X318,X320,X319] : (s(bool,'F_0') = s(bool,X320) | ~p(s(bool,X320)) | s('type.integer.int',X319) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X319,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X321),s('type.integer.int',X318))),s('type.integer.int',X321)))),s('type.integer.int',X318))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X318),s('type.integer.int',X319)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8959,f73])).
fof(f8959,plain,(
  ( ! [X321,X318,X320,X319] : (~p(s(bool,X320)) | s('type.integer.int',X319) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X319,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X321),s('type.integer.int',X318))),s('type.integer.int',X321)))),s('type.integer.int',X318))) | s(bool,X320) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X318),s('type.integer.int',X319)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1523,f2782])).
fof(f4531,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X62))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) | $spl66),
  inference(cnf_transformation,[],[f4531_D])).
fof(f4531_D,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X62))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f9605,plain,(
  ( ! [X315,X316,X314,X317] : (s(bool,'F_0') = s(bool,X316) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X315),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X317),s('type.integer.int',X314))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X315,X314)),s('type.integer.int',X317))),s('type.integer.int',X314))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X314),s('type.integer.int',X315)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9604,f1269])).
fof(f9604,plain,(
  ( ! [X315,X316,X314,X317] : (s(bool,'F_0') = s(bool,X316) | ~p(s(bool,X316)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X315),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X317),s('type.integer.int',X314))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X315,X314)),s('type.integer.int',X317))),s('type.integer.int',X314))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X314),s('type.integer.int',X315)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8958,f73])).
fof(f8958,plain,(
  ( ! [X315,X316,X314,X317] : (~p(s(bool,X316)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X315),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X317),s('type.integer.int',X314))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X315,X314)),s('type.integer.int',X317))),s('type.integer.int',X314))) | s(bool,X316) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X314),s('type.integer.int',X315)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1502,f2782])).
fof(f4462,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))),s('type.integer.int',X60))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X62),s('type.integer.int',sK3(X60,X59)))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) | $spl62),
  inference(cnf_transformation,[],[f4462_D])).
fof(f4462_D,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))),s('type.integer.int',X60))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X62),s('type.integer.int',sK3(X60,X59)))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f9603,plain,(
  ( ! [X313,X310,X312,X311] : (s(bool,'F_0') = s(bool,X312) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X313),s('type.integer.int',X310))),s('type.integer.int',X311))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X313),s('type.integer.int',sK3(X311,X310)))),s('type.integer.int',X310))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X310),s('type.integer.int',X311)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9602,f1269])).
fof(f9602,plain,(
  ( ! [X313,X310,X312,X311] : (s(bool,'F_0') = s(bool,X312) | ~p(s(bool,X312)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X313),s('type.integer.int',X310))),s('type.integer.int',X311))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X313),s('type.integer.int',sK3(X311,X310)))),s('type.integer.int',X310))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X310),s('type.integer.int',X311)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8957,f73])).
fof(f8957,plain,(
  ( ! [X313,X310,X312,X311] : (~p(s(bool,X312)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X313),s('type.integer.int',X310))),s('type.integer.int',X311))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X313),s('type.integer.int',sK3(X311,X310)))),s('type.integer.int',X310))) | s(bool,X312) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X310),s('type.integer.int',X311)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1501,f2782])).
fof(f4393,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X62))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) | $spl58),
  inference(cnf_transformation,[],[f4393_D])).
fof(f4393_D,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X62))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f9601,plain,(
  ( ! [X308,X306,X309,X307] : (s(bool,'F_0') = s(bool,X308) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X307),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X309),s('type.integer.int',X306))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X307,X306)),s('type.integer.int',X309))),s('type.integer.int',X306))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X306),s('type.integer.int',X307)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9600,f1269])).
fof(f9600,plain,(
  ( ! [X308,X306,X309,X307] : (s(bool,'F_0') = s(bool,X308) | ~p(s(bool,X308)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X307),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X309),s('type.integer.int',X306))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X307,X306)),s('type.integer.int',X309))),s('type.integer.int',X306))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X306),s('type.integer.int',X307)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8956,f73])).
fof(f8956,plain,(
  ( ! [X308,X306,X309,X307] : (~p(s(bool,X308)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X307),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X309),s('type.integer.int',X306))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X307,X306)),s('type.integer.int',X309))),s('type.integer.int',X306))) | s(bool,X308) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X306),s('type.integer.int',X307)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1500,f2782])).
fof(f4324,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))),s('type.integer.int',X60))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',sK3(X60,X59)))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) | $spl54),
  inference(cnf_transformation,[],[f4324_D])).
fof(f4324_D,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))),s('type.integer.int',X60))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',sK3(X60,X59)))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f9599,plain,(
  ( ! [X302,X304,X303,X305] : (s(bool,'F_0') = s(bool,X304) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X305),s('type.integer.int',X302))),s('type.integer.int',X303))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X305),s('type.integer.int',sK3(X303,X302)))),s('type.integer.int',X302))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X302),s('type.integer.int',X303)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9598,f1269])).
fof(f9598,plain,(
  ( ! [X302,X304,X303,X305] : (s(bool,'F_0') = s(bool,X304) | ~p(s(bool,X304)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X305),s('type.integer.int',X302))),s('type.integer.int',X303))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X305),s('type.integer.int',sK3(X303,X302)))),s('type.integer.int',X302))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X302),s('type.integer.int',X303)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8955,f73])).
fof(f8955,plain,(
  ( ! [X302,X304,X303,X305] : (~p(s(bool,X304)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X305),s('type.integer.int',X302))),s('type.integer.int',X303))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X305),s('type.integer.int',sK3(X303,X302)))),s('type.integer.int',X302))) | s(bool,X304) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X302),s('type.integer.int',X303)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1499,f2782])).
fof(f4954,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59)),X62)),s('type.integer.int',X62))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X62),s('type.integer.int',X60))))) ) | $spl78),
  inference(cnf_transformation,[],[f4954_D])).
fof(f4954_D,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59)),X62)),s('type.integer.int',X62))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X62),s('type.integer.int',X60))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f9597,plain,(
  ( ! [X300,X298,X301,X299] : (s(bool,'F_0') = s(bool,X300) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298)),X301)),s('type.integer.int',X301))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X301),s('type.integer.int',X299)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int',X299)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9596,f1269])).
fof(f9596,plain,(
  ( ! [X300,X298,X301,X299] : (s(bool,'F_0') = s(bool,X300) | ~p(s(bool,X300)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298)),X301)),s('type.integer.int',X301))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X301),s('type.integer.int',X299)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int',X299)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8954,f73])).
fof(f8954,plain,(
  ( ! [X300,X298,X301,X299] : (~p(s(bool,X300)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298)),X301)),s('type.integer.int',X301))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X301),s('type.integer.int',X299)))) | s(bool,X300) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int',X299)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1485,f2782])).
fof(f4684,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int',X62) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X62,'const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X60),s('type.integer.int',X62))))) ) | $spl72),
  inference(cnf_transformation,[],[f4684_D])).
fof(f4684_D,plain,(
  ( ! [X59,X62,X60] : (s('type.integer.int',X62) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X62,'const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X60),s('type.integer.int',X62))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f9595,plain,(
  ( ! [X294,X296,X295,X297] : (s(bool,'F_0') = s(bool,X296) | s('type.integer.int',X297) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X297,'const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X295),s('type.integer.int',X297)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X294),s('type.integer.int',X295)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9594,f1269])).
fof(f9594,plain,(
  ( ! [X294,X296,X295,X297] : (s(bool,'F_0') = s(bool,X296) | ~p(s(bool,X296)) | s('type.integer.int',X297) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X297,'const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X295),s('type.integer.int',X297)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X294),s('type.integer.int',X295)))) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8953,f73])).
fof(f8953,plain,(
  ( ! [X294,X296,X295,X297] : (~p(s(bool,X296)) | s('type.integer.int',X297) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X297,'const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X295),s('type.integer.int',X297)))) | s(bool,X296) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X294),s('type.integer.int',X295)))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1480,f2782])).
fof(f9571,plain,(
  ( ! [X204,X202,X203,X201] : (s(bool,'F_0') = s(bool,X202) | s(bool,'T_0') != s(bool,X203) | s(bool,X203) = s(bool,X204) | s(bool,X201) != s(bool,X204) | s(bool,'F_0') = s(bool,X201)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f9570,f1269])).
fof(f9570,plain,(
  ( ! [X204,X202,X203,X201] : (s(bool,'F_0') = s(bool,X202) | ~p(s(bool,X202)) | s(bool,'T_0') != s(bool,X203) | s(bool,X203) = s(bool,X204) | s(bool,X201) != s(bool,X204) | s(bool,'F_0') = s(bool,X201)) ) | ($spl6 | $spl16 | $spl22)),
  inference(forward_demodulation,[],[f8914,f73])).
fof(f8914,plain,(
  ( ! [X204,X202,X203,X201] : (~p(s(bool,X202)) | s(bool,'T_0') != s(bool,X203) | s(bool,X203) = s(bool,X204) | s(bool,X201) != s(bool,X204) | s(bool,X202) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X201)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f3843,f2782])).
fof(f5920,plain,(
  ( ! [X52,X51,X49] : (s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,X49) != s(bool,X51) | s(bool,'T_0') = s(bool,X52)) ) | $spl114),
  inference(cnf_transformation,[],[f5920_D])).
fof(f5920_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,X49) != s(bool,X51) | s(bool,'T_0') = s(bool,X52)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f9569,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,'F_0') = s(bool,X198) | s(bool,X199) != s(bool,X200) | s(bool,'T_0') = s(bool,X200) | s(bool,X197) != s(bool,X199) | s(bool,'F_0') = s(bool,X197)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f9568,f1269])).
fof(f9568,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,'F_0') = s(bool,X198) | ~p(s(bool,X198)) | s(bool,X199) != s(bool,X200) | s(bool,'T_0') = s(bool,X200) | s(bool,X197) != s(bool,X199) | s(bool,'F_0') = s(bool,X197)) ) | ($spl6 | $spl16 | $spl22)),
  inference(forward_demodulation,[],[f8913,f73])).
fof(f8913,plain,(
  ( ! [X198,X200,X199,X197] : (~p(s(bool,X198)) | s(bool,X199) != s(bool,X200) | s(bool,'T_0') = s(bool,X200) | s(bool,X197) != s(bool,X199) | s(bool,X198) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X197)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f2999,f2782])).
fof(f2999,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2967,f48])).
fof(f2967,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,'T_0')) | ~p(s(bool,X9))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f368,f40])).
fof(f9567,plain,(
  ( ! [X189,X187,X190,X188] : (s(bool,'F_0') = s(bool,X188) | s(bool,'F_0') = s(bool,X189) | s(bool,X189) = s(bool,X190) | s(bool,X187) != s(bool,X190) | s(bool,'F_0') = s(bool,X187)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f9566,f1269])).
fof(f9566,plain,(
  ( ! [X189,X187,X190,X188] : (s(bool,'F_0') = s(bool,X188) | ~p(s(bool,X188)) | s(bool,'F_0') = s(bool,X189) | s(bool,X189) = s(bool,X190) | s(bool,X187) != s(bool,X190) | s(bool,'F_0') = s(bool,X187)) ) | ($spl6 | $spl16 | $spl22)),
  inference(forward_demodulation,[],[f8910,f73])).
fof(f8910,plain,(
  ( ! [X189,X187,X190,X188] : (~p(s(bool,X188)) | s(bool,'F_0') = s(bool,X189) | s(bool,X189) = s(bool,X190) | s(bool,X187) != s(bool,X190) | s(bool,X188) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X187)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f2350,f2782])).
fof(f9556,plain,(
  ( ! [X140,X138,X136,X139,X137] : (s(bool,'F_0') = s(bool,X137) | s(bool,X136) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X138),s('type.integer.int',X139))) | s('type.integer.int',X139) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X140),s('type.integer.int',X138))) | s(bool,'F_0') = s(bool,X136)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9555,f1269])).
fof(f9555,plain,(
  ( ! [X140,X138,X136,X139,X137] : (s(bool,'F_0') = s(bool,X137) | ~p(s(bool,X137)) | s(bool,X136) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X138),s('type.integer.int',X139))) | s('type.integer.int',X139) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X140),s('type.integer.int',X138))) | s(bool,'F_0') = s(bool,X136)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8895,f73])).
fof(f8895,plain,(
  ( ! [X140,X138,X136,X139,X137] : (~p(s(bool,X137)) | s(bool,X136) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X138),s('type.integer.int',X139))) | s('type.integer.int',X139) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X140),s('type.integer.int',X138))) | s(bool,X137) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X136)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1311,f2782])).
fof(f9552,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,X120) | s(bool,'F_0') = s(bool,X121) | s(bool,X119) = s(bool,X122) | s(bool,X121) != s(bool,X122) | s(bool,'F_0') = s(bool,X119)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9551,f1269])).
fof(f9551,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,X120) | ~p(s(bool,X120)) | s(bool,'F_0') = s(bool,X121) | s(bool,X119) = s(bool,X122) | s(bool,X121) != s(bool,X122) | s(bool,'F_0') = s(bool,X119)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8889,f73])).
fof(f8889,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,'F_0') = s(bool,X121) | s(bool,X119) = s(bool,X122) | s(bool,X121) != s(bool,X122) | s(bool,X120) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X119)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f653,f2782])).
fof(f653,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X15)) | s(bool,'F_0') = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | ($spl6 | $spl22)),
  inference(resolution,[],[f641,f40])).
fof(f641,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f640,f73])).
fof(f640,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,X5) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,X4))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f623,f45])).
fof(f623,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,X5) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,'F_0')) | p(s(bool,X4))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f478,f37])).
fof(f5696,plain,(
  ( ! [X52,X51,X49] : (s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') != s(bool,X51) | s(bool,X49) = s(bool,X52)) ) | $spl106),
  inference(cnf_transformation,[],[f5696_D])).
fof(f5696_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') != s(bool,X51) | s(bool,X49) = s(bool,X52)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f9547,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'F_0') = s(bool,X96) | s(bool,X97) != s(bool,X98) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') != s(bool,X97) | s(bool,'F_0') = s(bool,X95)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f9546,f1269])).
fof(f9546,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'F_0') = s(bool,X96) | ~p(s(bool,X96)) | s(bool,X97) != s(bool,X98) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') != s(bool,X97) | s(bool,'F_0') = s(bool,X95)) ) | ($spl6 | $spl16 | $spl22)),
  inference(forward_demodulation,[],[f8881,f73])).
fof(f8881,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | s(bool,X97) != s(bool,X98) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') != s(bool,X97) | s(bool,X96) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X95)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f599,f2782])).
fof(f599,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X15)) | s(bool,X13) != s(bool,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') != s(bool,X13)) ) | ($spl16 | $spl22)),
  inference(resolution,[],[f387,f40])).
fof(f387,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X4) != s(bool,X5)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f371,f45])).
fof(f371,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,'F_0')) | p(s(bool,X4))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f312,f37])).
fof(f9541,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'F_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,X78) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X75)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9540,f1269])).
fof(f9540,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'F_0') = s(bool,X76) | ~p(s(bool,X76)) | s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,X78) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X75)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8874,f73])).
fof(f8874,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,X78) | s(bool,X77) = s(bool,X78) | s(bool,X76) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X75)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f521,f2782])).
fof(f4869,plain,(
  ( ! [X52,X51,X49] : (s(bool,'F_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X51) = s(bool,X52) | s(bool,X49) = s(bool,X52)) ) | $spl74),
  inference(cnf_transformation,[],[f4869_D])).
fof(f4869_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,'F_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X51) = s(bool,X52) | s(bool,X49) = s(bool,X52)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f9535,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') != s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,X59) = s(bool,X60) | s(bool,'F_0') = s(bool,X57)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f9534,f1269])).
fof(f9534,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,'F_0') = s(bool,X58) | ~p(s(bool,X58)) | s(bool,'F_0') != s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,X59) = s(bool,X60) | s(bool,'F_0') = s(bool,X57)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f8868,f73])).
fof(f8868,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | s(bool,'F_0') != s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,X59) = s(bool,X60) | s(bool,X58) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X57)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f461,f2782])).
fof(f461,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X7) | s(bool,X8) = s(bool,X9) | s(bool,X7) = s(bool,X8)) ) | $spl22),
  inference(resolution,[],[f313,f40])).
fof(f9515,plain,(
  ( ! [X356,X358,X360,X357,X359] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))),s('type.integer.int',X359))) | s('type.integer.int',X359) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X359,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X356),s('type.integer.int',X358))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X357),s('type.integer.int',X358)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))))) | s(bool,'F_0') = s(bool,X360)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9514,f73])).
fof(f9514,plain,(
  ( ! [X356,X358,X360,X357,X359] : (s('type.integer.int',X359) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X359,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X356),s('type.integer.int',X358))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X357),s('type.integer.int',X358)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))),s('type.integer.int',X359))) | s(bool,'F_0') = s(bool,X360)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8846,f1269])).
fof(f8846,plain,(
  ( ! [X356,X358,X360,X357,X359] : (~p(s(bool,X360)) | s('type.integer.int',X359) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X359,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X356),s('type.integer.int',X358))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X357),s('type.integer.int',X358)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X356),s('type.integer.int',X357))),s('type.integer.int',X358))),s('type.integer.int',X359))) | s(bool,'F_0') = s(bool,X360)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1524,f2782])).
fof(f9513,plain,(
  ( ! [X352,X354,X351,X353,X355] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))),s('type.integer.int',X354))) | s('type.integer.int',X354) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X354,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X351),s('type.integer.int',X353))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X352),s('type.integer.int',X353)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))))) | s(bool,'F_0') = s(bool,X355)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9512,f73])).
fof(f9512,plain,(
  ( ! [X352,X354,X351,X353,X355] : (s('type.integer.int',X354) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X354,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X351),s('type.integer.int',X353))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X352),s('type.integer.int',X353)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))),s('type.integer.int',X354))) | s(bool,'F_0') = s(bool,X355)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8845,f1269])).
fof(f8845,plain,(
  ( ! [X352,X354,X351,X353,X355] : (~p(s(bool,X355)) | s('type.integer.int',X354) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X354,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X351),s('type.integer.int',X353))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X352),s('type.integer.int',X353)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X351),s('type.integer.int',X352))),s('type.integer.int',X353))),s('type.integer.int',X354))) | s(bool,'F_0') = s(bool,X355)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1522,f2782])).
fof(f9511,plain,(
  ( ! [X350,X349,X347,X348] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348))),s('type.integer.int',X349))) | s('type.integer.int',X349) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X349,'const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348)))),s('type.integer.int',X347))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X348),s('type.integer.int',X347)))) | s(bool,'F_0') = s(bool,X350)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9510,f73])).
fof(f9510,plain,(
  ( ! [X350,X349,X347,X348] : (s('type.integer.int',X349) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X349,'const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348)))),s('type.integer.int',X347))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X348),s('type.integer.int',X347)))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348))),s('type.integer.int',X349))) | s(bool,'F_0') = s(bool,X350)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8844,f1269])).
fof(f8844,plain,(
  ( ! [X350,X349,X347,X348] : (~p(s(bool,X350)) | s('type.integer.int',X349) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X349,'const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348)))),s('type.integer.int',X347))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X348),s('type.integer.int',X347)))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X347,X348)),s('type.integer.int',X348))),s('type.integer.int',X349))) | s(bool,'F_0') = s(bool,X350)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1475,f2782])).
fof(f9509,plain,(
  ( ! [X344,X346,X343,X345,X342] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X342),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X343),s('type.integer.int',X345))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X344),s('type.integer.int',X345)))),X342)),s('type.integer.int',X342))) | s(bool,'F_0') = s(bool,X346)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9508,f73])).
fof(f9508,plain,(
  ( ! [X344,X346,X343,X345,X342] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X343),s('type.integer.int',X345))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X344),s('type.integer.int',X345)))),X342)),s('type.integer.int',X342))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X342),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))))) | s(bool,'F_0') = s(bool,X346)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8843,f1269])).
fof(f8843,plain,(
  ( ! [X344,X346,X343,X345,X342] : (~p(s(bool,X346)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X343),s('type.integer.int',X345))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X344),s('type.integer.int',X345)))),X342)),s('type.integer.int',X342))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X342),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X343),s('type.integer.int',X344))),s('type.integer.int',X345))))) | s(bool,'F_0') = s(bool,X346)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1532,f2782])).
fof(f9507,plain,(
  ( ! [X341,X337,X339,X340,X338] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X338),s('type.integer.int',X340))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X339),s('type.integer.int',X340)))),X337)),s('type.integer.int',X337))) | s(bool,'F_0') = s(bool,X341)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9506,f73])).
fof(f9506,plain,(
  ( ! [X341,X337,X339,X340,X338] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X338),s('type.integer.int',X340))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X339),s('type.integer.int',X340)))),X337)),s('type.integer.int',X337))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))))) | s(bool,'F_0') = s(bool,X341)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8842,f1269])).
fof(f8842,plain,(
  ( ! [X341,X337,X339,X340,X338] : (~p(s(bool,X341)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X338),s('type.integer.int',X340))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X339),s('type.integer.int',X340)))),X337)),s('type.integer.int',X337))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X337),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X338),s('type.integer.int',X339))),s('type.integer.int',X340))))) | s(bool,'F_0') = s(bool,X341)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1530,f2782])).
fof(f9505,plain,(
  ( ! [X325,X323,X324,X322] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X322),s('type.integer.int',X323))) | s('type.integer.int',X323) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X325),s('type.integer.int',X323))),s('type.integer.int',X325)),X322)),s('type.integer.int',X322))) | s(bool,'F_0') = s(bool,X324)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9504,f73])).
fof(f9504,plain,(
  ( ! [X325,X323,X324,X322] : (s('type.integer.int',X323) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X325),s('type.integer.int',X323))),s('type.integer.int',X325)),X322)),s('type.integer.int',X322))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X322),s('type.integer.int',X323))) | s(bool,'F_0') = s(bool,X324)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8838,f1269])).
fof(f8838,plain,(
  ( ! [X325,X323,X324,X322] : (~p(s(bool,X324)) | s('type.integer.int',X323) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X325),s('type.integer.int',X323))),s('type.integer.int',X325)),X322)),s('type.integer.int',X322))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X322),s('type.integer.int',X323))) | s(bool,'F_0') = s(bool,X324)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1531,f2782])).
fof(f9503,plain,(
  ( ! [X321,X318,X320,X319] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X318),s('type.integer.int',X319))) | s('type.integer.int',X319) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X319,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X321),s('type.integer.int',X318))),s('type.integer.int',X321)))),s('type.integer.int',X318))) | s(bool,'F_0') = s(bool,X320)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9502,f73])).
fof(f9502,plain,(
  ( ! [X321,X318,X320,X319] : (s('type.integer.int',X319) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X319,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X321),s('type.integer.int',X318))),s('type.integer.int',X321)))),s('type.integer.int',X318))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X318),s('type.integer.int',X319))) | s(bool,'F_0') = s(bool,X320)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8837,f1269])).
fof(f8837,plain,(
  ( ! [X321,X318,X320,X319] : (~p(s(bool,X320)) | s('type.integer.int',X319) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X319,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X321),s('type.integer.int',X318))),s('type.integer.int',X321)))),s('type.integer.int',X318))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X318),s('type.integer.int',X319))) | s(bool,'F_0') = s(bool,X320)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1523,f2782])).
fof(f9501,plain,(
  ( ! [X315,X316,X314,X317] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X314),s('type.integer.int',X315))) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X315),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X317),s('type.integer.int',X314))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X315,X314)),s('type.integer.int',X317))),s('type.integer.int',X314))) | s(bool,'F_0') = s(bool,X316)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9500,f73])).
fof(f9500,plain,(
  ( ! [X315,X316,X314,X317] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X315),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X317),s('type.integer.int',X314))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X315,X314)),s('type.integer.int',X317))),s('type.integer.int',X314))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X314),s('type.integer.int',X315))) | s(bool,'F_0') = s(bool,X316)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8836,f1269])).
fof(f8836,plain,(
  ( ! [X315,X316,X314,X317] : (~p(s(bool,X316)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X315),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X317),s('type.integer.int',X314))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X315,X314)),s('type.integer.int',X317))),s('type.integer.int',X314))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X314),s('type.integer.int',X315))) | s(bool,'F_0') = s(bool,X316)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1502,f2782])).
fof(f9499,plain,(
  ( ! [X313,X310,X312,X311] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X310),s('type.integer.int',X311))) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X313),s('type.integer.int',X310))),s('type.integer.int',X311))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X313),s('type.integer.int',sK3(X311,X310)))),s('type.integer.int',X310))) | s(bool,'F_0') = s(bool,X312)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9498,f73])).
fof(f9498,plain,(
  ( ! [X313,X310,X312,X311] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X313),s('type.integer.int',X310))),s('type.integer.int',X311))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X313),s('type.integer.int',sK3(X311,X310)))),s('type.integer.int',X310))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X310),s('type.integer.int',X311))) | s(bool,'F_0') = s(bool,X312)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8835,f1269])).
fof(f8835,plain,(
  ( ! [X313,X310,X312,X311] : (~p(s(bool,X312)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X313),s('type.integer.int',X310))),s('type.integer.int',X311))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X313),s('type.integer.int',sK3(X311,X310)))),s('type.integer.int',X310))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X310),s('type.integer.int',X311))) | s(bool,'F_0') = s(bool,X312)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1501,f2782])).
fof(f9497,plain,(
  ( ! [X308,X306,X309,X307] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X306),s('type.integer.int',X307))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X307),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X309),s('type.integer.int',X306))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X307,X306)),s('type.integer.int',X309))),s('type.integer.int',X306))) | s(bool,'F_0') = s(bool,X308)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9496,f73])).
fof(f9496,plain,(
  ( ! [X308,X306,X309,X307] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X307),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X309),s('type.integer.int',X306))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X307,X306)),s('type.integer.int',X309))),s('type.integer.int',X306))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X306),s('type.integer.int',X307))) | s(bool,'F_0') = s(bool,X308)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8834,f1269])).
fof(f8834,plain,(
  ( ! [X308,X306,X309,X307] : (~p(s(bool,X308)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X307),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X309),s('type.integer.int',X306))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X307,X306)),s('type.integer.int',X309))),s('type.integer.int',X306))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X306),s('type.integer.int',X307))) | s(bool,'F_0') = s(bool,X308)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1500,f2782])).
fof(f9495,plain,(
  ( ! [X302,X304,X303,X305] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X302),s('type.integer.int',X303))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X305),s('type.integer.int',X302))),s('type.integer.int',X303))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X305),s('type.integer.int',sK3(X303,X302)))),s('type.integer.int',X302))) | s(bool,'F_0') = s(bool,X304)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9494,f73])).
fof(f9494,plain,(
  ( ! [X302,X304,X303,X305] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X305),s('type.integer.int',X302))),s('type.integer.int',X303))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X305),s('type.integer.int',sK3(X303,X302)))),s('type.integer.int',X302))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X302),s('type.integer.int',X303))) | s(bool,'F_0') = s(bool,X304)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8833,f1269])).
fof(f8833,plain,(
  ( ! [X302,X304,X303,X305] : (~p(s(bool,X304)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X305),s('type.integer.int',X302))),s('type.integer.int',X303))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X305),s('type.integer.int',sK3(X303,X302)))),s('type.integer.int',X302))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X302),s('type.integer.int',X303))) | s(bool,'F_0') = s(bool,X304)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1499,f2782])).
fof(f9493,plain,(
  ( ! [X300,X298,X301,X299] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int',X299))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298)),X301)),s('type.integer.int',X301))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X301),s('type.integer.int',X299)))) | s(bool,'F_0') = s(bool,X300)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9492,f73])).
fof(f9492,plain,(
  ( ! [X300,X298,X301,X299] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298)),X301)),s('type.integer.int',X301))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X301),s('type.integer.int',X299)))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int',X299))) | s(bool,'F_0') = s(bool,X300)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8832,f1269])).
fof(f8832,plain,(
  ( ! [X300,X298,X301,X299] : (~p(s(bool,X300)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X299,X298)),s('type.integer.int',X298)),X301)),s('type.integer.int',X301))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X301),s('type.integer.int',X299)))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int',X299))) | s(bool,'F_0') = s(bool,X300)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1485,f2782])).
fof(f9491,plain,(
  ( ! [X294,X296,X295,X297] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X294),s('type.integer.int',X295))) | s('type.integer.int',X297) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X297,'const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X295),s('type.integer.int',X297)))) | s(bool,'F_0') = s(bool,X296)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9490,f73])).
fof(f9490,plain,(
  ( ! [X294,X296,X295,X297] : (s('type.integer.int',X297) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X297,'const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X295),s('type.integer.int',X297)))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X294),s('type.integer.int',X295))) | s(bool,'F_0') = s(bool,X296)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8831,f1269])).
fof(f8831,plain,(
  ( ! [X294,X296,X295,X297] : (~p(s(bool,X296)) | s('type.integer.int',X297) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X297,'const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X295,X294)),s('type.integer.int',X294))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X295),s('type.integer.int',X297)))) | s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X294),s('type.integer.int',X295))) | s(bool,'F_0') = s(bool,X296)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1480,f2782])).
fof(f9475,plain,(
  ( ! [X204,X203,X201] : (s(bool,'T_0') != s(bool,X203) | s(bool,X201) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X201) != s(bool,X204) | s(bool,X203) = s(bool,X204)) ) | $spl222),
  inference(cnf_transformation,[],[f9475_D])).
fof(f9475_D,plain,(
  ( ! [X204,X203,X201] : (s(bool,'T_0') != s(bool,X203) | s(bool,X201) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X201) != s(bool,X204) | s(bool,X203) = s(bool,X204)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f9474,plain,(
  ( ! [X204,X202,X203,X201] : (s(bool,'T_0') != s(bool,X203) | s(bool,X203) = s(bool,X204) | s(bool,X201) != s(bool,X204) | s(bool,X201) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X202)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f8792,f1269])).
fof(f8792,plain,(
  ( ! [X204,X202,X203,X201] : (~p(s(bool,X202)) | s(bool,'T_0') != s(bool,X203) | s(bool,X203) = s(bool,X204) | s(bool,X201) != s(bool,X204) | s(bool,X201) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X202)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f3843,f2782])).
fof(f9473,plain,(
  ( ! [X200,X199,X197] : (s(bool,X199) != s(bool,X200) | s(bool,X197) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X197) != s(bool,X199) | s(bool,'T_0') = s(bool,X200)) ) | $spl220),
  inference(cnf_transformation,[],[f9473_D])).
fof(f9473_D,plain,(
  ( ! [X200,X199,X197] : (s(bool,X199) != s(bool,X200) | s(bool,X197) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X197) != s(bool,X199) | s(bool,'T_0') = s(bool,X200)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f9472,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,X199) != s(bool,X200) | s(bool,'T_0') = s(bool,X200) | s(bool,X197) != s(bool,X199) | s(bool,X197) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X198)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f8791,f1269])).
fof(f8791,plain,(
  ( ! [X198,X200,X199,X197] : (~p(s(bool,X198)) | s(bool,X199) != s(bool,X200) | s(bool,'T_0') = s(bool,X200) | s(bool,X197) != s(bool,X199) | s(bool,X197) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X198)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f2999,f2782])).
fof(f9471,plain,(
  ( ! [X189,X187,X190] : (s(bool,'F_0') = s(bool,X189) | s(bool,X187) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X187) != s(bool,X190) | s(bool,X189) = s(bool,X190)) ) | $spl218),
  inference(cnf_transformation,[],[f9471_D])).
fof(f9471_D,plain,(
  ( ! [X189,X187,X190] : (s(bool,'F_0') = s(bool,X189) | s(bool,X187) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X187) != s(bool,X190) | s(bool,X189) = s(bool,X190)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f9470,plain,(
  ( ! [X189,X187,X190,X188] : (s(bool,'F_0') = s(bool,X189) | s(bool,X189) = s(bool,X190) | s(bool,X187) != s(bool,X190) | s(bool,X187) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X188)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f8788,f1269])).
fof(f8788,plain,(
  ( ! [X189,X187,X190,X188] : (~p(s(bool,X188)) | s(bool,'F_0') = s(bool,X189) | s(bool,X189) = s(bool,X190) | s(bool,X187) != s(bool,X190) | s(bool,X187) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X188)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f2350,f2782])).
fof(f9466,plain,(
  ( ! [X140,X138,X136,X139,X137] : (s(bool,'F_0') = s(bool,X136) | s(bool,X136) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X138),s('type.integer.int',X139))) | s('type.integer.int',X139) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X140),s('type.integer.int',X138))) | s(bool,'F_0') = s(bool,X137)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9465,f73])).
fof(f9465,plain,(
  ( ! [X140,X138,X136,X139,X137] : (s(bool,X136) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X138),s('type.integer.int',X139))) | s('type.integer.int',X139) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X140),s('type.integer.int',X138))) | s(bool,X136) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X137)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8773,f1269])).
fof(f8773,plain,(
  ( ! [X140,X138,X136,X139,X137] : (~p(s(bool,X137)) | s(bool,X136) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X138),s('type.integer.int',X139))) | s('type.integer.int',X139) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X140),s('type.integer.int',X138))) | s(bool,X136) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X137)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f1311,f2782])).
fof(f9463,plain,(
  ( ! [X121,X122,X119] : (s(bool,'F_0') = s(bool,X121) | s(bool,X119) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X121) != s(bool,X122) | s(bool,X119) = s(bool,X122)) ) | $spl216),
  inference(cnf_transformation,[],[f9463_D])).
fof(f9463_D,plain,(
  ( ! [X121,X122,X119] : (s(bool,'F_0') = s(bool,X121) | s(bool,X119) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X121) != s(bool,X122) | s(bool,X119) = s(bool,X122)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f9462,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,X121) | s(bool,X119) = s(bool,X122) | s(bool,X121) != s(bool,X122) | s(bool,X119) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X120)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8767,f1269])).
fof(f8767,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,'F_0') = s(bool,X121) | s(bool,X119) = s(bool,X122) | s(bool,X121) != s(bool,X122) | s(bool,X119) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X120)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f653,f2782])).
fof(f9458,plain,(
  ( ! [X97,X95,X98] : (s(bool,X97) != s(bool,X98) | s(bool,X95) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') != s(bool,X97) | s(bool,X95) = s(bool,X98)) ) | $spl214),
  inference(cnf_transformation,[],[f9458_D])).
fof(f9458_D,plain,(
  ( ! [X97,X95,X98] : (s(bool,X97) != s(bool,X98) | s(bool,X95) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') != s(bool,X97) | s(bool,X95) = s(bool,X98)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f9457,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,X97) != s(bool,X98) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') != s(bool,X97) | s(bool,X95) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X96)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f8759,f1269])).
fof(f8759,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | s(bool,X97) != s(bool,X98) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') != s(bool,X97) | s(bool,X95) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X96)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f599,f2782])).
fof(f9452,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,'T_0') != s(bool,X83) | s(bool,X81) = s(bool,X83) | s(bool,X81) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f8754,f312])).
fof(f8754,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,'T_0') != s(bool,X83) | s(bool,X81) = s(bool,X83) | s(bool,X81) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X82)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f544,f2782])).
fof(f9451,plain,(
  ( ! [X78,X77,X75] : (s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X77) = s(bool,X78) | s(bool,X75) = s(bool,X78)) ) | $spl212),
  inference(cnf_transformation,[],[f9451_D])).
fof(f9451_D,plain,(
  ( ! [X78,X77,X75] : (s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X77) = s(bool,X78) | s(bool,X75) = s(bool,X78)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f9450,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,X78) | s(bool,X77) = s(bool,X78) | s(bool,X75) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X76)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8752,f1269])).
fof(f8752,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,X78) | s(bool,X77) = s(bool,X78) | s(bool,X75) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X76)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f521,f2782])).
fof(f9445,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,'F_0') = s(bool,X57) | ~p(s(bool,X58)) | s(bool,'F_0') != s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,X59) = s(bool,X60)) ) | ($spl6 | $spl22)),
  inference(forward_demodulation,[],[f9444,f73])).
fof(f9444,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | s(bool,'F_0') != s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,X59) = s(bool,X60) | s(bool,X57) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8746,f309])).
fof(f8746,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | s(bool,'F_0') != s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,X59) = s(bool,X60) | s(bool,X57) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X58)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f461,f2782])).
fof(f9440,plain,(
  ( ! [X37,X35,X36] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') != s(bool,X37) | s(bool,X35) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X36)) ) | ($spl6 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f8738,f1269])).
fof(f8738,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X36)) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') != s(bool,X37) | s(bool,X35) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X36)) ) | ($spl6 | $spl16 | $spl22)),
  inference(superposition,[],[f206,f2782])).
fof(f9216,plain,(
  ( ! [X54,X52,X53] : (s(bool,'F_0') = s(bool,X53) | s(bool,X54) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X52) = s(bool,X54) | s(bool,X52) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | ($spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f8541])).
fof(f8541,plain,(
  ( ! [X54,X52,X53] : (s(bool,'F_0') = s(bool,X53) | s(bool,X54) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X52) = s(bool,X54) | s(bool,X52) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X53)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f2782,f2782])).
fof(f9221,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X56) | s(bool,X57) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | ($spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f8485])).
fof(f8485,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X56) | s(bool,X57) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X56)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f2782,f2782])).
fof(f9231,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,X54) | s(bool,X53) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X55) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X53) = s(bool,X55)) ) | ($spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f8446])).
fof(f8446,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,X54) | s(bool,X53) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X54) | s(bool,X55) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X53) = s(bool,X55)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f2782,f2782])).
fof(f5024,plain,(
  ( ! [X90,X91] : (s(bool,'F_0') = s(bool,X90) | s(bool,X90) = s(bool,X91) | s(bool,X91) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | $spl82),
  inference(cnf_transformation,[],[f5024_D])).
fof(f5024_D,plain,(
  ( ! [X90,X91] : (s(bool,'F_0') = s(bool,X90) | s(bool,X90) = s(bool,X91) | s(bool,X91) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8377,plain,(
  ( ! [X35,X33,X34] : (s(bool,X33) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8332,f44])).
fof(f8332,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,X34) | s(bool,X33) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f620,f1623])).
fof(f8376,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f8331,f44])).
fof(f8331,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X30) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f620,f1623])).
fof(f7632,plain,(
  ( ! [X304,X306,X303,X307] : (s(bool,'T_0') != s(bool,X306) | s(bool,X307) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X303),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X304),s('type.integer.int',X303))))) | s(bool,X306) = s(bool,X307)) ) | $spl206),
  inference(cnf_transformation,[],[f7632_D])).
fof(f7632_D,plain,(
  ( ! [X304,X306,X303,X307] : (s(bool,'T_0') != s(bool,X306) | s(bool,X307) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X303),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X304),s('type.integer.int',X303))))) | s(bool,X306) = s(bool,X307)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f7631,plain,(
  ( ! [X304,X306,X303,X305,X307] : (s(bool,'T_0') != s(bool,X306) | s(bool,X306) = s(bool,X307) | s(bool,X307) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X303),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X304),s('type.integer.int',X303))))) | s(bool,'F_0') = s(bool,X305)) ) | $spl16),
  inference(subsumption_resolution,[],[f7431,f1269])).
fof(f7431,plain,(
  ( ! [X304,X306,X303,X305,X307] : (~p(s(bool,X305)) | s(bool,'T_0') != s(bool,X306) | s(bool,X306) = s(bool,X307) | s(bool,X307) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X303),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X304),s('type.integer.int',X303))))) | s(bool,'F_0') = s(bool,X305)) ) | $spl16),
  inference(superposition,[],[f3843,f1625])).
fof(f7630,plain,(
  ( ! [X302,X298,X301,X299] : (s(bool,X301) != s(bool,X302) | s(bool,X301) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X299),s('type.integer.int',X298))))) | s(bool,'T_0') = s(bool,X302)) ) | $spl204),
  inference(cnf_transformation,[],[f7630_D])).
fof(f7630_D,plain,(
  ( ! [X302,X298,X301,X299] : (s(bool,X301) != s(bool,X302) | s(bool,X301) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X299),s('type.integer.int',X298))))) | s(bool,'T_0') = s(bool,X302)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7629,plain,(
  ( ! [X300,X302,X298,X301,X299] : (s(bool,X301) != s(bool,X302) | s(bool,'T_0') = s(bool,X302) | s(bool,X301) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X299),s('type.integer.int',X298))))) | s(bool,'F_0') = s(bool,X300)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f7430,f1269])).
fof(f7430,plain,(
  ( ! [X300,X302,X298,X301,X299] : (~p(s(bool,X300)) | s(bool,X301) != s(bool,X302) | s(bool,'T_0') = s(bool,X302) | s(bool,X301) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X298),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X299),s('type.integer.int',X298))))) | s(bool,'F_0') = s(bool,X300)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f2999,f1625])).
fof(f7628,plain,(
  ( ! [X294,X296,X293,X297] : (s(bool,'F_0') = s(bool,X296) | s(bool,X297) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X293),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X294),s('type.integer.int',X293))))) | s(bool,X296) = s(bool,X297)) ) | $spl202),
  inference(cnf_transformation,[],[f7628_D])).
fof(f7628_D,plain,(
  ( ! [X294,X296,X293,X297] : (s(bool,'F_0') = s(bool,X296) | s(bool,X297) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X293),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X294),s('type.integer.int',X293))))) | s(bool,X296) = s(bool,X297)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f7627,plain,(
  ( ! [X294,X296,X293,X295,X297] : (s(bool,'F_0') = s(bool,X296) | s(bool,X296) = s(bool,X297) | s(bool,X297) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X293),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X294),s('type.integer.int',X293))))) | s(bool,'F_0') = s(bool,X295)) ) | $spl16),
  inference(subsumption_resolution,[],[f7429,f1269])).
fof(f7429,plain,(
  ( ! [X294,X296,X293,X295,X297] : (~p(s(bool,X295)) | s(bool,'F_0') = s(bool,X296) | s(bool,X296) = s(bool,X297) | s(bool,X297) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X293),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X294),s('type.integer.int',X293))))) | s(bool,'F_0') = s(bool,X295)) ) | $spl16),
  inference(superposition,[],[f2350,f1625])).
fof(f7623,plain,(
  ( ! [X222,X225,X223,X224] : (s(bool,X225) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X222),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X223),s('type.integer.int',X222))))) | s(bool,'T_0') = s(bool,X225) | s(bool,'F_0') = s(bool,X224)) ) | $spl16),
  inference(subsumption_resolution,[],[f7412,f1269])).
fof(f7412,plain,(
  ( ! [X222,X225,X223,X224] : (~p(s(bool,X224)) | s(bool,X225) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X222),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X223),s('type.integer.int',X222))))) | s(bool,'T_0') = s(bool,X225) | s(bool,'F_0') = s(bool,X224)) ) | $spl16),
  inference(superposition,[],[f820,f1625])).
fof(f7621,plain,(
  ( ! [X218,X216,X215,X217,X214] : (s(bool,'F_0') = s(bool,X217) | s(bool,X218) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X214),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X215),s('type.integer.int',X214))))) | s(bool,X217) != s(bool,X218) | s(bool,'F_0') = s(bool,X216)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f7410,f1269])).
fof(f7410,plain,(
  ( ! [X218,X216,X215,X217,X214] : (~p(s(bool,X216)) | s(bool,'F_0') = s(bool,X217) | s(bool,X218) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X214),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X215),s('type.integer.int',X214))))) | s(bool,X217) != s(bool,X218) | s(bool,'F_0') = s(bool,X216)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f653,f1625])).
fof(f7619,plain,(
  ( ! [X191,X194,X190,X193] : (s(bool,X193) != s(bool,X194) | s(bool,'T_0') != s(bool,X193) | s(bool,X194) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X190),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X191),s('type.integer.int',X190)))))) ) | $spl198),
  inference(cnf_transformation,[],[f7619_D])).
fof(f7619_D,plain,(
  ( ! [X191,X194,X190,X193] : (s(bool,X193) != s(bool,X194) | s(bool,'T_0') != s(bool,X193) | s(bool,X194) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X190),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X191),s('type.integer.int',X190)))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f7618,plain,(
  ( ! [X191,X194,X192,X190,X193] : (s(bool,X193) != s(bool,X194) | s(bool,X194) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X190),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X191),s('type.integer.int',X190))))) | s(bool,'T_0') != s(bool,X193) | s(bool,'F_0') = s(bool,X192)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f7404,f1269])).
fof(f7404,plain,(
  ( ! [X191,X194,X192,X190,X193] : (~p(s(bool,X192)) | s(bool,X193) != s(bool,X194) | s(bool,X194) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X190),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X191),s('type.integer.int',X190))))) | s(bool,'T_0') != s(bool,X193) | s(bool,'F_0') = s(bool,X192)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f599,f1625])).
fof(f7605,plain,(
  ( ! [X112,X115,X113] : (s(bool,X115) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X112),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X113),s('type.integer.int',X112))))) | s(bool,'T_0') != s(bool,X115)) ) | $spl194),
  inference(cnf_transformation,[],[f7605_D])).
fof(f7605_D,plain,(
  ( ! [X112,X115,X113] : (s(bool,X115) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X112),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X113),s('type.integer.int',X112))))) | s(bool,'T_0') != s(bool,X115)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7615,plain,(
  ( ! [X177,X175,X176,X174] : (s(bool,'T_0') != s(bool,X176) | s(bool,'T_0') != s(bool,X177) | s(bool,X177) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X174),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X175),s('type.integer.int',X174)))))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f7400,f312])).
fof(f7400,plain,(
  ( ! [X177,X175,X176,X174] : (s(bool,'T_0') != s(bool,X176) | s(bool,'T_0') != s(bool,X177) | s(bool,X177) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X174),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X175),s('type.integer.int',X174))))) | s(bool,'F_0') = s(bool,X176)) ) | $spl16),
  inference(superposition,[],[f544,f1625])).
fof(f7614,plain,(
  ( ! [X167,X169,X166,X170] : (s(bool,'T_0') = s(bool,X169) | s(bool,X169) = s(bool,X170) | s(bool,X170) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X166),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X167),s('type.integer.int',X166)))))) ) | $spl196),
  inference(cnf_transformation,[],[f7614_D])).
fof(f7614_D,plain,(
  ( ! [X167,X169,X166,X170] : (s(bool,'T_0') = s(bool,X169) | s(bool,X169) = s(bool,X170) | s(bool,X170) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X166),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X167),s('type.integer.int',X166)))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f7613,plain,(
  ( ! [X167,X169,X166,X170,X168] : (s(bool,'T_0') = s(bool,X169) | s(bool,X170) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X166),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X167),s('type.integer.int',X166))))) | s(bool,X169) = s(bool,X170) | s(bool,'F_0') = s(bool,X168)) )),
  inference(subsumption_resolution,[],[f7398,f1269])).
fof(f7398,plain,(
  ( ! [X167,X169,X166,X170,X168] : (~p(s(bool,X168)) | s(bool,'T_0') = s(bool,X169) | s(bool,X170) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X166),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X167),s('type.integer.int',X166))))) | s(bool,X169) = s(bool,X170) | s(bool,'F_0') = s(bool,X168)) )),
  inference(superposition,[],[f521,f1625])).
fof(f7608,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'T_0') != s(bool,X128) | s(bool,X129) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X126),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X127),s('type.integer.int',X126))))) | s(bool,'T_0') = s(bool,X129)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f7388,f312])).
fof(f7388,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'T_0') != s(bool,X128) | s(bool,X129) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X126),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X127),s('type.integer.int',X126))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'F_0') = s(bool,X128)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f368,f1625])).
fof(f7604,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,X115) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X112),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X113),s('type.integer.int',X112))))) | s(bool,'T_0') != s(bool,X115) | s(bool,'F_0') = s(bool,X114)) ) | $spl16),
  inference(subsumption_resolution,[],[f7384,f1269])).
fof(f7384,plain,(
  ( ! [X114,X112,X115,X113] : (~p(s(bool,X114)) | s(bool,X115) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X112),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X113),s('type.integer.int',X112))))) | s(bool,'T_0') != s(bool,X115) | s(bool,'F_0') = s(bool,X114)) ) | $spl16),
  inference(superposition,[],[f206,f1625])).
fof(f7595,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61)))),'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X59),s('type.integer.int',X61))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X60),s('type.integer.int',X61)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61)))))) ) | $spl192),
  inference(cnf_transformation,[],[f7595_D])).
fof(f7595_D,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61)))),'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X59),s('type.integer.int',X61))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X60),s('type.integer.int',X61)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61)))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f7594,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61)))),'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X59),s('type.integer.int',X61))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X60),s('type.integer.int',X61)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61))))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f7370,f1269])).
fof(f7370,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,X63)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61)))),'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X59),s('type.integer.int',X61))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X60),s('type.integer.int',X61)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X59),s('type.integer.int',X60))),s('type.integer.int',X61))))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f1524,f1625])).
fof(f7593,plain,(
  ( ! [X57,X54,X56,X55] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X57),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X57),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56)))),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X54),s('type.integer.int',X56))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X55),s('type.integer.int',X56)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56)))))) ) | $spl190),
  inference(cnf_transformation,[],[f7593_D])).
fof(f7593_D,plain,(
  ( ! [X57,X54,X56,X55] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X57),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X57),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56)))),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X54),s('type.integer.int',X56))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X55),s('type.integer.int',X56)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56)))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7592,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X57),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X57),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56)))),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X54),s('type.integer.int',X56))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X55),s('type.integer.int',X56)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56))))) | s(bool,'F_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f7369,f1269])).
fof(f7369,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X58)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X57),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X57),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56)))),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X54),s('type.integer.int',X56))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X55),s('type.integer.int',X56)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X54),s('type.integer.int',X55))),s('type.integer.int',X56))))) | s(bool,'F_0') = s(bool,X58)) )),
  inference(superposition,[],[f1522,f1625])).
fof(f7591,plain,(
  ( ! [X52,X50,X51] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51)))),'const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51)))),s('type.integer.int',X50))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int',X50))))) ) | $spl188),
  inference(cnf_transformation,[],[f7591_D])).
fof(f7591_D,plain,(
  ( ! [X52,X50,X51] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51)))),'const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51)))),s('type.integer.int',X50))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int',X50))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7590,plain,(
  ( ! [X52,X50,X53,X51] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51)))),'const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51)))),s('type.integer.int',X50))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int',X50)))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f7368,f1269])).
fof(f7368,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X53)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X52),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51)))),'const.integer.int_mul_2'(s('type.integer.int',sK3(X50,X51)),s('type.integer.int',X51)))),s('type.integer.int',X50))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int',X50)))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f1475,f1625])).
fof(f7589,plain,(
  ( ! [X47,X46,X49] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X49),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))))),s('type.integer.int',X49)),X46)),s('type.integer.int',X46)))) ) | $spl186),
  inference(cnf_transformation,[],[f7589_D])).
fof(f7589_D,plain,(
  ( ! [X47,X46,X49] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X49),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))))),s('type.integer.int',X49)),X46)),s('type.integer.int',X46)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f7588,plain,(
  ( ! [X47,X48,X46,X49] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X49),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))))),s('type.integer.int',X49)),X46)),s('type.integer.int',X46))) | s(bool,'F_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f7367,f1269])).
fof(f7367,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X48)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X49),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X47),s('type.integer.int',X46))))),s('type.integer.int',X49)),X46)),s('type.integer.int',X46))) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f1531,f1625])).
fof(f7587,plain,(
  ( ! [X45,X43,X42] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42)),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X45),s('type.integer.int',X42))),s('type.integer.int',X45)))),s('type.integer.int',X42)))) ) | $spl184),
  inference(cnf_transformation,[],[f7587_D])).
fof(f7587_D,plain,(
  ( ! [X45,X43,X42] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42)),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X45),s('type.integer.int',X42))),s('type.integer.int',X45)))),s('type.integer.int',X42)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7586,plain,(
  ( ! [X45,X43,X44,X42] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42)),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X45),s('type.integer.int',X42))),s('type.integer.int',X45)))),s('type.integer.int',X42))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f7366,f1269])).
fof(f7366,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X44)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',X42)),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X45),s('type.integer.int',X42))),s('type.integer.int',X45)))),s('type.integer.int',X42))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(superposition,[],[f1523,f1625])).
fof(f7585,plain,(
  ( ! [X39,X41,X38] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X39),s('type.integer.int',X41))),s('type.integer.int',X38))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int',X38)),X38)),s('type.integer.int',X41))),s('type.integer.int',X38)))) ) | $spl182),
  inference(cnf_transformation,[],[f7585_D])).
fof(f7585_D,plain,(
  ( ! [X39,X41,X38] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X39),s('type.integer.int',X41))),s('type.integer.int',X38))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int',X38)),X38)),s('type.integer.int',X41))),s('type.integer.int',X38)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f7584,plain,(
  ( ! [X39,X41,X38,X40] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X39),s('type.integer.int',X41))),s('type.integer.int',X38))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int',X38)),X38)),s('type.integer.int',X41))),s('type.integer.int',X38))) | s(bool,'F_0') = s(bool,X40)) )),
  inference(forward_demodulation,[],[f7583,f33])).
fof(f7583,plain,(
  ( ! [X39,X41,X38,X40] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int',X38))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X41),s('type.integer.int',X38))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int',X38)),X38)),s('type.integer.int',X41))),s('type.integer.int',X38))) | s(bool,'F_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7365,f1269])).
fof(f7365,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X40)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int',X38))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X41),s('type.integer.int',X38))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int',X38)),X38)),s('type.integer.int',X41))),s('type.integer.int',X38))) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f1502,f1625])).
fof(f7582,plain,(
  ( ! [X37,X35,X34] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X37),s('type.integer.int',X35))),s('type.integer.int',X34))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X37),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X35),s('type.integer.int',X34)),X34)))),s('type.integer.int',X34)))) ) | $spl180),
  inference(cnf_transformation,[],[f7582_D])).
fof(f7582_D,plain,(
  ( ! [X37,X35,X34] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X37),s('type.integer.int',X35))),s('type.integer.int',X34))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X37),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X35),s('type.integer.int',X34)),X34)))),s('type.integer.int',X34)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f7581,plain,(
  ( ! [X37,X35,X36,X34] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X37),s('type.integer.int',X35))),s('type.integer.int',X34))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X37),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X35),s('type.integer.int',X34)),X34)))),s('type.integer.int',X34))) | s(bool,'F_0') = s(bool,X36)) )),
  inference(forward_demodulation,[],[f7580,f33])).
fof(f7580,plain,(
  ( ! [X37,X35,X36,X34] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int',X34))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X35),s('type.integer.int',X34))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X37),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X35),s('type.integer.int',X34)),X34)))),s('type.integer.int',X34))) | s(bool,'F_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f7364,f1269])).
fof(f7364,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X36)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int',X34))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X35),s('type.integer.int',X34))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X37),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X35),s('type.integer.int',X34)),X34)))),s('type.integer.int',X34))) | s(bool,'F_0') = s(bool,X36)) )),
  inference(superposition,[],[f1501,f1625])).
fof(f7579,plain,(
  ( ! [X30,X33,X31] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X31),s('type.integer.int',X33))),s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30)),X30)),s('type.integer.int',X33))),s('type.integer.int',X30)))) ) | $spl178),
  inference(cnf_transformation,[],[f7579_D])).
fof(f7579_D,plain,(
  ( ! [X30,X33,X31] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X31),s('type.integer.int',X33))),s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30)),X30)),s('type.integer.int',X33))),s('type.integer.int',X30)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f7578,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X31),s('type.integer.int',X33))),s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30)),X30)),s('type.integer.int',X33))),s('type.integer.int',X30))) | s(bool,'F_0') = s(bool,X32)) )),
  inference(forward_demodulation,[],[f7577,f34])).
fof(f7577,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X30))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30)),X30)),s('type.integer.int',X33))),s('type.integer.int',X30))) | s(bool,'F_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f7363,f1269])).
fof(f7363,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X32)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X30))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30)),X30)),s('type.integer.int',X33))),s('type.integer.int',X30))) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f1500,f1625])).
fof(f7576,plain,(
  ( ! [X26,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X29),s('type.integer.int',X27))),s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X29),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26)),X26)))),s('type.integer.int',X26)))) ) | $spl176),
  inference(cnf_transformation,[],[f7576_D])).
fof(f7576_D,plain,(
  ( ! [X26,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X29),s('type.integer.int',X27))),s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X29),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26)),X26)))),s('type.integer.int',X26)))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f7575,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X29),s('type.integer.int',X27))),s('type.integer.int',X26))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X29),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26)),X26)))),s('type.integer.int',X26))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(forward_demodulation,[],[f7574,f34])).
fof(f7574,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X29),s('type.integer.int',X26))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X29),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26)),X26)))),s('type.integer.int',X26))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7362,f1269])).
fof(f7362,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,X28)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X29),s('type.integer.int',X26))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X29),s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26)),X26)))),s('type.integer.int',X26))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f1499,f1625])).
fof(f7573,plain,(
  ( ! [X23,X25,X22] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)),X22)),s('type.integer.int',X22))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)),X22)),s('type.integer.int',X22)),X25)),s('type.integer.int',X25))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22))))))) ) | $spl174),
  inference(cnf_transformation,[],[f7573_D])).
fof(f7573_D,plain,(
  ( ! [X23,X25,X22] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)),X22)),s('type.integer.int',X22))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)),X22)),s('type.integer.int',X22)),X25)),s('type.integer.int',X25))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22))))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7572,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)),X22)),s('type.integer.int',X22))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)),X22)),s('type.integer.int',X22)),X25)),s('type.integer.int',X25))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)))))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f7361,f1269])).
fof(f7361,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,X24)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)),X22)),s('type.integer.int',X22))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)),X22)),s('type.integer.int',X22)),X25)),s('type.integer.int',X25))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X22)))))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f1485,f1625])).
fof(f7571,plain,(
  ( ! [X21,X19,X18] : (s('type.integer.int',X21) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X21,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18)),X18)),s('type.integer.int',X18)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18)),X18)),s('type.integer.int',X18))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18))),s('type.integer.int',X21))))) ) | $spl172),
  inference(cnf_transformation,[],[f7571_D])).
fof(f7571_D,plain,(
  ( ! [X21,X19,X18] : (s('type.integer.int',X21) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X21,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18)),X18)),s('type.integer.int',X18)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18)),X18)),s('type.integer.int',X18))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18))),s('type.integer.int',X21))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7570,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.integer.int',X21) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X21,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18)),X18)),s('type.integer.int',X18)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18)),X18)),s('type.integer.int',X18))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18))),s('type.integer.int',X21)))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f7360,f1269])).
fof(f7360,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,X20)) | s('type.integer.int',X21) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X21,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18)),X18)),s('type.integer.int',X18)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18)),X18)),s('type.integer.int',X18))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X18))),s('type.integer.int',X21)))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f1480,f1625])).
fof(f7569,plain,(
  ( ! [X12,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11)),X11)),s('type.integer.int',X11)))) ) | $spl170),
  inference(cnf_transformation,[],[f7569_D])).
fof(f7569_D,plain,(
  ( ! [X12,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11)),X11)),s('type.integer.int',X11)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7568,plain,(
  ( ! [X12,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11)),X11)),s('type.integer.int',X11))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f7358,f1269])).
fof(f7358,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X13)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11)),X11)),s('type.integer.int',X11))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f35,f1625])).
fof(f7567,plain,(
  ( ! [X8,X7,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X9))),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X7)))),X7)),s('type.integer.int',X7)))) ) | $spl168),
  inference(cnf_transformation,[],[f7567_D])).
fof(f7567_D,plain,(
  ( ! [X8,X7,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X9))),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X7)))),X7)),s('type.integer.int',X7)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f7566,plain,(
  ( ! [X10,X8,X7,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X9))),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X7)))),X7)),s('type.integer.int',X7))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f7357,f1269])).
fof(f7357,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,X10)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X9))),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X7)))),X7)),s('type.integer.int',X7))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f1532,f1625])).
fof(f7565,plain,(
  ( ! [X4,X5,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X3))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int',X3)))),X3)),s('type.integer.int',X3)))) ) | $spl166),
  inference(cnf_transformation,[],[f7565_D])).
fof(f7565_D,plain,(
  ( ! [X4,X5,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X3))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int',X3)))),X3)),s('type.integer.int',X3)))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f7564,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X3))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int',X3)))),X3)),s('type.integer.int',X3))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7356,f1269])).
fof(f7356,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,X6)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X4),s('type.integer.int',X5))),s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X3))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int',X3)))),X3)),s('type.integer.int',X3))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f1530,f1625])).
fof(f7445,plain,(
  ( ! [X146,X149,X147,X148] : (s(bool,X149) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X147),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X148),s('type.integer.int',X147))))) | s(bool,'F_0') = s(bool,X146) | s(bool,'T_0') = s(bool,X149)) ) | ($spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f7312])).
fof(f7312,plain,(
  ( ! [X146,X149,X147,X148] : (s(bool,X149) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X147),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X148),s('type.integer.int',X147))))) | s(bool,'F_0') = s(bool,X146) | s(bool,'T_0') = s(bool,X149) | s(bool,'F_0') = s(bool,X146)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f652,f1625])).
fof(f652,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl22)),
  inference(resolution,[],[f641,f106])).
fof(f6984,plain,(
  ( ! [X140,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X140) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,X139) = s(bool,X140)) ) | $spl162),
  inference(cnf_transformation,[],[f6984_D])).
fof(f6984_D,plain,(
  ( ! [X140,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X140) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,X139) = s(bool,X140)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f6983,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X139) = s(bool,X140) | s(bool,X140) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X138)) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f6805,f1269])).
fof(f6805,plain,(
  ( ! [X140,X138,X139] : (~p(s(bool,X138)) | s(bool,'T_0') != s(bool,X139) | s(bool,X139) = s(bool,X140) | s(bool,X140) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X138)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f3843,f1645])).
fof(f6982,plain,(
  ( ! [X136,X137] : (s(bool,X136) != s(bool,X137) | s(bool,X136) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X137)) ) | $spl160),
  inference(cnf_transformation,[],[f6982_D])).
fof(f6982_D,plain,(
  ( ! [X136,X137] : (s(bool,X136) != s(bool,X137) | s(bool,X136) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X137)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6981,plain,(
  ( ! [X136,X135,X137] : (s(bool,X136) != s(bool,X137) | s(bool,'T_0') = s(bool,X137) | s(bool,X136) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X135)) ) | ($spl4 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f6804,f1269])).
fof(f6804,plain,(
  ( ! [X136,X135,X137] : (~p(s(bool,X135)) | s(bool,X136) != s(bool,X137) | s(bool,'T_0') = s(bool,X137) | s(bool,X136) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X135)) ) | ($spl4 | $spl16 | $spl22)),
  inference(superposition,[],[f2999,f1645])).
fof(f6980,plain,(
  ( ! [X134,X133] : (s(bool,'F_0') = s(bool,X133) | s(bool,X134) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,X133) = s(bool,X134)) ) | $spl158),
  inference(cnf_transformation,[],[f6980_D])).
fof(f6980_D,plain,(
  ( ! [X134,X133] : (s(bool,'F_0') = s(bool,X133) | s(bool,X134) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,X133) = s(bool,X134)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f6979,plain,(
  ( ! [X134,X132,X133] : (s(bool,'F_0') = s(bool,X133) | s(bool,X133) = s(bool,X134) | s(bool,X134) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X132)) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f6803,f1269])).
fof(f6803,plain,(
  ( ! [X134,X132,X133] : (~p(s(bool,X132)) | s(bool,'F_0') = s(bool,X133) | s(bool,X133) = s(bool,X134) | s(bool,X134) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X132)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f2350,f1645])).
fof(f6882,plain,(
  ( ! [X71] : (s(bool,X71) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X71)) ) | $spl132),
  inference(cnf_transformation,[],[f6882_D])).
fof(f6882_D,plain,(
  ( ! [X71] : (s(bool,X71) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X71)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f6975,plain,(
  ( ! [X97,X96] : (s(bool,X97) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X97) | s(bool,'F_0') = s(bool,X96)) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f6787,f1269])).
fof(f6787,plain,(
  ( ! [X97,X96] : (~p(s(bool,X96)) | s(bool,X97) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X97) | s(bool,'F_0') = s(bool,X96)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f820,f1645])).
fof(f6974,plain,(
  ( ! [X94,X93] : (s(bool,'F_0') = s(bool,X93) | s(bool,X93) != s(bool,X94) | s(bool,X94) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl156),
  inference(cnf_transformation,[],[f6974_D])).
fof(f6974_D,plain,(
  ( ! [X94,X93] : (s(bool,'F_0') = s(bool,X93) | s(bool,X93) != s(bool,X94) | s(bool,X94) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f6973,plain,(
  ( ! [X94,X92,X93] : (s(bool,'F_0') = s(bool,X93) | s(bool,X94) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,X93) != s(bool,X94) | s(bool,'F_0') = s(bool,X92)) ) | ($spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f6785,f1269])).
fof(f6785,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,X92)) | s(bool,'F_0') = s(bool,X93) | s(bool,X94) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,X93) != s(bool,X94) | s(bool,'F_0') = s(bool,X92)) ) | ($spl4 | $spl6 | $spl22)),
  inference(superposition,[],[f653,f1645])).
fof(f6972,plain,(
  ( ! [X81,X82] : (s(bool,X81) != s(bool,X82) | s(bool,'T_0') != s(bool,X81) | s(bool,X82) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl154),
  inference(cnf_transformation,[],[f6972_D])).
fof(f6972_D,plain,(
  ( ! [X81,X82] : (s(bool,X81) != s(bool,X82) | s(bool,'T_0') != s(bool,X81) | s(bool,X82) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f6971,plain,(
  ( ! [X80,X81,X82] : (s(bool,X81) != s(bool,X82) | s(bool,X82) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') != s(bool,X81) | s(bool,'F_0') = s(bool,X80)) ) | ($spl4 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f6779,f1269])).
fof(f6779,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,X80)) | s(bool,X81) != s(bool,X82) | s(bool,X82) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') != s(bool,X81) | s(bool,'F_0') = s(bool,X80)) ) | ($spl4 | $spl16 | $spl22)),
  inference(superposition,[],[f599,f1645])).
fof(f6958,plain,(
  ( ! [X43] : (s(bool,X43) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') != s(bool,X43)) ) | $spl150),
  inference(cnf_transformation,[],[f6958_D])).
fof(f6958_D,plain,(
  ( ! [X43] : (s(bool,X43) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') != s(bool,X43)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f6968,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') != s(bool,X72) | s(bool,'T_0') != s(bool,X73) | s(bool,X73) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | ($spl4 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f6775,f312])).
fof(f6775,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') != s(bool,X72) | s(bool,'T_0') != s(bool,X73) | s(bool,X73) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X72)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f544,f1645])).
fof(f6967,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,X70) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl152),
  inference(cnf_transformation,[],[f6967_D])).
fof(f6967_D,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,X70) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6966,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X68)) ) | $spl4),
  inference(subsumption_resolution,[],[f6773,f1269])).
fof(f6773,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,X68)) | s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X68)) ) | $spl4),
  inference(superposition,[],[f521,f1645])).
fof(f6961,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X48) | s(bool,X49) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl4 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f6763,f312])).
fof(f6763,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X48) | s(bool,X49) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) ) | ($spl4 | $spl16 | $spl22)),
  inference(superposition,[],[f368,f1645])).
fof(f6957,plain,(
  ( ! [X43,X42] : (s(bool,X43) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') != s(bool,X43) | s(bool,'F_0') = s(bool,X42)) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f6759,f1269])).
fof(f6759,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | s(bool,X43) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'T_0') != s(bool,X43) | s(bool,'F_0') = s(bool,X42)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f206,f1645])).
fof(f6948,plain,(
  ( ! [X19] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int',X19)),sK0)),s('type.integer.int',sK0)))) ) | $spl148),
  inference(cnf_transformation,[],[f6948_D])).
fof(f6948_D,plain,(
  ( ! [X19] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int',X19)),sK0)),s('type.integer.int',sK0)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f6947,plain,(
  ( ! [X19,X18] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int',X19)),sK0)),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X18)) ) | $spl4),
  inference(subsumption_resolution,[],[f6745,f1269])).
fof(f6745,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int',X19)),sK0)),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X18)) ) | $spl4),
  inference(superposition,[],[f1531,f1645])).
fof(f6946,plain,(
  ( ! [X17] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK0))),s('type.integer.int',X17)))),s('type.integer.int',sK0)))) ) | $spl146),
  inference(cnf_transformation,[],[f6946_D])).
fof(f6946_D,plain,(
  ( ! [X17] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK0))),s('type.integer.int',X17)))),s('type.integer.int',sK0)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f6945,plain,(
  ( ! [X17,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK0))),s('type.integer.int',X17)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X16)) ) | $spl4),
  inference(subsumption_resolution,[],[f6744,f1269])).
fof(f6744,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK0))),s('type.integer.int',X17)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X16)) ) | $spl4),
  inference(superposition,[],[f1523,f1645])).
fof(f6944,plain,(
  ( ! [X15] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',X15))),s('type.integer.int',sK0)))) ) | $spl144),
  inference(cnf_transformation,[],[f6944_D])).
fof(f6944_D,plain,(
  ( ! [X15] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',X15))),s('type.integer.int',sK0)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f6943,plain,(
  ( ! [X14,X15] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',X15))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X14)) ) | $spl4),
  inference(subsumption_resolution,[],[f6743,f1269])).
fof(f6743,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',X15))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X14)) ) | $spl4),
  inference(superposition,[],[f1502,f1645])).
fof(f6942,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)))),s('type.integer.int',sK0)))) ) | $spl142),
  inference(cnf_transformation,[],[f6942_D])).
fof(f6942_D,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)))),s('type.integer.int',sK0)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f6941,plain,(
  ( ! [X12,X13] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X12)) ) | $spl4),
  inference(subsumption_resolution,[],[f6742,f1269])).
fof(f6742,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X13),s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X12)) ) | $spl4),
  inference(superposition,[],[f1501,f1645])).
fof(f6940,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',X11))),s('type.integer.int',sK0)))) ) | $spl140),
  inference(cnf_transformation,[],[f6940_D])).
fof(f6940_D,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',X11))),s('type.integer.int',sK0)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f6939,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',X11))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f6741,f1269])).
fof(f6741,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',X11))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X10)) ) | $spl4),
  inference(superposition,[],[f1500,f1645])).
fof(f6938,plain,(
  ( ! [X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)))),s('type.integer.int',sK0)))) ) | $spl138),
  inference(cnf_transformation,[],[f6938_D])).
fof(f6938_D,plain,(
  ( ! [X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)))),s('type.integer.int',sK0)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f6937,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X8)) ) | $spl4),
  inference(subsumption_resolution,[],[f6740,f1269])).
fof(f6740,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X8)) ) | $spl4),
  inference(superposition,[],[f1499,f1645])).
fof(f6936,plain,(
  ( ! [X7] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl136),
  inference(cnf_transformation,[],[f6936_D])).
fof(f6936_D,plain,(
  ( ! [X7] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f6935,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),X7)),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s(bool,'F_0') = s(bool,X6)) ) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f6934,f391])).
fof(f6933,plain,(
  ( ! [X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X5))))) ) | $spl134),
  inference(cnf_transformation,[],[f6933_D])).
fof(f6933_D,plain,(
  ( ! [X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X5))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f6932,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,'const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X5)))) | s(bool,'F_0') = s(bool,X4)) ) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f6931,f391])).
fof(f6818,plain,(
  ( ! [X70,X71] : (s(bool,X71) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X71)) ) | ($spl4 | $spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f6694])).
fof(f6694,plain,(
  ( ! [X70,X71] : (s(bool,X71) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s(bool,'F_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,X70)) ) | ($spl4 | $spl6 | $spl22)),
  inference(superposition,[],[f652,f1645])).
fof(f6447,plain,(
  ( ! [X90,X91] : (s(bool,'F_0') = s(bool,X90) | s(bool,X91) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X90) != s(bool,X91)) ) | $spl130),
  inference(cnf_transformation,[],[f6447_D])).
fof(f6447_D,plain,(
  ( ! [X90,X91] : (s(bool,'F_0') = s(bool,X90) | s(bool,X91) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X90) != s(bool,X91)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f6446,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') = s(bool,X90) | s(bool,X90) != s(bool,X91) | s(bool,X91) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f6427,f106])).
fof(f6427,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | s(bool,'F_0') = s(bool,X90) | s(bool,X90) != s(bool,X91) | s(bool,X91) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f4096,f539])).
fof(f539,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X21)) ) | $spl6),
  inference(subsumption_resolution,[],[f533,f45])).
fof(f533,plain,(
  ( ! [X21] : (p(s(bool,'F_0')) | s(bool,X21) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X21)) ) | $spl6),
  inference(superposition,[],[f110,f73])).
fof(f4096,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) != s(bool,X13) | s(bool,X11) != s(bool,X12)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f4050,f45])).
fof(f4050,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) != s(bool,X13) | p(s(bool,X11)) | p(s(bool,'F_0'))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f651,f37])).
fof(f651,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,'F_0') = s(bool,X10) | s(bool,X9) != s(bool,X10)) ) | ($spl6 | $spl22)),
  inference(resolution,[],[f641,f309])).
fof(f6434,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') = s(bool,X36) | s(bool,X36) != s(bool,X37) | s(bool,X34) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f6411,f106])).
fof(f6411,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X36) != s(bool,X37) | s(bool,X34) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f4096,f520])).
fof(f6433,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,X32) | s(bool,X32) != s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f6410,f106])).
fof(f6410,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X32) != s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f4096,f520])).
fof(f6358,plain,(
  ( ! [X30,X33,X32] : (s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') != s(bool,X33)) ) | $spl128),
  inference(cnf_transformation,[],[f6358_D])).
fof(f6358_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') != s(bool,X33)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f6359,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') != s(bool,X37) | s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f6335,f137])).
fof(f6335,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') != s(bool,X37) | s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f4004,f520])).
fof(f4004,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X13) | s(bool,X11) != s(bool,X12)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3955,f45])).
fof(f3955,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X13) | p(s(bool,X11)) | p(s(bool,'F_0'))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f597,f37])).
fof(f597,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,X9) != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) ) | ($spl16 | $spl22)),
  inference(resolution,[],[f387,f309])).
fof(f6357,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f6334,f137])).
fof(f6334,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f4004,f520])).
fof(f6295,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,X90) | s(bool,X90) != s(bool,X91) | s(bool,X91) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) | $spl126),
  inference(cnf_transformation,[],[f6295_D])).
fof(f6295_D,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,X90) | s(bool,X90) != s(bool,X91) | s(bool,X91) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f6294,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') = s(bool,X90) | s(bool,X91) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f6277,f106])).
fof(f6277,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | s(bool,'T_0') = s(bool,X90) | s(bool,X91) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f830,f539])).
fof(f5373,plain,(
  ( ! [X30,X33,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X33)) ) | $spl94),
  inference(cnf_transformation,[],[f5373_D])).
fof(f5373_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X33)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f6284,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl16),
  inference(subsumption_resolution,[],[f6261,f106])).
fof(f6261,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl16),
  inference(superposition,[],[f830,f520])).
fof(f6283,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl16),
  inference(subsumption_resolution,[],[f6260,f106])).
fof(f6260,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl16),
  inference(superposition,[],[f830,f520])).
fof(f6218,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') != s(bool,X83) | s(bool,X84) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X83) = s(bool,X84)) ) | $spl124),
  inference(cnf_transformation,[],[f6218_D])).
fof(f6218_D,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') != s(bool,X83) | s(bool,X84) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X83) = s(bool,X84)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6217,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X83) = s(bool,X84) | s(bool,X84) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X82)) ) | $spl16),
  inference(subsumption_resolution,[],[f6195,f1269])).
fof(f6195,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | s(bool,'T_0') != s(bool,X83) | s(bool,X83) = s(bool,X84) | s(bool,X84) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X82)) ) | $spl16),
  inference(superposition,[],[f3843,f1644])).
fof(f6209,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(subsumption_resolution,[],[f6187,f1269])).
fof(f6187,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') != s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(superposition,[],[f3843,f1623])).
fof(f6207,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(subsumption_resolution,[],[f6186,f1269])).
fof(f6186,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'T_0') != s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(superposition,[],[f3843,f1623])).
fof(f6136,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl119),
  inference(cnf_transformation,[],[f6136_D])).
fof(f6136_D,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f6138,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),sK0)),s('type.integer.int',sK0))) | $spl120),
  inference(cnf_transformation,[],[f6138_D])).
fof(f6138_D,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),sK0)),s('type.integer.int',sK0))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f6134,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),sK0)),s('type.integer.int',sK0))) | s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl8),
  inference(forward_demodulation,[],[f6067,f391])).
fof(f6067,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)),sK0)),s('type.integer.int',sK0)),sK0)),s('type.integer.int',sK0))) | $spl8),
  inference(superposition,[],[f1800,f391])).
fof(f1800,plain,(
  ( ! [X0] : (s('type.integer.int',sK1) != s('type.integer.int',X0) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X0,sK0)),s('type.integer.int',sK0)))) )),
  inference(superposition,[],[f1463,f49])).
fof(f1463,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X0,X1)),s('type.integer.int',X1)))) )),
  inference(resolution,[],[f35,f36])).
fof(f5929,plain,(
  ( ! [X83,X84] : (s(bool,X83) != s(bool,X84) | s(bool,X83) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X84)) ) | $spl116),
  inference(cnf_transformation,[],[f5929_D])).
fof(f5929_D,plain,(
  ( ! [X83,X84] : (s(bool,X83) != s(bool,X84) | s(bool,X83) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X84)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f5928,plain,(
  ( ! [X83,X84,X82] : (s(bool,X83) != s(bool,X84) | s(bool,'T_0') = s(bool,X84) | s(bool,X83) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X82)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f5907,f1269])).
fof(f5907,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | s(bool,X83) != s(bool,X84) | s(bool,'T_0') = s(bool,X84) | s(bool,X83) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X82)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f2999,f1644])).
fof(f5921,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X55) != s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X53) != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f5899,f1269])).
fof(f5899,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X55) != s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X53) != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f2999,f1623])).
fof(f5919,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f5898,f1269])).
fof(f5898,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f2999,f1623])).
fof(f5865,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X71),s('type.integer.int',X73))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73)))),X70)),s('type.integer.int',X70))) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X70),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73)))))) )),
  inference(subsumption_resolution,[],[f5832,f1269])).
fof(f5832,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,X74)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X71),s('type.integer.int',X73))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73)))),X70)),s('type.integer.int',X70))) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X70),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73)))))) )),
  inference(superposition,[],[f1532,f1623])).
fof(f5863,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))),X65)),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f5831,f1269])).
fof(f5831,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X69)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))),X65)),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f1532,f1623])).
fof(f5769,plain,(
  ( ! [X83,X84,X82] : (s(bool,'F_0') = s(bool,X83) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X83) != s(bool,X84) | s(bool,'F_0') = s(bool,X82)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f5749,f1269])).
fof(f5749,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | s(bool,'F_0') = s(bool,X83) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X83) != s(bool,X84) | s(bool,'F_0') = s(bool,X82)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f653,f1644])).
fof(f5762,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'F_0') = s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f5741,f1269])).
fof(f5741,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f653,f1623])).
fof(f5761,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f5740,f1269])).
fof(f5740,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'F_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f653,f1623])).
fof(f5706,plain,(
  ( ! [X83,X84] : (s(bool,X83) != s(bool,X84) | s(bool,'T_0') != s(bool,X83) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl108),
  inference(cnf_transformation,[],[f5706_D])).
fof(f5706_D,plain,(
  ( ! [X83,X84] : (s(bool,X83) != s(bool,X84) | s(bool,'T_0') != s(bool,X83) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f5705,plain,(
  ( ! [X83,X84,X82] : (s(bool,X83) != s(bool,X84) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') != s(bool,X83) | s(bool,'F_0') = s(bool,X82)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f5683,f1269])).
fof(f5683,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | s(bool,X83) != s(bool,X84) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') != s(bool,X83) | s(bool,'F_0') = s(bool,X82)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f599,f1644])).
fof(f5697,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X55) != s(bool,X56) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f5675,f1269])).
fof(f5675,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X55) != s(bool,X56) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f599,f1623])).
fof(f5695,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X51) != s(bool,X52) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f5674,f1269])).
fof(f5674,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X51) != s(bool,X52) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f599,f1623])).
fof(f5641,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X71),s('type.integer.int',X73))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73)))),X70)),s('type.integer.int',X70))) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X70),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73)))))) )),
  inference(subsumption_resolution,[],[f5607,f1269])).
fof(f5607,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,X74)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X71),s('type.integer.int',X73))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73)))),X70)),s('type.integer.int',X70))) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X70),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73)))))) )),
  inference(superposition,[],[f1530,f1623])).
fof(f5639,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))),X65)),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f5606,f1269])).
fof(f5606,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X69)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))),X65)),s('type.integer.int',X65))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f1530,f1623])).
fof(f5544,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') = s(bool,X83) | s(bool,X84) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X83) = s(bool,X84)) ) | $spl102),
  inference(cnf_transformation,[],[f5544_D])).
fof(f5544_D,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') = s(bool,X83) | s(bool,X84) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X83) = s(bool,X84)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f5543,plain,(
  ( ! [X83,X84,X82] : (s(bool,'F_0') = s(bool,X83) | s(bool,X83) = s(bool,X84) | s(bool,X84) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X82)) ) | $spl16),
  inference(subsumption_resolution,[],[f5521,f1269])).
fof(f5521,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | s(bool,'F_0') = s(bool,X83) | s(bool,X83) = s(bool,X84) | s(bool,X84) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'F_0') = s(bool,X82)) ) | $spl16),
  inference(superposition,[],[f2350,f1644])).
fof(f5535,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'F_0') = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(subsumption_resolution,[],[f5513,f1269])).
fof(f5513,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(superposition,[],[f2350,f1623])).
fof(f5533,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(subsumption_resolution,[],[f5512,f1269])).
fof(f5512,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(superposition,[],[f2350,f1623])).
fof(f5479,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s('type.integer.int',X74) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X74,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X71),s('type.integer.int',X73))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))),s('type.integer.int',X74)))) )),
  inference(subsumption_resolution,[],[f5446,f1269])).
fof(f5446,plain,(
  ( ! [X74,X72,X71,X75,X73] : (~p(s(bool,X75)) | s('type.integer.int',X74) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X74,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X71),s('type.integer.int',X73))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))),s('type.integer.int',X74)))) )),
  inference(superposition,[],[f1524,f1623])).
fof(f5477,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s('type.integer.int',X69) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X69,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))),s('type.integer.int',X69))) | s(bool,'F_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f5445,f1269])).
fof(f5445,plain,(
  ( ! [X70,X68,X66,X69,X67] : (~p(s(bool,X70)) | s('type.integer.int',X69) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X69,'const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))),s('type.integer.int',X69))) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f1524,f1623])).
fof(f5387,plain,(
  ( ! [X90,X91] : (s(bool,X90) = s(bool,X91) | s(bool,X90) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X91)) ) | $spl96),
  inference(cnf_transformation,[],[f5387_D])).
fof(f5387_D,plain,(
  ( ! [X90,X91] : (s(bool,X90) = s(bool,X91) | s(bool,X90) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X91)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f5386,plain,(
  ( ! [X90,X91,X89] : (s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91) | s(bool,X90) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f5366,f106])).
fof(f5366,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91) | s(bool,X90) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f2314,f539])).
fof(f2314,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f2289,f45])).
fof(f2289,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X3)) | p(s(bool,'F_0'))) ) | $spl22),
  inference(superposition,[],[f460,f37])).
fof(f460,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl22),
  inference(resolution,[],[f313,f106])).
fof(f5374,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl22),
  inference(subsumption_resolution,[],[f5350,f106])).
fof(f5350,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl22),
  inference(superposition,[],[f2314,f520])).
fof(f5372,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl22),
  inference(subsumption_resolution,[],[f5349,f106])).
fof(f5349,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl22),
  inference(superposition,[],[f2314,f520])).
fof(f5318,plain,(
  ( ! [X90,X91] : (s(bool,'F_0') != s(bool,X90) | s(bool,X91) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X90) = s(bool,X91)) ) | $spl92),
  inference(cnf_transformation,[],[f5318_D])).
fof(f5318_D,plain,(
  ( ! [X90,X91] : (s(bool,'F_0') != s(bool,X90) | s(bool,X91) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X90) = s(bool,X91)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f5317,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') != s(bool,X90) | s(bool,X90) = s(bool,X91) | s(bool,X91) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f5297,f106])).
fof(f5297,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | s(bool,'F_0') != s(bool,X90) | s(bool,X90) = s(bool,X91) | s(bool,X91) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f2281,f539])).
fof(f2281,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X5) | s(bool,X4) = s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f2257,f45])).
fof(f2257,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'F_0') != s(bool,X5) | s(bool,X4) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,'F_0'))) ) | $spl22),
  inference(superposition,[],[f459,f37])).
fof(f459,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) = s(bool,X4)) ) | $spl22),
  inference(resolution,[],[f313,f309])).
fof(f5304,plain,(
  ( ! [X30,X33,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) != s(bool,X33) | s(bool,X32) = s(bool,X33)) ) | $spl90),
  inference(cnf_transformation,[],[f5304_D])).
fof(f5304_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) != s(bool,X33) | s(bool,X32) = s(bool,X33)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f5305,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') != s(bool,X36) | s(bool,X36) = s(bool,X37) | s(bool,X34) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl22),
  inference(subsumption_resolution,[],[f5281,f106])).
fof(f5281,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') != s(bool,X36) | s(bool,X36) = s(bool,X37) | s(bool,X34) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl22),
  inference(superposition,[],[f2281,f520])).
fof(f5303,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,X32) = s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl22),
  inference(subsumption_resolution,[],[f5280,f106])).
fof(f5280,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') != s(bool,X32) | s(bool,X32) = s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl22),
  inference(superposition,[],[f2281,f520])).
fof(f5234,plain,(
  ( ! [X30,X33,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33)) ) | $spl88),
  inference(cnf_transformation,[],[f5234_D])).
fof(f5234_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f5235,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') != s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X34)) ) | $spl16),
  inference(subsumption_resolution,[],[f5211,f137])).
fof(f5211,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') != s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl16),
  inference(superposition,[],[f545,f520])).
fof(f545,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,X14) | s(bool,X15) = s(bool,X16) | s(bool,X14) = s(bool,X15)) ) | $spl16),
  inference(resolution,[],[f206,f37])).
fof(f5233,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') != s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | $spl16),
  inference(subsumption_resolution,[],[f5210,f137])).
fof(f5210,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') != s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl16),
  inference(superposition,[],[f545,f520])).
fof(f5177,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s('type.integer.int',X74) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X74,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X71),s('type.integer.int',X73))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))),s('type.integer.int',X74)))) )),
  inference(subsumption_resolution,[],[f5143,f1269])).
fof(f5143,plain,(
  ( ! [X74,X72,X71,X75,X73] : (~p(s(bool,X75)) | s('type.integer.int',X74) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X74,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X71),s('type.integer.int',X73))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X71),s('type.integer.int',X72))),s('type.integer.int',X73))),s('type.integer.int',X74)))) )),
  inference(superposition,[],[f1522,f1623])).
fof(f5175,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s('type.integer.int',X69) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X69,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))),s('type.integer.int',X69))) | s(bool,'F_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f5142,f1269])).
fof(f5142,plain,(
  ( ! [X70,X68,X66,X69,X67] : (~p(s(bool,X70)) | s('type.integer.int',X69) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X69,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X68))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68)))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',X67))),s('type.integer.int',X68))),s('type.integer.int',X69))) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f1522,f1623])).
fof(f5083,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,X83) | s(bool,X83) = s(bool,X84) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) | $spl84),
  inference(cnf_transformation,[],[f5083_D])).
fof(f5083_D,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,X83) | s(bool,X83) = s(bool,X84) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5082,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') = s(bool,X83) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X83) = s(bool,X84) | s(bool,'F_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f5061,f1269])).
fof(f5061,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | s(bool,'T_0') = s(bool,X83) | s(bool,X84) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,X83) = s(bool,X84) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f521,f1644])).
fof(f5074,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') = s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f5053,f1269])).
fof(f5053,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f521,f1623])).
fof(f5073,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f5052,f1269])).
fof(f5052,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f521,f1623])).
fof(f5023,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') = s(bool,X90) | s(bool,X91) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X89)) ) | $spl6),
  inference(subsumption_resolution,[],[f5003,f106])).
fof(f5003,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | s(bool,'F_0') = s(bool,X90) | s(bool,X91) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X89)) ) | $spl6),
  inference(superposition,[],[f1617,f539])).
fof(f1617,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s(bool,X4) = s(bool,X5) | s(bool,X3) = s(bool,X4)) )),
  inference(resolution,[],[f1369,f37])).
fof(f5011,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4987,f106])).
fof(f4987,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1617,f520])).
fof(f5009,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f4986,f106])).
fof(f4986,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f1617,f520])).
fof(f4955,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X63))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X63)),X66)),s('type.integer.int',X66))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X66),s('type.integer.int',X64)))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(subsumption_resolution,[],[f4928,f1269])).
fof(f4928,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X63))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X63)),X66)),s('type.integer.int',X66))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X66),s('type.integer.int',X64)))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f1485,f1623])).
fof(f4953,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59)),X62)),s('type.integer.int',X62))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X62),s('type.integer.int',X60)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f4927,f1269])).
fof(f4927,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59)),X62)),s('type.integer.int',X62))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X62),s('type.integer.int',X60)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1485,f1623])).
fof(f4944,plain,(
  ( ! [X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)),X1)),s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int',sK1))))) ) | $spl76),
  inference(cnf_transformation,[],[f4944_D])).
fof(f4944_D,plain,(
  ( ! [X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)),X1)),s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int',sK1))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4943,plain,(
  ( ! [X0,X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)),X1)),s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f4942,f49])).
fof(f4942,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)),X1)),s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4908,f1269])).
fof(f4908,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)),X1)),s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1485,f1644])).
fof(f4870,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'F_0') != s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | $spl22),
  inference(subsumption_resolution,[],[f4848,f309])).
fof(f4848,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'F_0') != s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl22),
  inference(superposition,[],[f461,f1623])).
fof(f4868,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'F_0') != s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | $spl22),
  inference(subsumption_resolution,[],[f4847,f309])).
fof(f4847,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'F_0') != s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl22),
  inference(superposition,[],[f461,f1623])).
fof(f4685,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.integer.int',X66) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X66,'const.integer.int_mul_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X63)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X63))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X64),s('type.integer.int',X66)))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(subsumption_resolution,[],[f4660,f1269])).
fof(f4660,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.integer.int',X66) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X66,'const.integer.int_mul_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X63)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X63))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X64),s('type.integer.int',X66)))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f1480,f1623])).
fof(f4683,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int',X62) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X62,'const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X60),s('type.integer.int',X62)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f4659,f1269])).
fof(f4659,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | s('type.integer.int',X62) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X62,'const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X59))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X60),s('type.integer.int',X62)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1480,f1623])).
fof(f4674,plain,(
  ( ! [X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X1,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)))),s('type.integer.int',sK1))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK1),s('type.integer.int',X1))))) ) | $spl70),
  inference(cnf_transformation,[],[f4674_D])).
fof(f4674_D,plain,(
  ( ! [X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X1,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)))),s('type.integer.int',sK1))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK1),s('type.integer.int',X1))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4673,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X1,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)))),s('type.integer.int',sK1))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK1),s('type.integer.int',X1)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f4672,f49])).
fof(f4672,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X1,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK1),s('type.integer.int',X1)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4640,f1269])).
fof(f4640,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X1,'const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0)))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',sK0))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK1),s('type.integer.int',X1)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1480,f1644])).
fof(f4602,plain,(
  ( ! [X59,X57,X58,X56] : (s('type.integer.int',X58) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X58,'const.integer.int_mul_2'(s('type.integer.int',sK3(X56,X57)),s('type.integer.int',X57)))),s('type.integer.int',X56))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X57),s('type.integer.int',X56)))) | s(bool,'F_0') = s(bool,X59) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X56,X57)),s('type.integer.int',X57))),s('type.integer.int',X58)))) )),
  inference(subsumption_resolution,[],[f4575,f1269])).
fof(f4575,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X59)) | s('type.integer.int',X58) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X58,'const.integer.int_mul_2'(s('type.integer.int',sK3(X56,X57)),s('type.integer.int',X57)))),s('type.integer.int',X56))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X57),s('type.integer.int',X56)))) | s(bool,'F_0') = s(bool,X59) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X56,X57)),s('type.integer.int',X57))),s('type.integer.int',X58)))) )),
  inference(superposition,[],[f1475,f1623])).
fof(f4600,plain,(
  ( ! [X54,X52,X55,X53] : (s('type.integer.int',X54) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X54,'const.integer.int_mul_2'(s('type.integer.int',sK3(X52,X53)),s('type.integer.int',X53)))),s('type.integer.int',X52))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X53),s('type.integer.int',X52)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X52,X53)),s('type.integer.int',X53))),s('type.integer.int',X54))) | s(bool,'F_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f4574,f1269])).
fof(f4574,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X55)) | s('type.integer.int',X54) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X54,'const.integer.int_mul_2'(s('type.integer.int',sK3(X52,X53)),s('type.integer.int',X53)))),s('type.integer.int',X52))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X53),s('type.integer.int',X52)))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X52,X53)),s('type.integer.int',X53))),s('type.integer.int',X54))) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f1475,f1623])).
fof(f4532,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X64),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X63))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X66))),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(subsumption_resolution,[],[f4508,f1269])).
fof(f4508,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X64),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X63))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X66))),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f1502,f1623])).
fof(f4530,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X62))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f4507,f1269])).
fof(f4507,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X62))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1502,f1623])).
fof(f4521,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X1))),s('type.integer.int',sK0)))) ) | $spl64),
  inference(cnf_transformation,[],[f4521_D])).
fof(f4521_D,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X1))),s('type.integer.int',sK0)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4520,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X1))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4488,f1269])).
fof(f4488,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X1))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1502,f1644])).
fof(f4463,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X63))),s('type.integer.int',X64))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',sK3(X64,X63)))),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(subsumption_resolution,[],[f4439,f1269])).
fof(f4439,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X63))),s('type.integer.int',X64))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X66),s('type.integer.int',sK3(X64,X63)))),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f1501,f1623])).
fof(f4461,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))),s('type.integer.int',X60))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X62),s('type.integer.int',sK3(X60,X59)))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f4438,f1269])).
fof(f4438,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))),s('type.integer.int',X60))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X62),s('type.integer.int',sK3(X60,X59)))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1501,f1623])).
fof(f4452,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X1),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0)))) ) | $spl60),
  inference(cnf_transformation,[],[f4452_D])).
fof(f4452_D,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X1),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4451,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X1),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4419,f1269])).
fof(f4419,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X1),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1501,f1644])).
fof(f4394,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X64),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X63))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X66))),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(subsumption_resolution,[],[f4370,f1269])).
fof(f4370,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X64),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X63))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X64,X63)),s('type.integer.int',X66))),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f1500,f1623])).
fof(f4392,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X62))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f4369,f1269])).
fof(f4369,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(X60,X59)),s('type.integer.int',X62))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1500,f1623])).
fof(f4383,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X1))),s('type.integer.int',sK0)))) ) | $spl56),
  inference(cnf_transformation,[],[f4383_D])).
fof(f4383_D,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X1))),s('type.integer.int',sK0)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4382,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X1))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4350,f1269])).
fof(f4350,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3(sK1,sK0)),s('type.integer.int',X1))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1500,f1644])).
fof(f4325,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X63))),s('type.integer.int',X64))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',sK3(X64,X63)))),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(subsumption_resolution,[],[f4301,f1269])).
fof(f4301,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X66),s('type.integer.int',X63))),s('type.integer.int',X64))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',sK3(X64,X63)))),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f1499,f1623])).
fof(f4323,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))),s('type.integer.int',X60))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',sK3(X60,X59)))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f4300,f1269])).
fof(f4300,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X62),s('type.integer.int',X59))),s('type.integer.int',X60))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',sK3(X60,X59)))),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1499,f1623])).
fof(f4314,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0)))) ) | $spl52),
  inference(cnf_transformation,[],[f4314_D])).
fof(f4314_D,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4313,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4281,f1269])).
fof(f4281,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK3(sK1,sK0)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1499,f1644])).
fof(f4105,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') = s(bool,X27) | s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f4073,f44])).
fof(f4073,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f651,f520])).
fof(f4104,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,X24) | s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f4072,f44])).
fof(f4072,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f651,f520])).
fof(f3919,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.integer.int',X64) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X64))),s('type.integer.int',X66)),X63)),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(subsumption_resolution,[],[f3895,f1269])).
fof(f3895,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.integer.int',X64) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X66),s('type.integer.int',X64))),s('type.integer.int',X66)),X63)),s('type.integer.int',X63))) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f1531,f1623])).
fof(f3917,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.integer.int',X60) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X60))),s('type.integer.int',X62)),X59)),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f3894,f1269])).
fof(f3894,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | s('type.integer.int',X60) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X60))),s('type.integer.int',X62)),X59)),s('type.integer.int',X59))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1531,f1623])).
fof(f3908,plain,(
  ( ! [X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK1))),s('type.integer.int',X1)),sK0)),s('type.integer.int',sK0)))) ) | $spl48),
  inference(cnf_transformation,[],[f3908_D])).
fof(f3908_D,plain,(
  ( ! [X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK1))),s('type.integer.int',X1)),sK0)),s('type.integer.int',sK0)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3907,plain,(
  ( ! [X0,X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK1))),s('type.integer.int',X1)),sK0)),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3875,f1269])).
fof(f3875,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3('const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK1))),s('type.integer.int',X1)),sK0)),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1531,f1644])).
fof(f3699,plain,(
  ( ! [X88,X87,X85] : (s(bool,X85) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))) | s(bool,'T_0') = s(bool,X85) | s('type.integer.int',X88) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X88,X87)),s('type.integer.int',X87)))) ) | $spl46),
  inference(cnf_transformation,[],[f3699_D])).
fof(f3699_D,plain,(
  ( ! [X88,X87,X85] : (s(bool,X85) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))) | s(bool,'T_0') = s(bool,X85) | s('type.integer.int',X88) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X88,X87)),s('type.integer.int',X87)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3719,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))) | s('type.integer.int',X88) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X88,X87)),s('type.integer.int',X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f3297,f106])).
fof(f3297,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | s(bool,X85) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))) | s('type.integer.int',X88) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X88,X87)),s('type.integer.int',X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f1465,f520])).
fof(f1465,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X7)) | s(bool,X7) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X6),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X5,X6)),s('type.integer.int',X6)))) )),
  inference(resolution,[],[f35,f37])).
fof(f3717,plain,(
  ( ! [X66,X67,X65] : (s(bool,X65) != s(bool,X67) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f3290,f106])).
fof(f3290,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | s(bool,X65) != s(bool,X67) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f641,f520])).
fof(f3698,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))) | s('type.integer.int',X88) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X88,X87)),s('type.integer.int',X87))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f3229,f106])).
fof(f3229,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | s(bool,X85) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))) | s('type.integer.int',X88) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X88,X87)),s('type.integer.int',X87))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f1465,f520])).
fof(f3696,plain,(
  ( ! [X66,X67,X65] : (s(bool,X65) != s(bool,X67) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X66)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f3222,f106])).
fof(f3222,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | s(bool,X65) != s(bool,X67) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X66)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f641,f520])).
fof(f3072,plain,(
  ( ! [X47,X50,X48,X49] : (s('type.integer.int',X48) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X48,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X50),s('type.integer.int',X47))),s('type.integer.int',X50)))),s('type.integer.int',X47))) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int',X48)))) )),
  inference(subsumption_resolution,[],[f3050,f1269])).
fof(f3050,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | s('type.integer.int',X48) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X48,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X50),s('type.integer.int',X47))),s('type.integer.int',X50)))),s('type.integer.int',X47))) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int',X48)))) )),
  inference(superposition,[],[f1523,f1623])).
fof(f3070,plain,(
  ( ! [X45,X43,X46,X44] : (s('type.integer.int',X44) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X44,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X46),s('type.integer.int',X43))),s('type.integer.int',X46)))),s('type.integer.int',X43))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int',X44))) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f3049,f1269])).
fof(f3049,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | s('type.integer.int',X44) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(X44,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X46),s('type.integer.int',X43))),s('type.integer.int',X46)))),s('type.integer.int',X43))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int',X44))) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f1523,f1623])).
fof(f3063,plain,(
  ( ! [X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',X1)))),s('type.integer.int',sK0)))) ) | $spl42),
  inference(cnf_transformation,[],[f3063_D])).
fof(f3063_D,plain,(
  ( ! [X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',X1)))),s('type.integer.int',sK0)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3062,plain,(
  ( ! [X0,X1] : (s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',X1)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3034,f1269])).
fof(f3034,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK3(sK1,'const.integer.int_sub_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK0))),s('type.integer.int',X1)))),s('type.integer.int',sK0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1523,f1644])).
fof(f3007,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X27)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2982,f312])).
fof(f2982,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X28) | s(bool,'F_0') = s(bool,X27)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f368,f1623])).
fof(f3005,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2981,f312])).
fof(f2981,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X25)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f368,f1623])).
fof(f2905,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,X74) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X76),s('type.integer.int',X77))) | s('type.integer.int',X77) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X78),s('type.integer.int',X76))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2542,f1269])).
fof(f2542,plain,(
  ( ! [X78,X76,X74,X77,X75] : (~p(s(bool,X75)) | s(bool,X74) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X76),s('type.integer.int',X77))) | s('type.integer.int',X77) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X78),s('type.integer.int',X76))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X74)) )),
  inference(superposition,[],[f1311,f1623])).
fof(f2904,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) ) | $spl16),
  inference(subsumption_resolution,[],[f2539,f1269])).
fof(f2539,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,X68)) | s(bool,X67) != s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) ) | $spl16),
  inference(superposition,[],[f820,f1623])).
fof(f2884,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,X74) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X76),s('type.integer.int',X77))) | s('type.integer.int',X77) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X78),s('type.integer.int',X76))) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2483,f1269])).
fof(f2483,plain,(
  ( ! [X78,X76,X74,X77,X75] : (~p(s(bool,X75)) | s(bool,X74) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X76),s('type.integer.int',X77))) | s('type.integer.int',X77) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X78),s('type.integer.int',X76))) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,X75)) )),
  inference(superposition,[],[f1311,f1623])).
fof(f2882,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X67) | s(bool,'F_0') = s(bool,X68)) ) | $spl16),
  inference(subsumption_resolution,[],[f2480,f1269])).
fof(f2480,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,X68)) | s(bool,X67) != s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X67) | s(bool,'F_0') = s(bool,X68)) ) | $spl16),
  inference(superposition,[],[f820,f1623])).
fof(f994,plain,(
  ( ! [X43] : (s(bool,X43) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X43)) ) | $spl30),
  inference(cnf_transformation,[],[f994_D])).
fof(f994_D,plain,(
  ( ! [X43] : (s(bool,X43) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X43)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2046,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,X39)) ) | $spl16),
  inference(subsumption_resolution,[],[f1991,f1269])).
fof(f1991,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | s(bool,X40) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,X39)) ) | $spl16),
  inference(superposition,[],[f820,f1644])).
fof(f1694,plain,(
  ( ! [X24,X25] : (s(bool,X25) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1685,f106])).
fof(f1685,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | s(bool,X25) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f641,f539])).
fof(f1587,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') != s(bool,X25)) ) | $spl34),
  inference(cnf_transformation,[],[f1587_D])).
fof(f1587_D,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') != s(bool,X25)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1586,plain,(
  ( ! [X24,X25] : (s(bool,X25) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1577,f106])).
fof(f1577,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | s(bool,X25) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f313,f539])).
fof(f1360,plain,(
  ( ! [X28,X26,X27] : (s('type.integer.int',X27) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',X26))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X26),s('type.integer.int',X27)))) ) | $spl32),
  inference(cnf_transformation,[],[f1360_D])).
fof(f1360_D,plain,(
  ( ! [X28,X26,X27] : (s('type.integer.int',X27) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',X26))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X26),s('type.integer.int',X27)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f136,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | $spl14),
  inference(cnf_transformation,[],[f136_D])).
fof(f136_D,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1334,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | s('type.integer.int',X27) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int',X26))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X26),s('type.integer.int',X27)))) ) | $spl6),
  inference(superposition,[],[f36,f539])).
fof(f135,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | $spl15),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1306,plain,(
  ( ! [X10] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1297,f44])).
fof(f1297,plain,(
  ( ! [X10] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f309,f539])).
fof(f1268,plain,(
  $false | ($spl16 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f1242,f793])).
fof(f1242,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33)) ) | ($spl16 | $spl22 | $spl24)),
  inference(duplicate_literal_removal,[],[f1233])).
fof(f1233,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') != s(bool,X33)) ) | ($spl16 | $spl22 | $spl24)),
  inference(superposition,[],[f312,f793])).
fof(f1253,plain,(
  $false | ($spl16 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f1252,f793])).
fof(f1252,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl16 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f1168,f137])).
fof(f1168,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl22 | $spl24)),
  inference(backward_demodulation,[],[f793,f309])).
fof(f1250,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1164,f47])).
fof(f1164,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f793,f45])).
fof(f1241,plain,(
  $false | ($spl3 | $spl24)),
  inference(trivial_inequality_removal,[],[f1234])).
fof(f1234,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl24)),
  inference(superposition,[],[f134,f793])).
fof(f134,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(subsumption_resolution,[],[f128,f70])).
fof(f128,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | $spl3),
  inference(superposition,[],[f70,f44])).
fof(f1238,plain,(
  $false | ($spl3 | $spl24)),
  inference(trivial_inequality_removal,[],[f1237])).
fof(f1237,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl24)),
  inference(superposition,[],[f70,f793])).
fof(f1131,plain,(
  $false | $spl28),
  inference(resolution,[],[f1049,f47])).
fof(f1049,plain,(
  ( ! [X5] : (~p(s(bool,X5))) ) | $spl28),
  inference(superposition,[],[f45,f993])).
fof(f1132,plain,(
  $false | $spl28),
  inference(resolution,[],[f1049,f46])).
fof(f46,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_LADD','thm.bool.TRUTH')).
fof(f1133,plain,(
  $false | $spl28),
  inference(resolution,[],[f1049,f48])).
fof(f1123,plain,(
  ( ! [X47,X48] : (~p(s(bool,X47)) | s(bool,'T_0') = s(bool,X48)) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1073,f44])).
fof(f1073,plain,(
  ( ! [X47,X48] : (~p(s(bool,X47)) | s(bool,'F_0') != s(bool,X48) | s(bool,'T_0') = s(bool,X48)) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f820,f993])).
fof(f992,plain,(
  ( ! [X43,X42] : (s(bool,X43) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X42)) ) | $spl16),
  inference(subsumption_resolution,[],[f937,f218])).
fof(f937,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | s(bool,X43) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X42)) ) | $spl16),
  inference(superposition,[],[f820,f592])).
fof(f792,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f740,f106])).
fof(f740,plain,(
  ( ! [X39,X40] : (p(s(bool,X39)) | s(bool,X40) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f641,f539])).
fof(f432,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f68,f264])).
fof(f264,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl6),
  inference(backward_demodulation,[],[f73,f32])).
fof(f426,plain,(
  $false | ($spl6 | $spl10 | $spl22)),
  inference(trivial_inequality_removal,[],[f424])).
fof(f424,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl10 | $spl22)),
  inference(superposition,[],[f408,f73])).
fof(f408,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl6 | $spl10 | $spl22)),
  inference(subsumption_resolution,[],[f407,f309])).
fof(f407,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f77,f73])).
fof(f425,plain,(
  $false | ($spl6 | $spl10 | $spl22)),
  inference(equality_resolution,[],[f408])).
fof(f309,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') != s(bool,X2)) ) | $spl22),
  inference(cnf_transformation,[],[f309_D])).
fof(f309_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f304,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ~p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f264,f40])).
fof(f303,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl6),
  inference(superposition,[],[f264,f40])).
fof(f308,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f302,f218])).
fof(f302,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl6),
  inference(superposition,[],[f264,f37])).
fof(f306,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | p(s(bool,X0))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f301,f218])).
fof(f301,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | p(s(bool,X0))) ) | $spl6),
  inference(superposition,[],[f264,f37])).
fof(f259,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f258,f70])).
fof(f258,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | $spl7),
  inference(trivial_inequality_removal,[],[f253])).
fof(f253,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | $spl7),
  inference(superposition,[],[f74,f44])).
fof(f226,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f32,f40])).
fof(f225,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) )),
  inference(superposition,[],[f32,f40])).
fof(f224,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) )),
  inference(superposition,[],[f32,f37])).
fof(f223,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f32,f37])).
fof(f222,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(superposition,[],[f32,f44])).
fof(f221,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(superposition,[],[f32,f44])).
fof(f181,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ~p(s(bool,X3))) ) | $spl2),
  inference(subsumption_resolution,[],[f177,f106])).
fof(f177,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ~p(s(bool,X3))) ) | $spl2),
  inference(superposition,[],[f141,f40])).
fof(f141,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl2),
  inference(backward_demodulation,[],[f69,f32])).
fof(f180,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f176,f106])).
fof(f176,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl2),
  inference(superposition,[],[f141,f40])).
fof(f137,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f175,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl2),
  inference(superposition,[],[f141,f37])).
fof(f174,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | p(s(bool,X0))) ) | $spl2),
  inference(superposition,[],[f141,f37])).
fof(f140,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | ~p(s(bool,X3))) ) | $spl3),
  inference(subsumption_resolution,[],[f132,f106])).
fof(f132,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | ~p(s(bool,X3))) ) | $spl3),
  inference(superposition,[],[f70,f40])).
fof(f138,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))))) ) | $spl3),
  inference(subsumption_resolution,[],[f131,f106])).
fof(f131,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))))) ) | $spl3),
  inference(superposition,[],[f70,f40])).
fof(f130,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))))) ) | $spl3),
  inference(superposition,[],[f70,f37])).
fof(f129,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2)))) | p(s(bool,X0))) ) | $spl3),
  inference(superposition,[],[f70,f37])).
fof(f66,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f32,f40])).
fof(f65,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) )),
  inference(superposition,[],[f32,f40])).
fof(f64,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) )),
  inference(superposition,[],[f32,f37])).
fof(f63,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f32,f37])).
fof(f62,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(superposition,[],[f32,f44])).
fof(f61,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',sK0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))),
  inference(superposition,[],[f32,f44])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 191
% Passive clauses: 2656
% Generated clauses: 13917
% Final active clauses: 76
% Final passive clauses: 1506
% Input formulas: 12
% Initial clauses: 17
% 
% Selected by SInE selection: 11
% SInE iterations: 1
% 
% Duplicate literals: 503
% Trivial inequalities: 239
% Fw subsumption resolutions: 1869
% Fw demodulations: 2160
% Bw demodulations: 51
% 
% Simple tautologies: 377
% Equational tautologies: 41
% Forward subsumptions: 5860
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 751
% Forward superposition: 3933
% Backward superposition: 3768
% Self superposition: 103
% Equality factoring: 80
% Equality resolution: 58
% 
% Split clauses: 325
% Split components: 650
% SAT solver clauses: 350
% SAT solver unit clauses: 8
% SAT solver binary clauses: 122
% 
% Sat splits: 325
% Sat splitting refutations: 24
% 
% Memory used [KB]: 7419
% Time elapsed: 0.361 s
% ------------------------------
% ------------------------------
% Success in time 12.083 s
