Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f1307,plain,(
  $false),
  inference(subsumption_resolution,[],[f1305,f321])).
fof(f321,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f232,f187])).
fof(f187,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  file('thm.numeral.numeral_lt|split|2','thm.arithmetic.ALT_ZERO')).
fof(f232,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))),
  file('thm.numeral.numeral_lt|split|2','thm.prim_rec.NOT_LESS_|48|')).
fof(f1305,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.ZERO_0'))))),
  inference(trivial_inequality_removal,[],[f1279])).
fof(f1279,plain,(
  s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.ZERO_0'))))),
  inference(superposition,[],[f205,f1173])).
fof(f1173,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(trivial_inequality_removal,[],[f1169])).
fof(f1169,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f184,f198])).
fof(f198,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.numeral.numeral_lt|split|2',aHL_BOOL_CASES)).
fof(f184,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f114])).
fof(f114,plain,(
  ? [X0] : s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ~! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.numeral.numeral_lt|split|2','thm.numeral.numeral_lt|split|2')).
fof(f205,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.numeral.numeral_lt|split|2','thm.bool.EQ_CLAUSES|split|0')).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 128
% Passive clauses: 271
% Generated clauses: 1118
% Final active clauses: 127
% Final passive clauses: 119
% Input formulas: 50
% Initial clauses: 123
% 
% Duplicate literals: 23
% Trivial inequalities: 8
% Fw subsumption resolutions: 79
% Fw demodulations: 4
% Bw demodulations: 1
% 
% Simple tautologies: 43
% Equational tautologies: 4
% Forward subsumptions: 632
% 
% Binary resolution: 560
% Factoring: 168
% Forward superposition: 26
% Backward superposition: 105
% Self superposition: 4
% Equality factoring: 2
% Equality resolution: 3
% 
% Split clauses: 30
% Split components: 60
% SAT solver clauses: 53
% SAT solver unit clauses: 6
% SAT solver binary clauses: 14
% SAT solver learnt clauses: 1
% 
% Sat splits: 30
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 52
% TWLsolver calls for satisfiability: 36
% 
% Memory used [KB]: 1023
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
% Success in time 0.032 s
