Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (90872)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 722
% Passive clauses: 9474
% Generated clauses: 20259
% Final active clauses: 660
% Final passive clauses: 8178
% Input formulas: 37
% Initial clauses: 142
% 
% Duplicate literals: 400
% Trivial inequalities: 97
% Fw subsumption resolutions: 1943
% Fw demodulations: 783
% Bw demodulations: 61
% 
% Simple tautologies: 186
% Equational tautologies: 61
% Forward subsumptions: 7111
% 
% Binary resolution: 9762
% Factoring: 1163
% Forward superposition: 2837
% Backward superposition: 2784
% Self superposition: 35
% Equality factoring: 26
% Equality resolution: 32
% 
% Split clauses: 341
% Split components: 689
% SAT solver clauses: 718
% SAT solver unit clauses: 16
% SAT solver binary clauses: 150
% 
% Sat splits: 341
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 716
% TWLsolver calls for satisfiability: 461
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (90892)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 306
% Passive clauses: 2969
% Generated clauses: 14145
% Final active clauses: 211
% Final passive clauses: 1250
% Input formulas: 37
% Initial clauses: 142
% 
% Duplicate literals: 206
% Trivial inequalities: 31
% Fw subsumption resolutions: 1322
% Fw demodulations: 1335
% Bw demodulations: 1217
% Condensations: 333
% Global subsumptions: 14
% 
% Simple tautologies: 223
% Equational tautologies: 20
% Forward subsumptions: 4209
% Bw demodulations to eq. taut.: 258
% 
% Binary resolution: 1875
% Factoring: 61
% Forward superposition: 1455
% Backward superposition: 1920
% Self superposition: 36
% Equality factoring: 32
% Equality resolution: 8
% 
% Split clauses: 274
% Split components: 566
% SAT solver clauses: 3246
% SAT solver unit clauses: 160
% SAT solver binary clauses: 74
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 10
% 
% Sat splits: 274
% Sat splitting refutations: 28
% 
% TWLsolver clauses: 3189
% TWLsolver calls for satisfiability: 1407
% 
% Memory used [KB]: 11897
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f26469,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f200,f352,f350,f348,f199,f351,f198,f349,f347,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2495,f2429,f2497,f2433,f2499,f2436,f2502,f2440,f2504,f2443,f2506,f2446,f2508,f2449,f2510,f2452,f2512,f2455,f2514,f2458,f2516,f2461,f2518,f2464,f2520,f2467,f2653,f2654,f2659,f2661,f2663,f2665,f2667,f2669,f2671,f2673,f2675,f2677,f2679,f2681,f2683,f2685,f2687,f2696,f2698,f2700,f2702,f2704,f2706,f2708,f2710,f2712,f2714,f2716,f2718,f2720,f2722,f2724,f2726,f2728,f2730,f2732,f2734,f2736,f2738,f2740,f2742,f2744,f2746,f2748,f2750,f2752,f2754,f4832,f4835,f4833,f4864,f4867,f4865,f4899,f4902,f4900,f4933,f4936,f4934,f5217,f5218,f5225,f5222,f5219,f5220,f5228,f5221,f5230,f5223,f5227,f5700,f5701,f5713,f5711,f5833,f5836,f5834,f5886,f5889,f5887,f5939,f5942,f5940,f6139,f6142,f6140,f7838,f7841,f7839,f7874,f7877,f7875,f8858,f8861,f8859,f8897,f8900,f8898,f8938,f8941,f8939,f8979,f8982,f8980,f9021,f9024,f2430,f9022,f11339,f11340,f11342,f11343,f13275,f14102,f14147,f26359,f26360,f26361,f26362,f26364,f11341,f26365,f26465])).
fof(f26465,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f26464,f353])).
fof(f353,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f204,f201])).
fof(f201,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.IN_RDOM_RRESTRICT',aHL_TRUTH)).
fof(f204,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.bool.NOT_CLAUSES|split|2')).
fof(f26464,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f26463,f16495])).
fof(f16495,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f396,f14728,f605])).
fof(f605,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f600,f353])).
fof(f600,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f244,f208])).
fof(f208,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.IN_RDOM_RRESTRICT',aHL_BOOL_CASES)).
fof(f244,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.bool.IMP_ANTISYM_AX')).
fof(f14728,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,X0))))) ) | $spl1),
  inference(forward_demodulation,[],[f14549,f1343])).
fof(f1343,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'const.relation.RRESTRICT_4'(s(fun(X7,fun(X6,bool)),X8),s(fun(X7,bool),X9),s(X7,X10),s(X6,X11))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),'const.relation.RRESTRICT_2'(s(fun(X7,fun(X6,bool)),X8),s(fun(X7,bool),X9))),s(X7,X10)))))) )),
  inference(superposition,[],[f279,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X5,X11] : s(bool,'const.bool.IN_2'(s(X7,X5),s(fun(X7,bool),X11))) = s(bool,happ(s(fun(X7,bool),X11),s(X7,X5)))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.bool.IN_DEF')).
fof(f279,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RRESTRICT_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RRESTRICT_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.RRESTRICT_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RRESTRICT_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3))),s(X0,X4))),s(X1,X5)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.RRESTRICT_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RRESTRICT_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3))),s(X0,X4))),s(X1,X5)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X7,X6] : ! [X19,X20,X21,X22] : s(bool,'const.relation.RRESTRICT_4'(s(fun(X7,fun(X6,bool)),X19),s(fun(X7,bool),X20),s(X7,X21),s(X6,X22))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),'const.relation.RRESTRICT_2'(s(fun(X7,fun(X6,bool)),X19),s(fun(X7,bool),X20))),s(X7,X21))),s(X6,X22)))),
  file('thm.relation.IN_RDOM_RRESTRICT','ARITY_2_4_const.relation.RRESTRICT')).
fof(f14549,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f14148,f355])).
fof(f355,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f278,f243])).
fof(f278,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,sK16(X3,X2,X1,X0)))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ? [X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X5))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(rectify,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7,X6] : ! [X10,X8] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X7,fun(X6,bool)),X10),s(X7,X8)))) <=> ? [X15] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X10),s(X7,X8))),s(X6,X15)))))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.relation.RDOM_DEF')).
fof(f14148,plain,(
  ~p(s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12)))) | $spl1),
  inference(forward_demodulation,[],[f348,f1157])).
fof(f1157,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'const.relation.RDOM_2'(s(fun(X4,fun(X5,bool)),X6),s(X4,X7))) = s(bool,'const.bool.IN_2'(s(X4,X7),s(fun(X4,bool),'const.relation.RDOM_1'(s(fun(X4,fun(X5,bool)),X6)))))) )),
  inference(superposition,[],[f275,f243])).
fof(f275,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6,X7] : ! [X19,X20] : s(bool,'const.relation.RDOM_2'(s(fun(X7,fun(X6,bool)),X19),s(X7,X20))) = s(bool,happ(s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),X19))),s(X7,X20)))),
  file('thm.relation.IN_RDOM_RRESTRICT','ARITY_1_2_const.relation.RDOM')).
fof(f396,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f353,f216])).
fof(f216,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.bool.EQ_CLAUSES|split|1')).
fof(f26463,plain,(
  p(s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,sK16(sK12,sK14,sK10,sK11))))) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f26429,f201])).
fof(f26429,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,sK16(sK12,sK14,sK10,sK11))))) | ($spl2 | $spl4)),
  inference(superposition,[],[f15424,f16934])).
fof(f16934,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK10,sK16(sK12,sK14,sK10,sK11)),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))))) | $spl2),
  inference(unit_resulting_resolution,[],[f201,f14346,f247])).
fof(f247,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f14346,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK16(sK12,sK14,sK10,sK11)),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12)))))) | $spl2),
  inference(unit_resulting_resolution,[],[f14149,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK16(X3,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))) )),
  inference(forward_demodulation,[],[f277,f243])).
fof(f277,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,sK16(X3,X2,X1,X0))))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f14149,plain,(
  p(s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12)))) | $spl2),
  inference(forward_demodulation,[],[f349,f1157])).
fof(f15424,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),happ(s(fun(sK11,fun(X6,bool)),X7),s(sK11,sK12)))))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(X6,bool)),X7),s(fun(sK11,bool),sK13),s(sK11,sK12),s(X6,X8))))) ) | $spl4),
  inference(subsumption_resolution,[],[f15272,f201])).
fof(f15272,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(X6,bool)),X7),s(fun(sK11,bool),sK13),s(sK11,sK12),s(X6,X8)))) | ~p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),happ(s(fun(sK11,fun(X6,bool)),X7),s(sK11,sK12))))))) ) | $spl4),
  inference(superposition,[],[f357,f14189])).
fof(f14189,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13))) | $spl4),
  inference(unit_resulting_resolution,[],[f201,f351,f247])).
fof(f357,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))))))) )),
  inference(forward_demodulation,[],[f282,f243])).
fof(f282,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5))))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3))))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5))))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : ! [X10,X9,X8,X15] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X7,fun(X6,bool)),X10),s(fun(X7,bool),X9),s(X7,X8),s(X6,X15)))) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X10),s(X7,X8))),s(X6,X15)))) & p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X9))))))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.relation.RRESTRICT_DEF')).
fof(f26365,plain,(
  ( ! [X112,X110,X113] : (sP1('const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,X110)),X112) | s(bool,'T_0') = s(bool,X113) | s(bool,X112) = s(bool,X113)) ) | $spl90),
  inference(cnf_transformation,[],[f26365_D])).
fof(f26365_D,plain,(
  ( ! [X112,X110,X113] : (sP1('const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,X110)),X112) | s(bool,'T_0') = s(bool,X113) | s(bool,X112) = s(bool,X113)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f11341,plain,(
  ( ! [X30] : (p(s(bool,X30))) ) | $spl82),
  inference(cnf_transformation,[],[f11341_D])).
fof(f11341_D,plain,(
  ( ! [X30] : (p(s(bool,X30))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f26364,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X111)) | sP1('const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,X110)),X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) ) | $spl1),
  inference(subsumption_resolution,[],[f26342,f216])).
fof(f26342,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X111)) | sP1('const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,X110)),X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,X111)) ) | $spl1),
  inference(superposition,[],[f2473,f16600])).
fof(f16600,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(subsumption_resolution,[],[f16535,f201])).
fof(f16535,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(superposition,[],[f14728,f553])).
fof(f553,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f208,f208])).
fof(f2473,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2105,f201])).
fof(f2105,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f270,f553])).
fof(f270,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f162])).
fof(f162,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f26362,plain,(
  ( ! [X107,X106] : (sP1('const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12)),X106) | s(bool,'T_0') = s(bool,X107) | s(bool,X106) = s(bool,X107)) ) | $spl88),
  inference(cnf_transformation,[],[f26362_D])).
fof(f26362_D,plain,(
  ( ! [X107,X106] : (sP1('const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12)),X106) | s(bool,'T_0') = s(bool,X107) | s(bool,X106) = s(bool,X107)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f26361,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,X105)) | sP1('const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12)),X106) | s(bool,X106) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) ) | $spl1),
  inference(subsumption_resolution,[],[f26340,f216])).
fof(f26340,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,X105)) | sP1('const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12)),X106) | s(bool,X106) = s(bool,X107) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,X105)) ) | $spl1),
  inference(superposition,[],[f2473,f14729])).
fof(f14729,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f14715,f201])).
fof(f14715,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f14148,f553])).
fof(f26360,plain,(
  ( ! [X35,X36] : (sP1('F_0',X35) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36)) ) | $spl86),
  inference(cnf_transformation,[],[f26360_D])).
fof(f26360_D,plain,(
  ( ! [X35,X36] : (sP1('F_0',X35) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f26359,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | sP1('F_0',X35) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f26323,f216])).
fof(f26323,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | sP1('F_0',X35) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f2473,f208])).
fof(f14147,plain,(
  $false | ($spl1 | $spl82)),
  inference(subsumption_resolution,[],[f348,f11341])).
fof(f14102,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f14100,f353])).
fof(f14100,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl5)),
  inference(backward_demodulation,[],[f14099,f11880])).
fof(f11880,plain,(
  p(s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13),s(sK11,sK12),s(sK10,sK16(sK12,'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13)),sK10,sK11))))) | $spl0),
  inference(forward_demodulation,[],[f11606,f1343])).
fof(f11606,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK16(sK12,'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13)),sK10,sK11)),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f11348,f356])).
fof(f11348,plain,(
  p(s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))),s(sK11,sK12)))) | $spl0),
  inference(forward_demodulation,[],[f347,f1157])).
fof(f14099,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(X0,bool)),X1),s(fun(sK11,bool),sK13),s(sK11,sK12),s(X0,X2)))) ) | $spl5),
  inference(forward_demodulation,[],[f14037,f13457])).
fof(f13457,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f396,f352,f605])).
fof(f14037,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13))) = s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(X0,bool)),X1),s(fun(sK11,bool),sK13),s(sK11,sK12),s(X0,X2)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f13318,f13317,f605])).
fof(f13317,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(X0,bool)),X1),s(fun(sK11,bool),sK13),s(sK11,sK12),s(X0,X2))))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f281])).
fof(f281,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f13318,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f216])).
fof(f13275,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f13273,f353])).
fof(f13273,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f13272,f11880])).
fof(f13272,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),X0),s(sK11,sK12),s(sK10,X1)))) ) | $spl3),
  inference(forward_demodulation,[],[f13187,f11509])).
fof(f11509,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))) | $spl3),
  inference(unit_resulting_resolution,[],[f396,f11347,f605])).
fof(f11347,plain,(
  ~p(s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12)))) | $spl3),
  inference(forward_demodulation,[],[f350,f1157])).
fof(f13187,plain,(
  ( ! [X0,X1] : (s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))) = s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),X0),s(sK11,sK12),s(sK10,X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f11350,f12122,f605])).
fof(f12122,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),X0),s(sK11,sK12),s(sK10,X1))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f11349,f358])).
fof(f358,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) | p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))))))) )),
  inference(forward_demodulation,[],[f280,f243])).
fof(f280,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f11349,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f11347,f355])).
fof(f11350,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))) | $spl3),
  inference(unit_resulting_resolution,[],[f11347,f216])).
fof(f11343,plain,(
  ( ! [X41,X40] : (sP1(X40,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))) | s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,X41)) ) | $spl84),
  inference(cnf_transformation,[],[f11343_D])).
fof(f11343_D,plain,(
  ( ! [X41,X40] : (sP1(X40,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))) | s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,X41)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f11342,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | sP1(X40,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl1),
  inference(subsumption_resolution,[],[f11317,f216])).
fof(f11317,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | sP1(X40,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X39)) ) | $spl1),
  inference(superposition,[],[f2472,f1754])).
fof(f1754,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl1),
  inference(subsumption_resolution,[],[f1749,f353])).
fof(f1749,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl1),
  inference(superposition,[],[f605,f675])).
fof(f675,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f353,f348,f244])).
fof(f2472,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2104,f201])).
fof(f2104,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f267,f553])).
fof(f267,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f163])).
fof(f11340,plain,(
  ( ! [X31,X32] : (sP1(X31,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) | $spl80),
  inference(cnf_transformation,[],[f11340_D])).
fof(f11340_D,plain,(
  ( ! [X31,X32] : (sP1(X31,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f11339,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1(X31,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f11313,f216])).
fof(f11313,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1(X31,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f2472,f208])).
fof(f9022,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP7(X7,'T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6))) | p(s(bool,X7))) ) | $spl78),
  inference(cnf_transformation,[],[f9022_D])).
fof(f9022_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP7(X7,'T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6))) | p(s(bool,X7))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2430,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f2430_D])).
fof(f2430_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f9024,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP7(X7,'T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)))) )),
  inference(condensation,[],[f9023])).
fof(f9023,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP7(X60,'T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59))) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f9004,f605])).
fof(f9004,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP7(X60,'T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59))) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1363,f553])).
fof(f1363,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (~p(s(bool,X131)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X126,fun(X125,bool)),X127),s(fun(X126,bool),X128),s(X126,X129),s(X125,X130)))) | p(s(bool,X132)) | ~sP7(X132,X131,happ(s(fun(X125,bool),happ(s(fun(X126,fun(X125,bool)),'const.relation.RRESTRICT_2'(s(fun(X126,fun(X125,bool)),X127),s(fun(X126,bool),X128))),s(X126,X129))),s(X125,X130)))) )),
  inference(superposition,[],[f322,f279])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f92,f129,f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f129,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f92,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X16,X17,X18] : ((p(s(bool,X18)) <=> (p(s(bool,X17)) & p(s(bool,X16)))) <=> ((p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X16))) & (p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X16)) | ~p(s(bool,X18)))))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.sat.sat_23_prove_ltb')).
fof(f9021,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP7(X7,'T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)))) )),
  inference(condensation,[],[f9020])).
fof(f9020,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP7(X51,'T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50))) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f9003,f605])).
fof(f9003,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP7(X51,'T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50))) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1363,f553])).
fof(f8980,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP6('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) | $spl76),
  inference(cnf_transformation,[],[f8980_D])).
fof(f8980_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP6('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f8982,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP6('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f8981])).
fof(f8981,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP6('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f8962,f605])).
fof(f8962,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP6('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1362,f553])).
fof(f1362,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119,X117] : (~p(s(bool,X123)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X118,fun(X117,bool)),X119),s(fun(X118,bool),X120),s(X118,X121),s(X117,X122)))) | p(s(bool,X124)) | ~sP6(X123,happ(s(fun(X117,bool),happ(s(fun(X118,fun(X117,bool)),'const.relation.RRESTRICT_2'(s(fun(X118,fun(X117,bool)),X119),s(fun(X118,bool),X120))),s(X118,X121))),s(X117,X122)),X124)) )),
  inference(superposition,[],[f313,f279])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f128])).
fof(f8979,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP6('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f8978])).
fof(f8978,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP6('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f8961,f605])).
fof(f8961,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP6('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1362,f553])).
fof(f8939,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP3('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) | $spl74),
  inference(cnf_transformation,[],[f8939_D])).
fof(f8939_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP3('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8941,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP3('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f8940])).
fof(f8940,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP3('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f8921,f605])).
fof(f8921,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP3('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1361,f553])).
fof(f1361,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (~p(s(bool,X115)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X110,fun(X109,bool)),X111),s(fun(X110,bool),X112),s(X110,X113),s(X109,X114)))) | p(s(bool,X116)) | ~sP3(X115,happ(s(fun(X109,bool),happ(s(fun(X110,fun(X109,bool)),'const.relation.RRESTRICT_2'(s(fun(X110,fun(X109,bool)),X111),s(fun(X110,bool),X112))),s(X110,X113))),s(X109,X114)),X116)) )),
  inference(superposition,[],[f294,f279])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f8938,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP3('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f8937])).
fof(f8937,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP3('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f8920,f605])).
fof(f8920,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP3('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1361,f553])).
fof(f8898,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP4('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) | $spl72),
  inference(cnf_transformation,[],[f8898_D])).
fof(f8898_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP4('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f8900,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP4('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f8899])).
fof(f8899,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP4('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f8880,f605])).
fof(f8880,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP4('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1360,f553])).
fof(f1360,plain,(
  ( ! [X103,X101,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X102,fun(X101,bool)),X103),s(fun(X102,bool),X104),s(X102,X105),s(X101,X106)))) | p(s(bool,X108)) | ~sP4(X107,happ(s(fun(X101,bool),happ(s(fun(X102,fun(X101,bool)),'const.relation.RRESTRICT_2'(s(fun(X102,fun(X101,bool)),X103),s(fun(X102,bool),X104))),s(X102,X105))),s(X101,X106)),X108)) )),
  inference(superposition,[],[f290,f279])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f175])).
fof(f175,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f8897,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP4('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f8896])).
fof(f8896,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP4('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f8879,f605])).
fof(f8879,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP4('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1360,f553])).
fof(f8859,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP5('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) | $spl70),
  inference(cnf_transformation,[],[f8859_D])).
fof(f8859_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | ~sP5('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f8861,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP5('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f8860])).
fof(f8860,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP5('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f8841,f605])).
fof(f8841,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP5('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1358,f553])).
fof(f1358,plain,(
  ( ! [X92,X90,X88,X87,X85,X91,X89,X86] : (~p(s(bool,X91)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X86,fun(X85,bool)),X87),s(fun(X86,bool),X88),s(X86,X89),s(X85,X90)))) | p(s(bool,X92)) | ~sP5(X91,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RRESTRICT_2'(s(fun(X86,fun(X85,bool)),X87),s(fun(X86,bool),X88))),s(X86,X89))),s(X85,X90)),X92)) )),
  inference(superposition,[],[f286,f279])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f172])).
fof(f172,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f8858,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP5('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f8857])).
fof(f8857,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP5('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f8840,f605])).
fof(f8840,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP5('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1358,f553])).
fof(f7875,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),'T_0',X7)) ) | $spl68),
  inference(cnf_transformation,[],[f7875_D])).
fof(f7875_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),'T_0',X7)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f7877,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),'T_0',X7)) )),
  inference(condensation,[],[f7876])).
fof(f7876,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | sP9(happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),'T_0',X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f7858,f605])).
fof(f7858,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | sP9(happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),'T_0',X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1366,f553])).
fof(f1366,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (~p(s(bool,X155)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X150,fun(X149,bool)),X151),s(fun(X150,bool),X152),s(X150,X153),s(X149,X154)))) | sP9(happ(s(fun(X149,bool),happ(s(fun(X150,fun(X149,bool)),'const.relation.RRESTRICT_2'(s(fun(X150,fun(X149,bool)),X151),s(fun(X150,bool),X152))),s(X150,X153))),s(X149,X154)),X155,X156)) )),
  inference(superposition,[],[f336,f279])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f94,f132,f131])).
fof(f131,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f132,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f94,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X16,X17,X18] : ((p(s(bool,X18)) <=> (p(s(bool,X17)) | p(s(bool,X16)))) <=> ((p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X16))) & (p(s(bool,X17)) | p(s(bool,X16)) | ~p(s(bool,X18)))))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.sat.sat_25_prove_ltb')).
fof(f7874,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),'T_0',X7)) )),
  inference(condensation,[],[f7873])).
fof(f7873,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | sP9(happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),'T_0',X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f7857,f605])).
fof(f7857,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | sP9(happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),'T_0',X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1366,f553])).
fof(f7839,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7,'T_0')) ) | $spl66),
  inference(cnf_transformation,[],[f7839_D])).
fof(f7839_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7,'T_0')) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7841,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7,'T_0')) )),
  inference(condensation,[],[f7840])).
fof(f7840,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | sP9(happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60,'T_0') | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f7822,f605])).
fof(f7822,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57),s(X54,X58),s(X55,X59)))) | sP9(happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RRESTRICT_2'(s(fun(X54,fun(X55,bool)),X56),s(fun(X54,bool),X57))),s(X54,X58))),s(X55,X59)),X60,'T_0') | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1365,f553])).
fof(f1365,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (~p(s(bool,X147)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X142,fun(X141,bool)),X143),s(fun(X142,bool),X144),s(X142,X145),s(X141,X146)))) | sP9(happ(s(fun(X141,bool),happ(s(fun(X142,fun(X141,bool)),'const.relation.RRESTRICT_2'(s(fun(X142,fun(X141,bool)),X143),s(fun(X142,bool),X144))),s(X142,X145))),s(X141,X146)),X148,X147)) )),
  inference(superposition,[],[f335,f279])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f7838,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X2,bool)),X3),s(fun(X1,bool),X4))),s(X1,X5))),s(X2,X6)),X7,'T_0')) )),
  inference(condensation,[],[f7837])).
fof(f7837,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | sP9(happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51,'T_0') | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f7821,f605])).
fof(f7821,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48),s(X45,X49),s(X46,X50)))) | sP9(happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RRESTRICT_2'(s(fun(X45,fun(X46,bool)),X47),s(fun(X45,bool),X48))),s(X45,X49))),s(X46,X50)),X51,'T_0') | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1365,f553])).
fof(f6140,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP7(X5,'T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4))) | p(s(bool,X5))) ) | $spl64),
  inference(cnf_transformation,[],[f6140_D])).
fof(f6140_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP7(X5,'T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4))) | p(s(bool,X5))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f6142,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP7(X5,'T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)))) )),
  inference(condensation,[],[f6141])).
fof(f6141,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP7(X46,'T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f6124,f605])).
fof(f6124,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP7(X46,'T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1177,f553])).
fof(f1177,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (~p(s(bool,X93)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X89,fun(X90,bool)),X91),s(X89,X92)))) | p(s(bool,X94)) | ~sP7(X94,X93,happ(s(fun(X89,bool),'const.relation.RDOM_1'(s(fun(X89,fun(X90,bool)),X91))),s(X89,X92)))) )),
  inference(superposition,[],[f322,f275])).
fof(f6139,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP7(X5,'T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)))) )),
  inference(condensation,[],[f6138])).
fof(f6138,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP7(X39,'T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f6123,f605])).
fof(f6123,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP7(X39,'T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1177,f553])).
fof(f5940,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP6('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl62),
  inference(cnf_transformation,[],[f5940_D])).
fof(f5940_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP6('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5942,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP6('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f5941])).
fof(f5941,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP6('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f5925,f605])).
fof(f5925,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP6('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1176,f553])).
fof(f1176,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (~p(s(bool,X87)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X83,fun(X84,bool)),X85),s(X83,X86)))) | p(s(bool,X88)) | ~sP6(X87,happ(s(fun(X83,bool),'const.relation.RDOM_1'(s(fun(X83,fun(X84,bool)),X85))),s(X83,X86)),X88)) )),
  inference(superposition,[],[f313,f275])).
fof(f5939,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP6('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f5938])).
fof(f5938,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP6('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f5924,f605])).
fof(f5924,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP6('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1176,f553])).
fof(f5887,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP3('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl60),
  inference(cnf_transformation,[],[f5887_D])).
fof(f5887_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP3('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f5889,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP3('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f5888])).
fof(f5888,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP3('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f5872,f605])).
fof(f5872,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP3('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1175,f553])).
fof(f1175,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (~p(s(bool,X81)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X77,fun(X78,bool)),X79),s(X77,X80)))) | p(s(bool,X82)) | ~sP3(X81,happ(s(fun(X77,bool),'const.relation.RDOM_1'(s(fun(X77,fun(X78,bool)),X79))),s(X77,X80)),X82)) )),
  inference(superposition,[],[f294,f275])).
fof(f5886,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP3('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f5885])).
fof(f5885,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP3('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f5871,f605])).
fof(f5871,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP3('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1175,f553])).
fof(f5834,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP4('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl58),
  inference(cnf_transformation,[],[f5834_D])).
fof(f5834_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP4('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5836,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP4('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f5835])).
fof(f5835,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP4('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f5819,f605])).
fof(f5819,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP4('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1174,f553])).
fof(f1174,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,X75)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X71,fun(X72,bool)),X73),s(X71,X74)))) | p(s(bool,X76)) | ~sP4(X75,happ(s(fun(X71,bool),'const.relation.RDOM_1'(s(fun(X71,fun(X72,bool)),X73))),s(X71,X74)),X76)) )),
  inference(superposition,[],[f290,f275])).
fof(f5833,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP4('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f5832])).
fof(f5832,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f5818,f605])).
fof(f5818,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1174,f553])).
fof(f5711,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool))))) ) | $spl54),
  inference(cnf_transformation,[],[f5711_D])).
fof(f5711_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5713,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl57),
  inference(cnf_transformation,[],[f5713_D])).
fof(f5713_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f5701,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK15(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1205,f553])).
fof(f1205,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK15(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK15(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f276,f208])).
fof(f276,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK15(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK15(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK15(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK15(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.relation.IN_RDOM_RRESTRICT',aHL_EXT)).
fof(f5700,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1205,f553])).
fof(f5227,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl48),
  inference(cnf_transformation,[],[f5227_D])).
fof(f5227_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5223,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool))) ) | $spl44),
  inference(cnf_transformation,[],[f5223_D])).
fof(f5223_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5230,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl53),
  inference(cnf_transformation,[],[f5230_D])).
fof(f5230_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f5221,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK15(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1204,f208])).
fof(f1204,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,X2,bool))) )),
  inference(superposition,[],[f276,f208])).
fof(f5228,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl50),
  inference(cnf_transformation,[],[f5228_D])).
fof(f5228_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f5220,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK15(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1204,f553])).
fof(f5219,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK15(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1204,f553])).
fof(f5222,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f5222_D])).
fof(f5222_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5225,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl47),
  inference(cnf_transformation,[],[f5225_D])).
fof(f5225_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f5218,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK15(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1204,f553])).
fof(f5217,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1204,f553])).
fof(f4934,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP7(X4,'T_0',happ(s(fun(X1,bool),X3),s(X1,X2))) | p(s(bool,X4))) ) | $spl40),
  inference(cnf_transformation,[],[f4934_D])).
fof(f4934_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP7(X4,'T_0',happ(s(fun(X1,bool),X3),s(X1,X2))) | p(s(bool,X4))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4936,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP7(X4,'T_0',happ(s(fun(X1,bool),X3),s(X1,X2)))) )),
  inference(condensation,[],[f4935])).
fof(f4935,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP7(X39,'T_0',happ(s(fun(X36,bool),X38),s(X36,X37))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f4922,f605])).
fof(f4922,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP7(X39,'T_0',happ(s(fun(X36,bool),X38),s(X36,X37))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1124,f553])).
fof(f1124,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,X71)) | ~p(s(bool,'const.bool.IN_2'(s(X68,X70),s(fun(X68,bool),X69)))) | p(s(bool,X72)) | ~sP7(X72,X71,happ(s(fun(X68,bool),X69),s(X68,X70)))) )),
  inference(superposition,[],[f322,f243])).
fof(f4933,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP7(X4,'T_0',happ(s(fun(X1,bool),X3),s(X1,X2)))) )),
  inference(condensation,[],[f4932])).
fof(f4932,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP7(X33,'T_0',happ(s(fun(X30,bool),X32),s(X30,X31))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4921,f605])).
fof(f4921,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP7(X33,'T_0',happ(s(fun(X30,bool),X32),s(X30,X31))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1124,f553])).
fof(f4900,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP6('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) | $spl38),
  inference(cnf_transformation,[],[f4900_D])).
fof(f4900_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP6('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4902,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP6('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f4901])).
fof(f4901,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP6('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f4888,f605])).
fof(f4888,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP6('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1123,f553])).
fof(f1123,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X66)) | ~p(s(bool,'const.bool.IN_2'(s(X63,X65),s(fun(X63,bool),X64)))) | p(s(bool,X67)) | ~sP6(X66,happ(s(fun(X63,bool),X64),s(X63,X65)),X67)) )),
  inference(superposition,[],[f313,f243])).
fof(f4899,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP6('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f4898])).
fof(f4898,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP6('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4887,f605])).
fof(f4887,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP6('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1123,f553])).
fof(f4865,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP3('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) | $spl36),
  inference(cnf_transformation,[],[f4865_D])).
fof(f4865_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP3('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4867,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP3('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f4866])).
fof(f4866,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP3('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f4853,f605])).
fof(f4853,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP3('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1122,f553])).
fof(f1122,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | ~p(s(bool,'const.bool.IN_2'(s(X58,X60),s(fun(X58,bool),X59)))) | p(s(bool,X62)) | ~sP3(X61,happ(s(fun(X58,bool),X59),s(X58,X60)),X62)) )),
  inference(superposition,[],[f294,f243])).
fof(f4864,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP3('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f4863])).
fof(f4863,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP3('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4852,f605])).
fof(f4852,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP3('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1122,f553])).
fof(f4833,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP4('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) | $spl34),
  inference(cnf_transformation,[],[f4833_D])).
fof(f4833_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP4('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4835,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP4('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f4834])).
fof(f4834,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f4821,f605])).
fof(f4821,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1121,f553])).
fof(f1121,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X56)) | ~p(s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),X54)))) | p(s(bool,X57)) | ~sP4(X56,happ(s(fun(X53,bool),X54),s(X53,X55)),X57)) )),
  inference(superposition,[],[f290,f243])).
fof(f4832,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP4('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f4831])).
fof(f4831,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP4('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4820,f605])).
fof(f4820,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP4('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1121,f553])).
fof(f2754,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737,f2739,f2741,f2743,f2745,f2747,f2749,f2751,f2753])).
fof(f2753,plain,(
  ( ! [X167,X165,X171,X169,X166,X172,X170,X168] : (~sP8(happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RRESTRICT_2'(s(fun(X166,fun(X165,bool)),X167),s(fun(X166,bool),X168))),s(X166,X169))),s(X165,X170)),X171,X172) | sP9(happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RRESTRICT_2'(s(fun(X166,fun(X165,bool)),X167),s(fun(X166,bool),X168))),s(X166,X169))),s(X165,X170)),X171,X172)) ) | $spl8),
  inference(subsumption_resolution,[],[f2631,f201])).
fof(f2631,plain,(
  ( ! [X167,X165,X171,X169,X166,X172,X170,X168] : (~p(s(bool,'T_0')) | ~sP8(happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RRESTRICT_2'(s(fun(X166,fun(X165,bool)),X167),s(fun(X166,bool),X168))),s(X166,X169))),s(X165,X170)),X171,X172) | sP9(happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RRESTRICT_2'(s(fun(X166,fun(X165,bool)),X167),s(fun(X166,bool),X168))),s(X166,X169))),s(X165,X170)),X171,X172)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1368])).
fof(f2751,plain,(
  ( ! [X158,X163,X161,X159,X157,X164,X162,X160] : (p(s(bool,X163)) | p(s(bool,X164)) | ~sP9(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RRESTRICT_2'(s(fun(X158,fun(X157,bool)),X159),s(fun(X158,bool),X160))),s(X158,X161))),s(X157,X162)),X164,X163)) ) | $spl8),
  inference(subsumption_resolution,[],[f2630,f201])).
fof(f2630,plain,(
  ( ! [X158,X163,X161,X159,X157,X164,X162,X160] : (~p(s(bool,'T_0')) | p(s(bool,X163)) | p(s(bool,X164)) | ~sP9(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RRESTRICT_2'(s(fun(X158,fun(X157,bool)),X159),s(fun(X158,bool),X160))),s(X158,X161))),s(X157,X162)),X164,X163)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1367])).
fof(f2749,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (~p(s(bool,X155)) | sP9(happ(s(fun(X149,bool),happ(s(fun(X150,fun(X149,bool)),'const.relation.RRESTRICT_2'(s(fun(X150,fun(X149,bool)),X151),s(fun(X150,bool),X152))),s(X150,X153))),s(X149,X154)),X155,X156)) ) | $spl8),
  inference(subsumption_resolution,[],[f2629,f201])).
fof(f2629,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (~p(s(bool,'T_0')) | ~p(s(bool,X155)) | sP9(happ(s(fun(X149,bool),happ(s(fun(X150,fun(X149,bool)),'const.relation.RRESTRICT_2'(s(fun(X150,fun(X149,bool)),X151),s(fun(X150,bool),X152))),s(X150,X153))),s(X149,X154)),X155,X156)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1366])).
fof(f2747,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (~p(s(bool,X147)) | sP9(happ(s(fun(X141,bool),happ(s(fun(X142,fun(X141,bool)),'const.relation.RRESTRICT_2'(s(fun(X142,fun(X141,bool)),X143),s(fun(X142,bool),X144))),s(X142,X145))),s(X141,X146)),X148,X147)) ) | $spl8),
  inference(subsumption_resolution,[],[f2628,f201])).
fof(f2628,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (~p(s(bool,'T_0')) | ~p(s(bool,X147)) | sP9(happ(s(fun(X141,bool),happ(s(fun(X142,fun(X141,bool)),'const.relation.RRESTRICT_2'(s(fun(X142,fun(X141,bool)),X143),s(fun(X142,bool),X144))),s(X142,X145))),s(X141,X146)),X148,X147)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1365])).
fof(f2745,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (p(s(bool,X139)) | p(s(bool,X140)) | ~sP8(happ(s(fun(X133,bool),happ(s(fun(X134,fun(X133,bool)),'const.relation.RRESTRICT_2'(s(fun(X134,fun(X133,bool)),X135),s(fun(X134,bool),X136))),s(X134,X137))),s(X133,X138)),X139,X140)) ) | $spl8),
  inference(subsumption_resolution,[],[f2627,f201])).
fof(f2627,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (~p(s(bool,'T_0')) | p(s(bool,X139)) | p(s(bool,X140)) | ~sP8(happ(s(fun(X133,bool),happ(s(fun(X134,fun(X133,bool)),'const.relation.RRESTRICT_2'(s(fun(X134,fun(X133,bool)),X135),s(fun(X134,bool),X136))),s(X134,X137))),s(X133,X138)),X139,X140)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1364])).
fof(f2743,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (~p(s(bool,X131)) | p(s(bool,X132)) | ~sP7(X132,X131,happ(s(fun(X125,bool),happ(s(fun(X126,fun(X125,bool)),'const.relation.RRESTRICT_2'(s(fun(X126,fun(X125,bool)),X127),s(fun(X126,bool),X128))),s(X126,X129))),s(X125,X130)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2626,f201])).
fof(f2626,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X131)) | p(s(bool,X132)) | ~sP7(X132,X131,happ(s(fun(X125,bool),happ(s(fun(X126,fun(X125,bool)),'const.relation.RRESTRICT_2'(s(fun(X126,fun(X125,bool)),X127),s(fun(X126,bool),X128))),s(X126,X129))),s(X125,X130)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1363])).
fof(f2741,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119,X117] : (~p(s(bool,X123)) | p(s(bool,X124)) | ~sP6(X123,happ(s(fun(X117,bool),happ(s(fun(X118,fun(X117,bool)),'const.relation.RRESTRICT_2'(s(fun(X118,fun(X117,bool)),X119),s(fun(X118,bool),X120))),s(X118,X121))),s(X117,X122)),X124)) ) | $spl8),
  inference(subsumption_resolution,[],[f2625,f201])).
fof(f2625,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,X123)) | p(s(bool,X124)) | ~sP6(X123,happ(s(fun(X117,bool),happ(s(fun(X118,fun(X117,bool)),'const.relation.RRESTRICT_2'(s(fun(X118,fun(X117,bool)),X119),s(fun(X118,bool),X120))),s(X118,X121))),s(X117,X122)),X124)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1362])).
fof(f2739,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (~p(s(bool,X115)) | p(s(bool,X116)) | ~sP3(X115,happ(s(fun(X109,bool),happ(s(fun(X110,fun(X109,bool)),'const.relation.RRESTRICT_2'(s(fun(X110,fun(X109,bool)),X111),s(fun(X110,bool),X112))),s(X110,X113))),s(X109,X114)),X116)) ) | $spl8),
  inference(subsumption_resolution,[],[f2624,f201])).
fof(f2624,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X115)) | p(s(bool,X116)) | ~sP3(X115,happ(s(fun(X109,bool),happ(s(fun(X110,fun(X109,bool)),'const.relation.RRESTRICT_2'(s(fun(X110,fun(X109,bool)),X111),s(fun(X110,bool),X112))),s(X110,X113))),s(X109,X114)),X116)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1361])).
fof(f2737,plain,(
  ( ! [X103,X101,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | p(s(bool,X108)) | ~sP4(X107,happ(s(fun(X101,bool),happ(s(fun(X102,fun(X101,bool)),'const.relation.RRESTRICT_2'(s(fun(X102,fun(X101,bool)),X103),s(fun(X102,bool),X104))),s(X102,X105))),s(X101,X106)),X108)) ) | $spl8),
  inference(subsumption_resolution,[],[f2623,f201])).
fof(f2623,plain,(
  ( ! [X103,X101,X107,X105,X102,X108,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X107)) | p(s(bool,X108)) | ~sP4(X107,happ(s(fun(X101,bool),happ(s(fun(X102,fun(X101,bool)),'const.relation.RRESTRICT_2'(s(fun(X102,fun(X101,bool)),X103),s(fun(X102,bool),X104))),s(X102,X105))),s(X101,X106)),X108)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1360])).
fof(f2735,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (sP5(X99,X100,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98))) | ~sP4(X99,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X100) | ~sP3(X100,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X99) | ~sP2(happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X100,X99)) ) | $spl8),
  inference(subsumption_resolution,[],[f2622,f201])).
fof(f2622,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (~p(s(bool,'T_0')) | sP5(X99,X100,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98))) | ~sP4(X99,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X100) | ~sP3(X100,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X99) | ~sP2(happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X100,X99)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1359])).
fof(f2733,plain,(
  ( ! [X92,X90,X88,X87,X85,X91,X89,X86] : (~p(s(bool,X91)) | p(s(bool,X92)) | ~sP5(X91,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RRESTRICT_2'(s(fun(X86,fun(X85,bool)),X87),s(fun(X86,bool),X88))),s(X86,X89))),s(X85,X90)),X92)) ) | $spl8),
  inference(subsumption_resolution,[],[f2621,f201])).
fof(f2621,plain,(
  ( ! [X92,X90,X88,X87,X85,X91,X89,X86] : (~p(s(bool,'T_0')) | ~p(s(bool,X91)) | p(s(bool,X92)) | ~sP5(X91,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RRESTRICT_2'(s(fun(X86,fun(X85,bool)),X87),s(fun(X86,bool),X88))),s(X86,X89))),s(X85,X90)),X92)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1358])).
fof(f2731,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (p(s(bool,X84)) | sP1(X84,happ(s(fun(X78,bool),happ(s(fun(X79,fun(X78,bool)),'const.relation.RRESTRICT_2'(s(fun(X79,fun(X78,bool)),X80),s(fun(X79,bool),X81))),s(X79,X82))),s(X78,X83)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2620,f201])).
fof(f2620,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (~p(s(bool,'T_0')) | p(s(bool,X84)) | sP1(X84,happ(s(fun(X78,bool),happ(s(fun(X79,fun(X78,bool)),'const.relation.RRESTRICT_2'(s(fun(X79,fun(X78,bool)),X80),s(fun(X79,bool),X81))),s(X79,X82))),s(X78,X83)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1357])).
fof(f2729,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (p(s(bool,X77)) | sP1(happ(s(fun(X71,bool),happ(s(fun(X72,fun(X71,bool)),'const.relation.RRESTRICT_2'(s(fun(X72,fun(X71,bool)),X73),s(fun(X72,bool),X74))),s(X72,X75))),s(X71,X76)),X77)) ) | $spl8),
  inference(subsumption_resolution,[],[f2619,f201])).
fof(f2619,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X77)) | sP1(happ(s(fun(X71,bool),happ(s(fun(X72,fun(X71,bool)),'const.relation.RRESTRICT_2'(s(fun(X72,fun(X71,bool)),X73),s(fun(X72,bool),X74))),s(X72,X75))),s(X71,X76)),X77)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1356])).
fof(f2727,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (sP0(happ(s(fun(X64,bool),happ(s(fun(X65,fun(X64,bool)),'const.relation.RRESTRICT_2'(s(fun(X65,fun(X64,bool)),X66),s(fun(X65,bool),X67))),s(X65,X68))),s(X64,X69)),X70)) ) | $spl8),
  inference(subsumption_resolution,[],[f2618,f201])).
fof(f2618,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X64,bool),happ(s(fun(X65,fun(X64,bool)),'const.relation.RRESTRICT_2'(s(fun(X65,fun(X64,bool)),X66),s(fun(X65,bool),X67))),s(X65,X68))),s(X64,X69)),X70)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1355])).
fof(f2725,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (sP0(X63,happ(s(fun(X57,bool),happ(s(fun(X58,fun(X57,bool)),'const.relation.RRESTRICT_2'(s(fun(X58,fun(X57,bool)),X59),s(fun(X58,bool),X60))),s(X58,X61))),s(X57,X62)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2617,f201])).
fof(f2617,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'T_0')) | sP0(X63,happ(s(fun(X57,bool),happ(s(fun(X58,fun(X57,bool)),'const.relation.RRESTRICT_2'(s(fun(X58,fun(X57,bool)),X59),s(fun(X58,bool),X60))),s(X58,X61))),s(X57,X62)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1354])).
fof(f2723,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (~sP8(happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)),X123,X124) | sP9(happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)),X123,X124)) ) | $spl8),
  inference(subsumption_resolution,[],[f2615,f201])).
fof(f2615,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (~p(s(bool,'T_0')) | ~sP8(happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)),X123,X124) | sP9(happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)),X123,X124)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1182])).
fof(f2721,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (p(s(bool,X117)) | p(s(bool,X118)) | ~sP9(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X118,X117)) ) | $spl8),
  inference(subsumption_resolution,[],[f2614,f201])).
fof(f2614,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,X117)) | p(s(bool,X118)) | ~sP9(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X118,X117)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1181])).
fof(f2719,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (~p(s(bool,X111)) | sP9(happ(s(fun(X107,bool),'const.relation.RDOM_1'(s(fun(X107,fun(X108,bool)),X109))),s(X107,X110)),X111,X112)) ) | $spl8),
  inference(subsumption_resolution,[],[f2613,f201])).
fof(f2613,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (~p(s(bool,'T_0')) | ~p(s(bool,X111)) | sP9(happ(s(fun(X107,bool),'const.relation.RDOM_1'(s(fun(X107,fun(X108,bool)),X109))),s(X107,X110)),X111,X112)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1180])).
fof(f2717,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (~p(s(bool,X105)) | sP9(happ(s(fun(X101,bool),'const.relation.RDOM_1'(s(fun(X101,fun(X102,bool)),X103))),s(X101,X104)),X106,X105)) ) | $spl8),
  inference(subsumption_resolution,[],[f2612,f201])).
fof(f2612,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X105)) | sP9(happ(s(fun(X101,bool),'const.relation.RDOM_1'(s(fun(X101,fun(X102,bool)),X103))),s(X101,X104)),X106,X105)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1179])).
fof(f2715,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (p(s(bool,X99)) | p(s(bool,X100)) | ~sP8(happ(s(fun(X95,bool),'const.relation.RDOM_1'(s(fun(X95,fun(X96,bool)),X97))),s(X95,X98)),X99,X100)) ) | $spl8),
  inference(subsumption_resolution,[],[f2611,f201])).
fof(f2611,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X99)) | p(s(bool,X100)) | ~sP8(happ(s(fun(X95,bool),'const.relation.RDOM_1'(s(fun(X95,fun(X96,bool)),X97))),s(X95,X98)),X99,X100)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1178])).
fof(f2713,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (~p(s(bool,X93)) | p(s(bool,X94)) | ~sP7(X94,X93,happ(s(fun(X89,bool),'const.relation.RDOM_1'(s(fun(X89,fun(X90,bool)),X91))),s(X89,X92)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2610,f201])).
fof(f2610,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X93)) | p(s(bool,X94)) | ~sP7(X94,X93,happ(s(fun(X89,bool),'const.relation.RDOM_1'(s(fun(X89,fun(X90,bool)),X91))),s(X89,X92)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1177])).
fof(f2711,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (~p(s(bool,X87)) | p(s(bool,X88)) | ~sP6(X87,happ(s(fun(X83,bool),'const.relation.RDOM_1'(s(fun(X83,fun(X84,bool)),X85))),s(X83,X86)),X88)) ) | $spl8),
  inference(subsumption_resolution,[],[f2609,f201])).
fof(f2609,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X87)) | p(s(bool,X88)) | ~sP6(X87,happ(s(fun(X83,bool),'const.relation.RDOM_1'(s(fun(X83,fun(X84,bool)),X85))),s(X83,X86)),X88)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1176])).
fof(f2709,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~sP3(X81,happ(s(fun(X77,bool),'const.relation.RDOM_1'(s(fun(X77,fun(X78,bool)),X79))),s(X77,X80)),X82)) ) | $spl8),
  inference(subsumption_resolution,[],[f2608,f201])).
fof(f2608,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP3(X81,happ(s(fun(X77,bool),'const.relation.RDOM_1'(s(fun(X77,fun(X78,bool)),X79))),s(X77,X80)),X82)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1175])).
fof(f2707,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,X75)) | p(s(bool,X76)) | ~sP4(X75,happ(s(fun(X71,bool),'const.relation.RDOM_1'(s(fun(X71,fun(X72,bool)),X73))),s(X71,X74)),X76)) ) | $spl8),
  inference(subsumption_resolution,[],[f2607,f201])).
fof(f2607,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | p(s(bool,X76)) | ~sP4(X75,happ(s(fun(X71,bool),'const.relation.RDOM_1'(s(fun(X71,fun(X72,bool)),X73))),s(X71,X74)),X76)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1174])).
fof(f2705,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (sP5(X69,X70,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68))) | ~sP4(X69,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70) | ~sP3(X70,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X69) | ~sP2(happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70,X69)) ) | $spl8),
  inference(subsumption_resolution,[],[f2606,f201])).
fof(f2606,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | sP5(X69,X70,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68))) | ~sP4(X69,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70) | ~sP3(X70,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X69) | ~sP2(happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70,X69)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1173])).
fof(f2703,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | ~sP5(X63,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64)) ) | $spl8),
  inference(subsumption_resolution,[],[f2605,f201])).
fof(f2605,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | p(s(bool,X64)) | ~sP5(X63,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1172])).
fof(f2701,plain,(
  ( ! [X57,X54,X58,X56,X55] : (p(s(bool,X58)) | sP1(X58,happ(s(fun(X54,bool),'const.relation.RDOM_1'(s(fun(X54,fun(X55,bool)),X56))),s(X54,X57)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2604,f201])).
fof(f2604,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | sP1(X58,happ(s(fun(X54,bool),'const.relation.RDOM_1'(s(fun(X54,fun(X55,bool)),X56))),s(X54,X57)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1171])).
fof(f2699,plain,(
  ( ! [X52,X50,X53,X51,X49] : (p(s(bool,X53)) | sP1(happ(s(fun(X49,bool),'const.relation.RDOM_1'(s(fun(X49,fun(X50,bool)),X51))),s(X49,X52)),X53)) ) | $spl8),
  inference(subsumption_resolution,[],[f2603,f201])).
fof(f2603,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | sP1(happ(s(fun(X49,bool),'const.relation.RDOM_1'(s(fun(X49,fun(X50,bool)),X51))),s(X49,X52)),X53)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1170])).
fof(f2697,plain,(
  ( ! [X47,X45,X48,X46,X44] : (sP0(happ(s(fun(X44,bool),'const.relation.RDOM_1'(s(fun(X44,fun(X45,bool)),X46))),s(X44,X47)),X48)) ) | $spl8),
  inference(subsumption_resolution,[],[f2602,f201])).
fof(f2602,plain,(
  ( ! [X47,X45,X48,X46,X44] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X44,bool),'const.relation.RDOM_1'(s(fun(X44,fun(X45,bool)),X46))),s(X44,X47)),X48)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1169])).
fof(f2695,plain,(
  ( ! [X39,X43,X41,X42,X40] : (sP0(X43,happ(s(fun(X39,bool),'const.relation.RDOM_1'(s(fun(X39,fun(X40,bool)),X41))),s(X39,X42)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2601,f201])).
fof(f2601,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'T_0')) | sP0(X43,happ(s(fun(X39,bool),'const.relation.RDOM_1'(s(fun(X39,fun(X40,bool)),X41))),s(X39,X42)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1168])).
fof(f2686,plain,(
  ( ! [X94,X97,X95,X93,X96] : (~sP8(happ(s(fun(X93,bool),X94),s(X93,X95)),X96,X97) | sP9(happ(s(fun(X93,bool),X94),s(X93,X95)),X96,X97)) ) | $spl8),
  inference(subsumption_resolution,[],[f2562,f201])).
fof(f2562,plain,(
  ( ! [X94,X97,X95,X93,X96] : (~p(s(bool,'T_0')) | ~sP8(happ(s(fun(X93,bool),X94),s(X93,X95)),X96,X97) | sP9(happ(s(fun(X93,bool),X94),s(X93,X95)),X96,X97)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1129])).
fof(f2684,plain,(
  ( ! [X92,X90,X88,X91,X89] : (p(s(bool,X91)) | p(s(bool,X92)) | ~sP9(happ(s(fun(X88,bool),X89),s(X88,X90)),X92,X91)) ) | $spl8),
  inference(subsumption_resolution,[],[f2561,f201])).
fof(f2561,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | p(s(bool,X92)) | ~sP9(happ(s(fun(X88,bool),X89),s(X88,X90)),X92,X91)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1128])).
fof(f2682,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,X86)) | sP9(happ(s(fun(X83,bool),X84),s(X83,X85)),X86,X87)) ) | $spl8),
  inference(subsumption_resolution,[],[f2560,f201])).
fof(f2560,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X86)) | sP9(happ(s(fun(X83,bool),X84),s(X83,X85)),X86,X87)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1127])).
fof(f2680,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,X81)) | sP9(happ(s(fun(X78,bool),X79),s(X78,X80)),X82,X81)) ) | $spl8),
  inference(subsumption_resolution,[],[f2559,f201])).
fof(f2559,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | sP9(happ(s(fun(X78,bool),X79),s(X78,X80)),X82,X81)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1126])).
fof(f2678,plain,(
  ( ! [X76,X74,X77,X75,X73] : (p(s(bool,X76)) | p(s(bool,X77)) | ~sP8(happ(s(fun(X73,bool),X74),s(X73,X75)),X76,X77)) ) | $spl8),
  inference(subsumption_resolution,[],[f2558,f201])).
fof(f2558,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | p(s(bool,X77)) | ~sP8(happ(s(fun(X73,bool),X74),s(X73,X75)),X76,X77)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1125])).
fof(f2676,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~sP7(X72,X71,happ(s(fun(X68,bool),X69),s(X68,X70)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2557,f201])).
fof(f2557,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X71)) | p(s(bool,X72)) | ~sP7(X72,X71,happ(s(fun(X68,bool),X69),s(X68,X70)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1124])).
fof(f2674,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~sP6(X66,happ(s(fun(X63,bool),X64),s(X63,X65)),X67)) ) | $spl8),
  inference(subsumption_resolution,[],[f2556,f201])).
fof(f2556,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP6(X66,happ(s(fun(X63,bool),X64),s(X63,X65)),X67)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1123])).
fof(f2672,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP3(X61,happ(s(fun(X58,bool),X59),s(X58,X60)),X62)) ) | $spl8),
  inference(subsumption_resolution,[],[f2555,f201])).
fof(f2555,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP3(X61,happ(s(fun(X58,bool),X59),s(X58,X60)),X62)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1122])).
fof(f2670,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP4(X56,happ(s(fun(X53,bool),X54),s(X53,X55)),X57)) ) | $spl8),
  inference(subsumption_resolution,[],[f2554,f201])).
fof(f2554,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP4(X56,happ(s(fun(X53,bool),X54),s(X53,X55)),X57)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1121])).
fof(f2668,plain,(
  ( ! [X52,X50,X48,X51,X49] : (sP5(X51,X52,happ(s(fun(X48,bool),X49),s(X48,X50))) | ~sP4(X51,happ(s(fun(X48,bool),X49),s(X48,X50)),X52) | ~sP3(X52,happ(s(fun(X48,bool),X49),s(X48,X50)),X51) | ~sP2(happ(s(fun(X48,bool),X49),s(X48,X50)),X52,X51)) ) | $spl8),
  inference(subsumption_resolution,[],[f2553,f201])).
fof(f2553,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,'T_0')) | sP5(X51,X52,happ(s(fun(X48,bool),X49),s(X48,X50))) | ~sP4(X51,happ(s(fun(X48,bool),X49),s(X48,X50)),X52) | ~sP3(X52,happ(s(fun(X48,bool),X49),s(X48,X50)),X51) | ~sP2(happ(s(fun(X48,bool),X49),s(X48,X50)),X52,X51)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1120])).
fof(f2666,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,happ(s(fun(X43,bool),X44),s(X43,X45)),X47)) ) | $spl8),
  inference(subsumption_resolution,[],[f2552,f201])).
fof(f2552,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,happ(s(fun(X43,bool),X44),s(X43,X45)),X47)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1119])).
fof(f2664,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2551,f201])).
fof(f2551,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1118])).
fof(f2662,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),X36),s(X35,X37)),X38)) ) | $spl8),
  inference(subsumption_resolution,[],[f2550,f201])).
fof(f2550,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),X36),s(X35,X37)),X38)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1117])).
fof(f2660,plain,(
  ( ! [X33,X31,X34,X32] : (sP0(happ(s(fun(X31,bool),X32),s(X31,X33)),X34)) ) | $spl8),
  inference(subsumption_resolution,[],[f2549,f201])).
fof(f2549,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X31,bool),X32),s(X31,X33)),X34)) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1116])).
fof(f2658,plain,(
  ( ! [X30,X28,X29,X27] : (sP0(X30,happ(s(fun(X27,bool),X28),s(X27,X29)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2548,f201])).
fof(f2548,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | sP0(X30,happ(s(fun(X27,bool),X28),s(X27,X29)))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f1115])).
fof(f2652,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f2521,f2430])).
fof(f2521,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f2430,f220])).
fof(f2150,plain,(
  ( ! [X218,X220,X219,X217] : (s(X219,f(s(bool,'T_0'))) != s(X219,g(s(bool,'T_0'))) | s(fun(bool,X219),X217) = s(fun(bool,X219),X218) | s(bool,X220) = s(bool,sK15(X217,X218,X219,bool)) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f276,f553])).
fof(f2149,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,'T_0') != s(bool,f(s(X213,sK15(X214,X215,bool,X213)))) | s(fun(X213,bool),X214) = s(fun(X213,bool),X215) | s(bool,X216) = s(bool,g(s(X213,sK15(X214,X215,bool,X213)))) | s(bool,'T_0') = s(bool,X216)) )),
  inference(superposition,[],[f276,f553])).
fof(f2144,plain,(
  ( ! [X196,X194,X192,X197,X195,X193] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X193,fun(bool,bool)),X194),s(fun(X193,bool),X195),s(X193,X196),s(bool,'T_0')))) | p(s(bool,'const.bool.IN_2'(s(bool,X192),s(fun(bool,bool),happ(s(fun(X193,fun(bool,bool)),X194),s(X193,X196)))))) | s(bool,X192) = s(bool,X197) | s(bool,'T_0') = s(bool,X197)) )),
  inference(superposition,[],[f358,f553])).
fof(f2143,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X187,bool)),X188),s(fun(bool,bool),X189),s(bool,'T_0'),s(X187,X190)))) | p(s(bool,'const.bool.IN_2'(s(X187,X190),s(fun(X187,bool),happ(s(fun(bool,fun(X187,bool)),X188),s(bool,X186)))))) | s(bool,X186) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f358,f553])).
fof(f2142,plain,(
  ( ! [X185,X182,X180,X184,X183,X181] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X181)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X182,bool)),X183),s(fun(bool,bool),X181),s(bool,X180),s(X182,X184)))) | ~p(s(bool,'const.bool.IN_2'(s(X182,X184),s(fun(X182,bool),happ(s(fun(bool,fun(X182,bool)),X183),s(bool,X180)))))) | s(bool,X180) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f357,f553])).
fof(f2141,plain,(
  ( ! [X177,X178,X176,X179] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X177,bool)),X178),s(bool,'T_0')))) | p(s(bool,'const.bool.IN_2'(s(X177,sK16(X176,X178,X177,bool)),s(fun(X177,bool),happ(s(fun(bool,fun(X177,bool)),X178),s(bool,X176)))))) | s(bool,X176) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f356,f553])).
fof(f2139,plain,(
  ( ! [X167,X169,X166,X170,X168] : (~p(s(bool,'const.bool.IN_2'(s(X167,X168),s(fun(X167,bool),happ(s(fun(bool,fun(X167,bool)),X169),s(bool,'T_0')))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X167,bool)),X169),s(bool,X166)))) | s(bool,X166) = s(bool,X170) | s(bool,'T_0') = s(bool,X170)) )),
  inference(superposition,[],[f355,f553])).
fof(f2137,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X155,bool)),X156),s(fun(bool,bool),X157),s(bool,'T_0'),s(X155,X158)))) | p(s(bool,'const.bool.IN_2'(s(bool,X154),s(fun(bool,bool),X157)))) | s(bool,X154) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f281,f553])).
fof(f2489,plain,(
  ( ! [X134,X132,X130,X128,X133,X131,X129] : (p(s(bool,'const.bool.IN_2'(s(X128,X132),s(fun(X128,bool),X131)))) | s(bool,X134) = s(bool,'const.relation.RRESTRICT_4'(s(fun(X128,fun(X129,bool)),X130),s(fun(X128,bool),X131),s(X128,X132),s(X129,X133))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(subsumption_resolution,[],[f2132,f201])).
fof(f2132,plain,(
  ( ! [X134,X132,X130,X128,X133,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X128,X132),s(fun(X128,bool),X131)))) | s(bool,X134) = s(bool,'const.relation.RRESTRICT_4'(s(fun(X128,fun(X129,bool)),X130),s(fun(X128,bool),X131),s(X128,X132),s(X129,X133))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f281,f553])).
fof(f2488,plain,(
  ( ! [X127,X125,X123,X121,X126,X124,X122] : (p(s(bool,'const.bool.IN_2'(s(X122,X126),s(fun(X122,bool),happ(s(fun(X121,fun(X122,bool)),X123),s(X121,X125)))))) | s(bool,X127) = s(bool,'const.relation.RRESTRICT_4'(s(fun(X121,fun(X122,bool)),X123),s(fun(X121,bool),X124),s(X121,X125),s(X122,X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(subsumption_resolution,[],[f2131,f201])).
fof(f2131,plain,(
  ( ! [X127,X125,X123,X121,X126,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X122,X126),s(fun(X122,bool),happ(s(fun(X121,fun(X122,bool)),X123),s(X121,X125)))))) | s(bool,X127) = s(bool,'const.relation.RRESTRICT_4'(s(fun(X121,fun(X122,bool)),X123),s(fun(X121,bool),X124),s(X121,X125),s(X122,X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f358,f553])).
fof(f2487,plain,(
  ( ! [X118,X116,X120,X119,X117] : (p(s(bool,'const.bool.IN_2'(s(X117,sK16(X119,X118,X117,X116)),s(fun(X117,bool),happ(s(fun(X116,fun(X117,bool)),X118),s(X116,X119)))))) | s(bool,X120) = s(bool,'const.relation.RDOM_2'(s(fun(X116,fun(X117,bool)),X118),s(X116,X119))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f2130,f201])).
fof(f2130,plain,(
  ( ! [X118,X116,X120,X119,X117] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X117,sK16(X119,X118,X117,X116)),s(fun(X117,bool),happ(s(fun(X116,fun(X117,bool)),X118),s(X116,X119)))))) | s(bool,X120) = s(bool,'const.relation.RDOM_2'(s(fun(X116,fun(X117,bool)),X118),s(X116,X119))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f356,f553])).
fof(f2486,plain,(
  ( ! [X94,X99,X97,X95,X93,X98,X96] : (p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X93,fun(X96,bool)),X97),s(fun(X93,bool),X95),s(X93,X94),s(X96,X98)))) | ~p(s(bool,'const.bool.IN_2'(s(X96,X98),s(fun(X96,bool),happ(s(fun(X93,fun(X96,bool)),X97),s(X93,X94)))))) | s(bool,X99) = s(bool,'const.bool.IN_2'(s(X93,X94),s(fun(X93,bool),X95))) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f2126,f201])).
fof(f2126,plain,(
  ( ! [X94,X99,X97,X95,X93,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X93,fun(X96,bool)),X97),s(fun(X93,bool),X95),s(X93,X94),s(X96,X98)))) | ~p(s(bool,'const.bool.IN_2'(s(X96,X98),s(fun(X96,bool),happ(s(fun(X93,fun(X96,bool)),X97),s(X93,X94)))))) | s(bool,X99) = s(bool,'const.bool.IN_2'(s(X93,X94),s(fun(X93,bool),X95))) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f357,f553])).
fof(f2485,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X89,fun(X87,bool)),X90),s(X89,X91)))) | s(bool,X92) = s(bool,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),happ(s(fun(X89,fun(X87,bool)),X90),s(X89,X91))))) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f2125,f201])).
fof(f2125,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RDOM_2'(s(fun(X89,fun(X87,bool)),X90),s(X89,X91)))) | s(bool,X92) = s(bool,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),happ(s(fun(X89,fun(X87,bool)),X90),s(X89,X91))))) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f355,f553])).
fof(f2484,plain,(
  ( ! [X74,X72,X71,X73] : (~sP8(X71,X72,X73) | sP9(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2116,f201])).
fof(f2116,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~sP8(X71,X72,X73) | sP9(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f340,f553])).
fof(f2483,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2115,f201])).
fof(f2115,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f337,f553])).
fof(f2482,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2114,f201])).
fof(f2114,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f336,f553])).
fof(f2481,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2113,f201])).
fof(f2113,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f335,f553])).
fof(f2480,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2112,f201])).
fof(f2112,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f330,f553])).
fof(f2479,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2111,f201])).
fof(f2111,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f322,f553])).
fof(f2478,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2110,f201])).
fof(f2110,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f313,f553])).
fof(f2477,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2109,f201])).
fof(f2109,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f294,f553])).
fof(f2476,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f2108,f201])).
fof(f2108,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f290,f553])).
fof(f2475,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2107,f201])).
fof(f2107,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f287,f553])).
fof(f2474,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2106,f201])).
fof(f2106,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f286,f553])).
fof(f2471,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f2103,f201])).
fof(f2103,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f252,f553])).
fof(f2470,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2102,f201])).
fof(f2102,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f251,f553])).
fof(f2469,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2101,f201])).
fof(f2101,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f247,f553])).
fof(f2468,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2097,f201])).
fof(f2097,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f220,f553])).
fof(f2089,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X141,fun(bool,bool)),X142),s(fun(X141,bool),X143),s(X141,X144),s(bool,X140)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X141,fun(bool,bool)),X142),s(X141,X144)))))) | s(bool,X140) = s(bool,X145) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f358,f553])).
fof(f2088,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X135,bool)),X136),s(fun(bool,bool),X137),s(bool,X134),s(X135,X138)))) | p(s(bool,'const.bool.IN_2'(s(X135,X138),s(fun(X135,bool),happ(s(fun(bool,fun(X135,bool)),X136),s(bool,'T_0')))))) | s(bool,X134) = s(bool,X139) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f358,f553])).
fof(f2087,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (~p(s(bool,'const.bool.IN_2'(s(bool,X128),s(fun(bool,bool),X129)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X130,bool)),X131),s(fun(bool,bool),X129),s(bool,'T_0'),s(X130,X132)))) | ~p(s(bool,'const.bool.IN_2'(s(X130,X132),s(fun(X130,bool),happ(s(fun(bool,fun(X130,bool)),X131),s(bool,'T_0')))))) | s(bool,X128) = s(bool,X133) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f357,f553])).
fof(f2086,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X125,bool)),X126),s(bool,X124)))) | p(s(bool,'const.bool.IN_2'(s(X125,sK16('T_0',X126,X125,bool)),s(fun(X125,bool),happ(s(fun(bool,fun(X125,bool)),X126),s(bool,'T_0')))))) | s(bool,X124) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f356,f553])).
fof(f2084,plain,(
  ( ! [X118,X116,X114,X117,X115] : (~p(s(bool,'const.bool.IN_2'(s(X115,X116),s(fun(X115,bool),happ(s(fun(bool,fun(X115,bool)),X117),s(bool,X114)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X115,bool)),X117),s(bool,'T_0')))) | s(bool,X114) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f355,f553])).
fof(f2082,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X103,bool)),X104),s(fun(bool,bool),X105),s(bool,X102),s(X103,X106)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X105)))) | s(bool,X102) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f281,f553])).
fof(f2053,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f216,f553])).
fof(f1993,plain,(
  ( ! [X222,X225,X223,X224] : (s(X224,f(s(bool,X225))) != s(X224,g(s(bool,X225))) | s(fun(bool,X224),X222) = s(fun(bool,X224),X223) | s(bool,'T_0') = s(bool,X225) | s(bool,'T_0') = s(bool,sK15(X222,X223,X224,bool))) )),
  inference(superposition,[],[f276,f553])).
fof(f1992,plain,(
  ( ! [X218,X220,X219,X221] : (s(bool,X221) != s(bool,f(s(X218,sK15(X219,X220,bool,X218)))) | s(fun(X218,bool),X219) = s(fun(X218,bool),X220) | s(bool,'T_0') = s(bool,X221) | s(bool,'T_0') = s(bool,g(s(X218,sK15(X219,X220,bool,X218))))) )),
  inference(superposition,[],[f276,f553])).
fof(f1987,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X198,fun(bool,bool)),X199),s(fun(X198,bool),X200),s(X198,X201),s(bool,X197)))) | p(s(bool,'const.bool.IN_2'(s(bool,X196),s(fun(bool,bool),happ(s(fun(X198,fun(bool,bool)),X199),s(X198,X201)))))) | s(bool,'T_0') = s(bool,X197) | s(bool,'T_0') = s(bool,X196)) )),
  inference(superposition,[],[f358,f553])).
fof(f1986,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X192,bool)),X193),s(fun(bool,bool),X194),s(bool,X191),s(X192,X195)))) | p(s(bool,'const.bool.IN_2'(s(X192,X195),s(fun(X192,bool),happ(s(fun(bool,fun(X192,bool)),X193),s(bool,X190)))))) | s(bool,'T_0') = s(bool,X191) | s(bool,'T_0') = s(bool,X190)) )),
  inference(superposition,[],[f358,f553])).
fof(f1985,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (~p(s(bool,'const.bool.IN_2'(s(bool,X185),s(fun(bool,bool),X186)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X187,bool)),X188),s(fun(bool,bool),X186),s(bool,X184),s(X187,X189)))) | ~p(s(bool,'const.bool.IN_2'(s(X187,X189),s(fun(X187,bool),happ(s(fun(bool,fun(X187,bool)),X188),s(bool,X184)))))) | s(bool,'T_0') = s(bool,X185) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f357,f553])).
fof(f1984,plain,(
  ( ! [X182,X180,X183,X181] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X182,bool)),X183),s(bool,X181)))) | p(s(bool,'const.bool.IN_2'(s(X182,sK16(X180,X183,X182,bool)),s(fun(X182,bool),happ(s(fun(bool,fun(X182,bool)),X183),s(bool,X180)))))) | s(bool,'T_0') = s(bool,X181) | s(bool,'T_0') = s(bool,X180)) )),
  inference(superposition,[],[f356,f553])).
fof(f1982,plain,(
  ( ! [X173,X171,X174,X172,X170] : (~p(s(bool,'const.bool.IN_2'(s(X172,X173),s(fun(X172,bool),happ(s(fun(bool,fun(X172,bool)),X174),s(bool,X171)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X172,bool)),X174),s(bool,X170)))) | s(bool,'T_0') = s(bool,X171) | s(bool,'T_0') = s(bool,X170)) )),
  inference(superposition,[],[f355,f553])).
fof(f1980,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X160,bool)),X161),s(fun(bool,bool),X162),s(bool,X159),s(X160,X163)))) | p(s(bool,'const.bool.IN_2'(s(bool,X158),s(fun(bool,bool),X162)))) | s(bool,'T_0') = s(bool,X159) | s(bool,'T_0') = s(bool,X158)) )),
  inference(superposition,[],[f281,f553])).
fof(f1920,plain,(
  ( ! [X30,X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') = s(bool,'const.relation.RRESTRICT_4'(s(fun(X25,fun(X24,bool)),X26),s(fun(X25,bool),X27),s(X25,X28),s(X24,X29))) | s(bool,X30) = s(bool,happ(s(fun(X24,bool),happ(s(fun(X25,fun(X24,bool)),'const.relation.RRESTRICT_2'(s(fun(X25,fun(X24,bool)),X26),s(fun(X25,bool),X27))),s(X25,X28))),s(X24,X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f553,f279])).
fof(f1919,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.relation.RDOM_2'(s(fun(X19,fun(X20,bool)),X21),s(X19,X22))) | s(bool,X23) = s(bool,happ(s(fun(X19,bool),'const.relation.RDOM_1'(s(fun(X19,fun(X20,bool)),X21))),s(X19,X22))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f553,f275])).
fof(f1918,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X15,X17),s(fun(X15,bool),X16))) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f553,f243])).
fof(f1907,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f553,f553])).
fof(f1892,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,X38) = s(bool,'const.relation.RRESTRICT_4'(s(fun(X33,fun(X32,bool)),X34),s(fun(X33,bool),X35),s(X33,X36),s(X32,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,happ(s(fun(X32,bool),happ(s(fun(X33,fun(X32,bool)),'const.relation.RRESTRICT_2'(s(fun(X33,fun(X32,bool)),X34),s(fun(X33,bool),X35))),s(X33,X36))),s(X32,X37)))) )),
  inference(superposition,[],[f553,f279])).
fof(f1891,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X31) = s(bool,'const.relation.RDOM_2'(s(fun(X27,fun(X28,bool)),X29),s(X27,X30))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,happ(s(fun(X27,bool),'const.relation.RDOM_1'(s(fun(X27,fun(X28,bool)),X29))),s(X27,X30)))) )),
  inference(superposition,[],[f553,f275])).
fof(f1890,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,'const.bool.IN_2'(s(X23,X25),s(fun(X23,bool),X24))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25)))) )),
  inference(superposition,[],[f553,f243])).
fof(f1845,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',happ(s(fun(X0,bool),X1),s(X0,X2)),'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f1100,f285])).
fof(f1796,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)),happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f1098,f285])).
fof(f1416,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X6)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X7,bool)),X8),s(fun(bool,bool),X6),s(bool,'F_0'),s(X7,X9)))) | ~p(s(bool,'const.bool.IN_2'(s(X7,X9),s(fun(X7,bool),happ(s(fun(bool,fun(X7,bool)),X8),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f357,f208])).
fof(f1415,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X2,bool)),X3),s(fun(bool,bool),X1),s(bool,X0),s(X2,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),happ(s(fun(bool,fun(X2,bool)),X3),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f357,f208])).
fof(f1397,plain,(
  ( ! [X231,X229,X227,X233,X230,X228,X226,X232] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X232,bool)),X233),s(bool,'const.relation.RRESTRICT_4'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229),s(X227,X230),s(X226,X231)))))) | p(s(bool,'const.bool.IN_2'(s(X232,sK16(happ(s(fun(X226,bool),happ(s(fun(X227,fun(X226,bool)),'const.relation.RRESTRICT_2'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229))),s(X227,X230))),s(X226,X231)),X233,X232,bool)),s(fun(X232,bool),happ(s(fun(bool,fun(X232,bool)),X233),s(bool,'const.relation.RRESTRICT_4'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229),s(X227,X230),s(X226,X231))))))))) )),
  inference(forward_demodulation,[],[f1396,f1343])).
fof(f1396,plain,(
  ( ! [X231,X229,X227,X233,X230,X228,X226,X232] : (p(s(bool,'const.bool.IN_2'(s(X232,sK16(happ(s(fun(X226,bool),happ(s(fun(X227,fun(X226,bool)),'const.relation.RRESTRICT_2'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229))),s(X227,X230))),s(X226,X231)),X233,X232,bool)),s(fun(X232,bool),happ(s(fun(bool,fun(X232,bool)),X233),s(bool,'const.bool.IN_2'(s(X226,X231),s(fun(X226,bool),happ(s(fun(X227,fun(X226,bool)),'const.relation.RRESTRICT_2'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229))),s(X227,X230)))))))))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X232,bool)),X233),s(bool,'const.relation.RRESTRICT_4'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229),s(X227,X230),s(X226,X231))))))) )),
  inference(forward_demodulation,[],[f1375,f243])).
fof(f1375,plain,(
  ( ! [X231,X229,X227,X233,X230,X228,X226,X232] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X232,bool)),X233),s(bool,'const.relation.RRESTRICT_4'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229),s(X227,X230),s(X226,X231)))))) | p(s(bool,'const.bool.IN_2'(s(X232,sK16(happ(s(fun(X226,bool),happ(s(fun(X227,fun(X226,bool)),'const.relation.RRESTRICT_2'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229))),s(X227,X230))),s(X226,X231)),X233,X232,bool)),s(fun(X232,bool),happ(s(fun(bool,fun(X232,bool)),X233),s(bool,happ(s(fun(X226,bool),happ(s(fun(X227,fun(X226,bool)),'const.relation.RRESTRICT_2'(s(fun(X227,fun(X226,bool)),X228),s(fun(X227,bool),X229))),s(X227,X230))),s(X226,X231))))))))) )),
  inference(superposition,[],[f356,f279])).
fof(f1368,plain,(
  ( ! [X167,X165,X171,X169,X166,X172,X170,X168] : (sP9(happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RRESTRICT_2'(s(fun(X166,fun(X165,bool)),X167),s(fun(X166,bool),X168))),s(X166,X169))),s(X165,X170)),X171,X172) | ~sP8(happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RRESTRICT_2'(s(fun(X166,fun(X165,bool)),X167),s(fun(X166,bool),X168))),s(X166,X169))),s(X165,X170)),X171,X172) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X166,fun(X165,bool)),X167),s(fun(X166,bool),X168),s(X166,X169),s(X165,X170))))) )),
  inference(superposition,[],[f340,f279])).
fof(f1367,plain,(
  ( ! [X158,X163,X161,X159,X157,X164,X162,X160] : (~sP9(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RRESTRICT_2'(s(fun(X158,fun(X157,bool)),X159),s(fun(X158,bool),X160))),s(X158,X161))),s(X157,X162)),X164,X163) | p(s(bool,X163)) | p(s(bool,X164)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X158,fun(X157,bool)),X159),s(fun(X158,bool),X160),s(X158,X161),s(X157,X162))))) )),
  inference(superposition,[],[f337,f279])).
fof(f1364,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X134,fun(X133,bool)),X135),s(fun(X134,bool),X136),s(X134,X137),s(X133,X138)))) | p(s(bool,X139)) | p(s(bool,X140)) | ~sP8(happ(s(fun(X133,bool),happ(s(fun(X134,fun(X133,bool)),'const.relation.RRESTRICT_2'(s(fun(X134,fun(X133,bool)),X135),s(fun(X134,bool),X136))),s(X134,X137))),s(X133,X138)),X139,X140)) )),
  inference(superposition,[],[f330,f279])).
fof(f1359,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (~sP3(X100,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X99) | sP5(X99,X100,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98))) | ~sP4(X99,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X100) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96),s(X94,X97),s(X93,X98)))) | ~sP2(happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RRESTRICT_2'(s(fun(X94,fun(X93,bool)),X95),s(fun(X94,bool),X96))),s(X94,X97))),s(X93,X98)),X100,X99)) )),
  inference(superposition,[],[f287,f279])).
fof(f1357,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (sP1(X84,happ(s(fun(X78,bool),happ(s(fun(X79,fun(X78,bool)),'const.relation.RRESTRICT_2'(s(fun(X79,fun(X78,bool)),X80),s(fun(X79,bool),X81))),s(X79,X82))),s(X78,X83))) | p(s(bool,X84)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X79,fun(X78,bool)),X80),s(fun(X79,bool),X81),s(X79,X82),s(X78,X83))))) )),
  inference(superposition,[],[f270,f279])).
fof(f1356,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (sP1(happ(s(fun(X71,bool),happ(s(fun(X72,fun(X71,bool)),'const.relation.RRESTRICT_2'(s(fun(X72,fun(X71,bool)),X73),s(fun(X72,bool),X74))),s(X72,X75))),s(X71,X76)),X77) | p(s(bool,X77)) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X72,fun(X71,bool)),X73),s(fun(X72,bool),X74),s(X72,X75),s(X71,X76))))) )),
  inference(superposition,[],[f267,f279])).
fof(f1355,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (sP0(happ(s(fun(X64,bool),happ(s(fun(X65,fun(X64,bool)),'const.relation.RRESTRICT_2'(s(fun(X65,fun(X64,bool)),X66),s(fun(X65,bool),X67))),s(X65,X68))),s(X64,X69)),X70) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X65,fun(X64,bool)),X66),s(fun(X65,bool),X67),s(X65,X68),s(X64,X69))))) )),
  inference(superposition,[],[f252,f279])).
fof(f1354,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (sP0(X63,happ(s(fun(X57,bool),happ(s(fun(X58,fun(X57,bool)),'const.relation.RRESTRICT_2'(s(fun(X58,fun(X57,bool)),X59),s(fun(X58,bool),X60))),s(X58,X61))),s(X57,X62))) | ~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X58,fun(X57,bool)),X59),s(fun(X58,bool),X60),s(X58,X61),s(X57,X62))))) )),
  inference(superposition,[],[f251,f279])).
fof(f1344,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),happ(s(fun(X13,fun(X12,bool)),'const.relation.RRESTRICT_2'(s(fun(X13,fun(X12,bool)),X14),s(fun(X13,bool),X15))),s(X13,X16))),s(X12,X17))) | s(bool,'F_0') = s(bool,'const.relation.RRESTRICT_4'(s(fun(X13,fun(X12,bool)),X14),s(fun(X13,bool),X15),s(X13,X16),s(X12,X17)))) )),
  inference(superposition,[],[f279,f208])).
fof(f1325,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3))),s(X1,X4))),s(X0,X5)),'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) )),
  inference(unit_resulting_resolution,[],[f353,f279,f305])).
fof(f1326,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3))),s(X1,X4))),s(X0,X5)),'const.relation.RRESTRICT_4'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(X0,X5)))) )),
  inference(unit_resulting_resolution,[],[f201,f279,f303])).
fof(f1327,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5('F_0','const.relation.RRESTRICT_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3),s(X0,X4),s(X1,X5)),happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RRESTRICT_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(unit_resulting_resolution,[],[f353,f279,f305])).
fof(f1328,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5('T_0','const.relation.RRESTRICT_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3),s(X0,X4),s(X1,X5)),happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RRESTRICT_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,bool),X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(unit_resulting_resolution,[],[f201,f279,f303])).
fof(f1301,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X6,fun(bool,bool)),X7),s(fun(X6,bool),X8),s(X6,X9),s(bool,X5)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),X7),s(X6,X9)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f358,f208])).
fof(f1300,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(X1,fun(bool,bool)),X2),s(fun(X1,bool),X3),s(X1,X4),s(bool,'F_0')))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),happ(s(fun(X1,fun(bool,bool)),X2),s(X1,X4)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f358,f208])).
fof(f1295,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X6,bool)),X7),s(fun(bool,bool),X8),s(bool,X5),s(X6,X9)))) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),happ(s(fun(bool,fun(X6,bool)),X7),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f358,f208])).
fof(f1294,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X1,bool)),X2),s(fun(bool,bool),X3),s(bool,'F_0'),s(X1,X4)))) | p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X2),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f358,f208])).
fof(f1281,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X19,bool)),X20),s(bool,'const.relation.RDOM_2'(s(fun(X15,fun(X16,bool)),X17),s(X15,X18)))))) | p(s(bool,'const.bool.IN_2'(s(X19,sK16(happ(s(fun(X15,bool),'const.relation.RDOM_1'(s(fun(X15,fun(X16,bool)),X17))),s(X15,X18)),X20,X19,bool)),s(fun(X19,bool),happ(s(fun(bool,fun(X19,bool)),X20),s(bool,'const.relation.RDOM_2'(s(fun(X15,fun(X16,bool)),X17),s(X15,X18))))))))) )),
  inference(forward_demodulation,[],[f1280,f1157])).
fof(f1280,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X19,sK16(happ(s(fun(X15,bool),'const.relation.RDOM_1'(s(fun(X15,fun(X16,bool)),X17))),s(X15,X18)),X20,X19,bool)),s(fun(X19,bool),happ(s(fun(bool,fun(X19,bool)),X20),s(bool,'const.bool.IN_2'(s(X15,X18),s(fun(X15,bool),'const.relation.RDOM_1'(s(fun(X15,fun(X16,bool)),X17)))))))))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X19,bool)),X20),s(bool,'const.relation.RDOM_2'(s(fun(X15,fun(X16,bool)),X17),s(X15,X18))))))) )),
  inference(forward_demodulation,[],[f1275,f243])).
fof(f1275,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X19,bool)),X20),s(bool,'const.relation.RDOM_2'(s(fun(X15,fun(X16,bool)),X17),s(X15,X18)))))) | p(s(bool,'const.bool.IN_2'(s(X19,sK16(happ(s(fun(X15,bool),'const.relation.RDOM_1'(s(fun(X15,fun(X16,bool)),X17))),s(X15,X18)),X20,X19,bool)),s(fun(X19,bool),happ(s(fun(bool,fun(X19,bool)),X20),s(bool,happ(s(fun(X15,bool),'const.relation.RDOM_1'(s(fun(X15,fun(X16,bool)),X17))),s(X15,X18))))))))) )),
  inference(superposition,[],[f356,f275])).
fof(f1279,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X13,bool)),X14),s(bool,'const.bool.IN_2'(s(X10,X12),s(fun(X10,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(X13,sK16(happ(s(fun(X10,bool),X11),s(X10,X12)),X14,X13,bool)),s(fun(X13,bool),happ(s(fun(bool,fun(X13,bool)),X14),s(bool,'const.bool.IN_2'(s(X10,X12),s(fun(X10,bool),X11))))))))) )),
  inference(forward_demodulation,[],[f1274,f243])).
fof(f1274,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X13,bool)),X14),s(bool,'const.bool.IN_2'(s(X10,X12),s(fun(X10,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(X13,sK16(happ(s(fun(X10,bool),X11),s(X10,X12)),X14,X13,bool)),s(fun(X13,bool),happ(s(fun(bool,fun(X13,bool)),X14),s(bool,happ(s(fun(X10,bool),X11),s(X10,X12))))))))) )),
  inference(superposition,[],[f356,f243])).
fof(f1271,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X4,bool)),X5),s(bool,X3)))) | p(s(bool,'const.bool.IN_2'(s(X4,sK16('F_0',X5,X4,bool)),s(fun(X4,bool),happ(s(fun(bool,fun(X4,bool)),X5),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f356,f208])).
fof(f1270,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X1,bool)),X2),s(bool,'F_0')))) | p(s(bool,'const.bool.IN_2'(s(X1,sK16(X0,X2,X1,bool)),s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X2),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f356,f208])).
fof(f1251,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),happ(s(fun(bool,fun(X5,bool)),X7),s(bool,X4)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X5,bool)),X7),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f355,f208])).
fof(f1250,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X3),s(bool,'F_0')))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X1,bool)),X3),s(bool,X0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f355,f208])).
fof(f1217,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X6,bool)),X7),s(fun(bool,bool),X8),s(bool,X5),s(X6,X9)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f281,f208])).
fof(f1216,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RRESTRICT_4'(s(fun(bool,fun(X1,bool)),X2),s(fun(bool,bool),X3),s(bool,'F_0'),s(X1,X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f281,f208])).
fof(f1182,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (sP9(happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)),X123,X124) | ~sP8(happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)),X123,X124) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X119,fun(X120,bool)),X121),s(X119,X122))))) )),
  inference(superposition,[],[f340,f275])).
fof(f1181,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (~sP9(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X118,X117) | p(s(bool,X117)) | p(s(bool,X118)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X113,fun(X114,bool)),X115),s(X113,X116))))) )),
  inference(superposition,[],[f337,f275])).
fof(f1180,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (sP9(happ(s(fun(X107,bool),'const.relation.RDOM_1'(s(fun(X107,fun(X108,bool)),X109))),s(X107,X110)),X111,X112) | ~p(s(bool,X111)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X107,fun(X108,bool)),X109),s(X107,X110))))) )),
  inference(superposition,[],[f336,f275])).
fof(f1179,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (sP9(happ(s(fun(X101,bool),'const.relation.RDOM_1'(s(fun(X101,fun(X102,bool)),X103))),s(X101,X104)),X106,X105) | ~p(s(bool,X105)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X101,fun(X102,bool)),X103),s(X101,X104))))) )),
  inference(superposition,[],[f335,f275])).
fof(f1178,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (~sP8(happ(s(fun(X95,bool),'const.relation.RDOM_1'(s(fun(X95,fun(X96,bool)),X97))),s(X95,X98)),X99,X100) | p(s(bool,X99)) | p(s(bool,X100)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X95,fun(X96,bool)),X97),s(X95,X98))))) )),
  inference(superposition,[],[f330,f275])).
fof(f1173,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (~sP4(X69,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70) | sP5(X69,X70,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X65,fun(X66,bool)),X67),s(X65,X68)))) | ~sP3(X70,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X69) | ~sP2(happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70,X69)) )),
  inference(superposition,[],[f287,f275])).
fof(f1172,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (~sP5(X63,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64) | ~p(s(bool,X63)) | p(s(bool,X64)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X59,fun(X60,bool)),X61),s(X59,X62))))) )),
  inference(superposition,[],[f286,f275])).
fof(f1171,plain,(
  ( ! [X57,X54,X58,X56,X55] : (sP1(X58,happ(s(fun(X54,bool),'const.relation.RDOM_1'(s(fun(X54,fun(X55,bool)),X56))),s(X54,X57))) | p(s(bool,X58)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))))) )),
  inference(superposition,[],[f270,f275])).
fof(f1170,plain,(
  ( ! [X52,X50,X53,X51,X49] : (sP1(happ(s(fun(X49,bool),'const.relation.RDOM_1'(s(fun(X49,fun(X50,bool)),X51))),s(X49,X52)),X53) | p(s(bool,X53)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X49,fun(X50,bool)),X51),s(X49,X52))))) )),
  inference(superposition,[],[f267,f275])).
fof(f1169,plain,(
  ( ! [X47,X45,X48,X46,X44] : (sP0(happ(s(fun(X44,bool),'const.relation.RDOM_1'(s(fun(X44,fun(X45,bool)),X46))),s(X44,X47)),X48) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X44,fun(X45,bool)),X46),s(X44,X47))))) )),
  inference(superposition,[],[f252,f275])).
fof(f1168,plain,(
  ( ! [X39,X43,X41,X42,X40] : (sP0(X43,happ(s(fun(X39,bool),'const.relation.RDOM_1'(s(fun(X39,fun(X40,bool)),X41))),s(X39,X42))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X39,fun(X40,bool)),X41),s(X39,X42))))) )),
  inference(superposition,[],[f251,f275])).
fof(f1158,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),'const.relation.RDOM_1'(s(fun(X8,fun(X9,bool)),X10))),s(X8,X11))) | s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(X8,fun(X9,bool)),X10),s(X8,X11)))) )),
  inference(superposition,[],[f275,f208])).
fof(f1147,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3)),'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f353,f275,f305])).
fof(f1148,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3)),'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f201,f275,f303])).
fof(f1149,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('F_0','const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)),happ(s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f353,f275,f305])).
fof(f1150,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('T_0','const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)),happ(s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f201,f275,f303])).
fof(f1129,plain,(
  ( ! [X94,X97,X95,X93,X96] : (sP9(happ(s(fun(X93,bool),X94),s(X93,X95)),X96,X97) | ~sP8(happ(s(fun(X93,bool),X94),s(X93,X95)),X96,X97) | ~p(s(bool,'const.bool.IN_2'(s(X93,X95),s(fun(X93,bool),X94))))) )),
  inference(superposition,[],[f340,f243])).
fof(f1128,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~sP9(happ(s(fun(X88,bool),X89),s(X88,X90)),X92,X91) | p(s(bool,X91)) | p(s(bool,X92)) | ~p(s(bool,'const.bool.IN_2'(s(X88,X90),s(fun(X88,bool),X89))))) )),
  inference(superposition,[],[f337,f243])).
fof(f1127,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP9(happ(s(fun(X83,bool),X84),s(X83,X85)),X86,X87) | ~p(s(bool,X86)) | ~p(s(bool,'const.bool.IN_2'(s(X83,X85),s(fun(X83,bool),X84))))) )),
  inference(superposition,[],[f336,f243])).
fof(f1126,plain,(
  ( ! [X80,X78,X81,X79,X82] : (sP9(happ(s(fun(X78,bool),X79),s(X78,X80)),X82,X81) | ~p(s(bool,X81)) | ~p(s(bool,'const.bool.IN_2'(s(X78,X80),s(fun(X78,bool),X79))))) )),
  inference(superposition,[],[f335,f243])).
fof(f1125,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~sP8(happ(s(fun(X73,bool),X74),s(X73,X75)),X76,X77) | p(s(bool,X76)) | p(s(bool,X77)) | ~p(s(bool,'const.bool.IN_2'(s(X73,X75),s(fun(X73,bool),X74))))) )),
  inference(superposition,[],[f330,f243])).
fof(f1120,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~sP4(X51,happ(s(fun(X48,bool),X49),s(X48,X50)),X52) | sP5(X51,X52,happ(s(fun(X48,bool),X49),s(X48,X50))) | ~p(s(bool,'const.bool.IN_2'(s(X48,X50),s(fun(X48,bool),X49)))) | ~sP3(X52,happ(s(fun(X48,bool),X49),s(X48,X50)),X51) | ~sP2(happ(s(fun(X48,bool),X49),s(X48,X50)),X52,X51)) )),
  inference(superposition,[],[f287,f243])).
fof(f1119,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~sP5(X46,happ(s(fun(X43,bool),X44),s(X43,X45)),X47) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~p(s(bool,'const.bool.IN_2'(s(X43,X45),s(fun(X43,bool),X44))))) )),
  inference(superposition,[],[f286,f243])).
fof(f1118,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X39,X41),s(fun(X39,bool),X40)))) | p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),X40),s(X39,X41)))) )),
  inference(superposition,[],[f270,f243])).
fof(f1117,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'const.bool.IN_2'(s(X35,X37),s(fun(X35,bool),X36)))) | p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),X36),s(X35,X37)),X38)) )),
  inference(superposition,[],[f267,f243])).
fof(f1116,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X31,X33),s(fun(X31,bool),X32)))) | sP0(happ(s(fun(X31,bool),X32),s(X31,X33)),X34)) )),
  inference(superposition,[],[f252,f243])).
fof(f1115,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(X27,X29),s(fun(X27,bool),X28)))) | sP0(X30,happ(s(fun(X27,bool),X28),s(X27,X29)))) )),
  inference(superposition,[],[f251,f243])).
fof(f1106,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X3,X5),s(fun(X3,bool),X4)))) )),
  inference(superposition,[],[f243,f208])).
fof(f1097,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',happ(s(fun(X0,bool),X1),s(X0,X2)),'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f353,f243,f305])).
fof(f1098,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',happ(s(fun(X0,bool),X1),s(X0,X2)),'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f201,f243,f303])).
fof(f1099,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)),happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f353,f243,f305])).
fof(f1100,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)),happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f201,f243,f303])).
fof(f1040,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f371,f367,f365,f353,f288])).
fof(f998,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f375,f367,f363,f201,f287])).
fof(f945,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f931,f325])).
fof(f931,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f353,f201,f330])).
fof(f915,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f385,f353,f353,f329])).
fof(f846,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f834,f306])).
fof(f834,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f201,f313])).
fof(f822,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f377,f201,f201,f309])).
fof(f818,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f805,f283])).
fof(f805,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f353,f353,f298])).
fof(f797,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f785,f284])).
fof(f785,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f201,f294])).
fof(f765,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f201,f290])).
fof(f652,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f201,f396,f304])).
fof(f654,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f201,f396,f304])).
fof(f631,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f353,f396,f302])).
fof(f632,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f396,f302])).
fof(f505,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f201,f201,f336])).
fof(f494,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f201,f335])).
fof(f489,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f201,f353,f324])).
fof(f487,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f353,f323])).
fof(f481,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f353,f353,f318])).
fof(f474,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f353,f317])).
fof(f455,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f201,f353,f308])).
fof(f452,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f201,f353,f307])).
fof(f444,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f353,f354])).
fof(f431,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f270])).
fof(f423,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f267])).
fof(f420,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f201,f201,f266])).
fof(f419,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f353,f265])).
fof(f385,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f353,f333])).
fof(f383,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f201,f332])).
fof(f381,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f331])).
fof(f380,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f353,f316])).
fof(f379,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f353,f315])).
fof(f377,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f314])).
fof(f375,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f201,f301])).
fof(f373,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f201,f300])).
fof(f371,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f299])).
fof(f370,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f353,f297])).
fof(f369,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f353,f296])).
fof(f367,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f295])).
fof(f366,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f353,f293])).
fof(f365,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f353,f292])).
fof(f363,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f291])).
fof(f361,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f201,f252])).
fof(f359,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f251])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f197])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f131])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f195])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f195])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f191])).
fof(f191,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f132])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f187])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f187])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f184])).
fof(f184,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f129])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f312])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f90,f126,f125,f124,f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f90,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X16,X17,X18] : ((p(s(bool,X18)) <=> s(bool,X16) = s(bool,X17)) <=> ((p(s(bool,X18)) | p(s(bool,X17)) | p(s(bool,X16))) & (p(s(bool,X18)) | ~p(s(bool,X16)) | ~p(s(bool,X17))) & (p(s(bool,X17)) | ~p(s(bool,X16)) | ~p(s(bool,X18))) & (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)))))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.sat.sat_2|49|_prove_ltb')).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f123])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f181])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f181])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f179])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f176])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f176])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f265,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f163])).
fof(f266,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f163])).
fof(f354,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f250,f353])).
fof(f250,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f251,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f252,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f220,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.bool.EQ_CLAUSES|split|2')).
fof(f221,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f217,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f139])).
fof(f2752,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737,f2739,f2741,f2743,f2745,f2747,f2749,f2751])).
fof(f2750,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737,f2739,f2741,f2743,f2745,f2747,f2749])).
fof(f2748,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737,f2739,f2741,f2743,f2745,f2747])).
fof(f2746,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737,f2739,f2741,f2743,f2745])).
fof(f2744,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737,f2739,f2741,f2743])).
fof(f2742,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737,f2739,f2741])).
fof(f2740,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737,f2739])).
fof(f2738,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735,f2737])).
fof(f2736,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733,f2735])).
fof(f2734,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731,f2733])).
fof(f2732,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729,f2731])).
fof(f2730,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727,f2729])).
fof(f2728,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725,f2727])).
fof(f2726,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723,f2725])).
fof(f2724,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721,f2723])).
fof(f2722,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719,f2721])).
fof(f2720,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717,f2719])).
fof(f2718,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715,f2717])).
fof(f2716,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713,f2715])).
fof(f2714,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711,f2713])).
fof(f2712,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709,f2711])).
fof(f2710,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707,f2709])).
fof(f2708,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705,f2707])).
fof(f2706,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703,f2705])).
fof(f2704,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701,f2703])).
fof(f2702,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699,f2701])).
fof(f2700,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697,f2699])).
fof(f2698,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695,f2697])).
fof(f2696,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686,f2695])).
fof(f2687,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684,f2686])).
fof(f2685,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682,f2684])).
fof(f2683,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680,f2682])).
fof(f2681,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678,f2680])).
fof(f2679,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2678])).
fof(f2677,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676])).
fof(f2675,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674])).
fof(f2673,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672])).
fof(f2671,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668,f2670])).
fof(f2669,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666,f2668])).
fof(f2667,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664,f2666])).
fof(f2665,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662,f2664])).
fof(f2663,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660,f2662])).
fof(f2661,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658,f2660])).
fof(f2659,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652,f2658])).
fof(f2654,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2523,f201])).
fof(f2523,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f2430,f353])).
fof(f2653,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f361,f363,f365,f366,f367,f369,f370,f371,f373,f375,f377,f379,f380,f381,f383,f385,f396,f419,f420,f423,f431,f444,f452,f455,f474,f481,f487,f489,f494,f505,f553,f605,f632,f631,f654,f652,f765,f785,f797,f805,f818,f822,f834,f846,f915,f931,f945,f998,f1040,f1100,f1099,f1098,f1097,f1106,f1115,f1116,f1117,f1118,f1119,f1120,f1121,f1122,f1123,f1124,f1125,f1126,f1127,f1128,f1129,f1150,f1149,f1148,f1147,f1157,f1158,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1204,f1205,f1216,f1217,f1250,f1251,f1270,f1271,f1279,f1281,f1294,f1295,f1300,f1301,f1328,f1327,f1326,f1325,f1343,f1344,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1362,f1363,f1364,f1365,f1366,f1367,f1368,f1397,f1415,f1416,f1796,f1845,f1890,f1891,f1892,f1907,f1918,f1919,f1920,f1980,f1982,f1984,f1985,f1986,f1987,f1992,f1993,f2053,f2428,f2432,f2435,f2439,f2442,f2445,f2448,f2451,f2454,f2457,f2460,f2463,f2466,f2082,f2084,f2086,f2087,f2088,f2089,f2468,f2469,f2470,f2471,f2472,f2473,f2474,f2475,f2476,f2477,f2478,f2479,f2480,f2481,f2482,f2483,f2484,f2485,f2486,f2487,f2488,f2489,f2137,f2139,f2141,f2142,f2143,f2144,f2149,f2150,f2495,f2497,f2499,f2502,f2504,f2506,f2508,f2510,f2512,f2514,f2516,f2518,f2520,f2430,f2652])).
fof(f2467,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl32),
  inference(cnf_transformation,[],[f2467_D])).
fof(f2467_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2520,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2519])).
fof(f2519,plain,(
  ( ! [X78,X76,X79,X77] : (~sP8('T_0',X77,X78) | sP9('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2174,f605])).
fof(f2174,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | ~sP8('T_0',X77,X78) | sP9('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f340,f553])).
fof(f2464,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f2464_D])).
fof(f2464_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2518,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2517])).
fof(f2517,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2173,f605])).
fof(f2173,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f337,f553])).
fof(f2461,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f2461_D])).
fof(f2461_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2516,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2515])).
fof(f2515,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f2172,f605])).
fof(f2172,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | ~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f336,f553])).
fof(f2458,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f2458_D])).
fof(f2458_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2514,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2513])).
fof(f2513,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f2171,f605])).
fof(f2171,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f335,f553])).
fof(f2455,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f2455_D])).
fof(f2455_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2512,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f2511])).
fof(f2511,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2170,f605])).
fof(f2170,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f330,f553])).
fof(f2452,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f2452_D])).
fof(f2452_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2510,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f2509])).
fof(f2509,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2169,f605])).
fof(f2169,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f322,f553])).
fof(f2449,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f2449_D])).
fof(f2449_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2508,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f2507])).
fof(f2507,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f2168,f605])).
fof(f2168,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f313,f553])).
fof(f2446,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f2446_D])).
fof(f2446_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2506,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2505])).
fof(f2505,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2167,f605])).
fof(f2167,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f294,f553])).
fof(f2443,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f2443_D])).
fof(f2443_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2504,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2503])).
fof(f2503,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f2166,f605])).
fof(f2166,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f290,f553])).
fof(f2440,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) | $spl14),
  inference(cnf_transformation,[],[f2440_D])).
fof(f2440_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2502,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2501,f375])).
fof(f2501,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2500])).
fof(f2500,plain,(
  ( ! [X43,X41,X42,X40] : (sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f2165,f605])).
fof(f2165,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f287,f553])).
fof(f2436,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f2436_D])).
fof(f2436_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2499,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2498])).
fof(f2498,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f2164,f605])).
fof(f2164,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f286,f553])).
fof(f2433,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl10),
  inference(cnf_transformation,[],[f2433_D])).
fof(f2433_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2497,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2496])).
fof(f2496,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2163,f605])).
fof(f2163,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f270,f553])).
fof(f2429,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl6),
  inference(cnf_transformation,[],[f2429_D])).
fof(f2429_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2495,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2494])).
fof(f2494,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2162,f605])).
fof(f2162,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f267,f553])).
fof(f2466,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2465])).
fof(f2465,plain,(
  ( ! [X78,X76,X79,X77] : (~sP8('T_0',X77,X78) | sP9('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2076,f605])).
fof(f2076,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | ~sP8('T_0',X77,X78) | sP9('T_0',X77,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f340,f553])).
fof(f2463,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2462])).
fof(f2462,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2075,f605])).
fof(f2075,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f337,f553])).
fof(f2460,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2459])).
fof(f2459,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f2074,f605])).
fof(f2074,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | ~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f336,f553])).
fof(f2457,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2456])).
fof(f2456,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f2073,f605])).
fof(f2073,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f335,f553])).
fof(f2454,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f2453])).
fof(f2453,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2072,f605])).
fof(f2072,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f330,f553])).
fof(f2451,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f2450])).
fof(f2450,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2071,f605])).
fof(f2071,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f322,f553])).
fof(f2448,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f2447])).
fof(f2447,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f2070,f605])).
fof(f2070,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f313,f553])).
fof(f2445,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2444])).
fof(f2444,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2069,f605])).
fof(f2069,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f294,f553])).
fof(f2442,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2441])).
fof(f2441,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f2068,f605])).
fof(f2068,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f290,f553])).
fof(f2439,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2438,f375])).
fof(f2438,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2437])).
fof(f2437,plain,(
  ( ! [X43,X41,X42,X40] : (sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f2067,f605])).
fof(f2067,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f287,f553])).
fof(f2435,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2434])).
fof(f2434,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f2066,f605])).
fof(f2066,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f286,f553])).
fof(f2432,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2431])).
fof(f2431,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2065,f605])).
fof(f2065,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f270,f553])).
fof(f2428,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2427])).
fof(f2427,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2064,f605])).
fof(f2064,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f267,f553])).
fof(f347,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13)))))))) | $spl0),
  inference(cnf_transformation,[],[f347_D])).
fof(f347_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13)))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f349,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | $spl2),
  inference(cnf_transformation,[],[f349_D])).
fof(f349_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f198,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  (p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13)))))))) | (p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) & p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13)))))) & (~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13)))))))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13,sK14])],[f135])).
fof(f135,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))))) & (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))))) & (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) | (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))))))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) <~> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RRESTRICT_2'(s(fun(X1,fun(X0,bool)),X4),s(fun(X1,bool),X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),'const.relation.RRESTRICT_2'(s(fun(X7,fun(X6,bool)),X10),s(fun(X7,bool),X9)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),X10)))))) & p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),'const.relation.RRESTRICT_2'(s(fun(X7,fun(X6,bool)),X10),s(fun(X7,bool),X9)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),X10)))))) & p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X9))))))),
  file('thm.relation.IN_RDOM_RRESTRICT','thm.relation.IN_RDOM_RRESTRICT')).
fof(f351,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13)))) | $spl4),
  inference(cnf_transformation,[],[f351_D])).
fof(f351_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f199,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))))),
  inference(cnf_transformation,[],[f136])).
fof(f348,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13)))))))) | $spl1),
  inference(cnf_transformation,[],[f348_D])).
fof(f348_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f350,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | $spl3),
  inference(cnf_transformation,[],[f350_D])).
fof(f350_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f352,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13)))) | $spl5),
  inference(cnf_transformation,[],[f352_D])).
fof(f352_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f200,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RRESTRICT_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(fun(sK11,bool),sK13))))))))),
  inference(cnf_transformation,[],[f136])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1457
% Passive clauses: 3512
% Generated clauses: 26218
% Final active clauses: 945
% Final passive clauses: 1188
% Input formulas: 37
% Initial clauses: 142
% 
% Duplicate literals: 379
% Trivial inequalities: 44
% Fw subsumption resolutions: 727
% Fw demodulations: 5963
% Bw demodulations: 136
% Condensations: 668
% Global subsumptions: 49
% 
% Simple tautologies: 70
% Equational tautologies: 208
% Forward subsumptions: 14476
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 8974
% Factoring: 1
% Forward superposition: 4889
% Backward superposition: 3999
% Self superposition: 45
% Equality factoring: 32
% Equality resolution: 2
% 
% Split clauses: 71
% Split components: 148
% SAT solver clauses: 3756
% SAT solver unit clauses: 134
% SAT solver binary clauses: 1435
% 
% Sat splits: 71
% Sat splitting refutations: 51
% 
% Memory used [KB]: 20340
% Time elapsed: 1.172 s
% ------------------------------
% ------------------------------
% Success in time 2.493 s
