Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (107983)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 691
% Passive clauses: 7850
% Generated clauses: 18077
% Final active clauses: 678
% Final passive clauses: 6750
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 423
% Trivial inequalities: 57
% Fw subsumption resolutions: 1377
% Fw demodulations: 356
% Bw demodulations: 9
% 
% Simple tautologies: 221
% Equational tautologies: 63
% Forward subsumptions: 7540
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 10081
% Factoring: 2312
% Forward superposition: 1076
% Backward superposition: 1934
% Self superposition: 40
% Equality factoring: 26
% Equality resolution: 24
% 
% Split clauses: 287
% Split components: 575
% SAT solver clauses: 569
% SAT solver unit clauses: 18
% SAT solver binary clauses: 93
% 
% Sat splits: 287
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 569
% TWLsolver calls for satisfiability: 394
% 
% Memory used [KB]: 7547
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (107999)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 554
% Passive clauses: 2608
% Generated clauses: 17191
% Final active clauses: 369
% Final passive clauses: 1169
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 256
% Trivial inequalities: 196
% Fw subsumption resolutions: 1200
% Fw demodulations: 2141
% Bw demodulations: 847
% Condensations: 421
% Global subsumptions: 22
% 
% Simple tautologies: 329
% Equational tautologies: 33
% Forward subsumptions: 7002
% Fw demodulations to eq. taut.: 27
% Bw demodulations to eq. taut.: 100
% 
% Binary resolution: 2762
% Factoring: 255
% Forward superposition: 3115
% Backward superposition: 2379
% Self superposition: 62
% Equality factoring: 28
% Equality resolution: 50
% 
% Split clauses: 255
% Split components: 526
% SAT solver clauses: 2868
% SAT solver unit clauses: 254
% SAT solver binary clauses: 118
% SAT solver learnt clauses: 47
% SAT solver learnt literals: 14
% 
% Sat splits: 255
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 2777
% TWLsolver calls for satisfiability: 1959
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f21373,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f306,f545,f543,f305,f544,f542,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4996,f4912,f4998,f4916,f5000,f4919,f5002,f4922,f5004,f4925,f5006,f4928,f5008,f4931,f5010,f4934,f5012,f4937,f5014,f4940,f5016,f4943,f5018,f4946,f5020,f4949,f5022,f4952,f5024,f4955,f5026,f4958,f5028,f4913,f4961,f5251,f5252,f5260,f5262,f5264,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348,f5350,f5351,f5352,f5353,f5354,f5356,f5358,f5681,f5683,f5716,f5718,f5719,f5721,f5726,f5728,f5730,f5732,f5736,f5738,f5740,f5742,f5744,f5746,f5752,f5754,f5760,f5762,f5768,f5774,f5787,f5811,f5835,f5859,f5917,f5919,f5921,f5924,f5926,f5931,f5933,f5938,f5940,f5953,f5955,f5957,f5960,f5962,f5966,f5968,f5999,f6010,f6026,f6037,f6039,f6048,f6055,f6064,f6071,f6073,f6075,f6077,f6079,f6604,f6605,f6124,f6126,f6128,f6130,f6137,f6144,f6154,f6162,f6164,f6175,f6183,f6185,f6195,f6202,f6211,f6218,f6220,f6229,f6231,f6236,f6238,f6247,f6249,f6254,f6256,f6265,f6267,f6272,f6274,f6283,f6290,f6299,f6306,f6312,f6318,f6327,f6337,f6339,f6348,f6350,f6352,f6357,f6359,f6364,f6366,f6371,f6373,f6378,f6380,f6388,f6397,f6407,f6415,f6417,f6433,f6435,f6437,f6439,f6450,f6453,f6464,f6467,f6482,f6484,f6498,f6500,f6502,f6508,f6513,f6518,f6520,f6525,f6527,f6538,f6547,f6558,f6567,f6572,f6574,f6576,f6578,f6584,f6586,f6588,f6606,f6607,f6608,f6609,f6610,f6611,f6612,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6620,f6621,f6622,f6623,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6633,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6644,f6645,f6646,f6647,f6648,f6649,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f6670,f6671,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6680,f6681,f6682,f6683,f6684,f6685,f6686,f6687,f6688,f6689,f6690,f6691,f6692,f6693,f6694,f6695,f6696,f6697,f6698,f6699,f6700,f6701,f6702,f6703,f6704,f6705,f6706,f6707,f6708,f6709,f6710,f6711,f6712,f6713,f6714,f6715,f6716,f6717,f6718,f6719,f6721,f6722,f6723,f6724,f6725,f6726,f6727,f6728,f6729,f6730,f6731,f6732,f6733,f6734,f6735,f6736,f6737,f6738,f6739,f6740,f6741,f6742,f6743,f6744,f6745,f6746,f6747,f6748,f6749,f6750,f6751,f6752,f6753,f6754,f6755,f6756,f6757,f6758,f6759,f6760,f6761,f6762,f6763,f6764,f6765,f6766,f6767,f6768,f6769,f6770,f6771,f6772,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6781,f6782,f6783,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6791,f6792,f6793,f6794,f6795,f6796,f6797,f6798,f6799,f6800,f6801,f6802,f6803,f6804,f6805,f6806,f6807,f6808,f6809,f6810,f6811,f6812,f6813,f6814,f8655,f8657,f8656,f12661,f12663,f12662,f12861,f12866,f12867,f13010,f13014,f13017,f13019,f13012,f13524,f13525,f13532,f13529,f13526,f13527,f13535,f13528,f13537,f13530,f13534,f13544,f13545,f13557,f13555,f20468,f20470,f20469,f20472,f20473,f20504,f20506,f20505,f20508,f5684,f20509,f21061,f21082,f21095,f21099,f21101,f21105,f21109,f21113,f21116,f21120,f21124,f21128,f21132,f21136,f21140,f21144,f21148,f21152,f21156,f21160,f21164,f21168,f21172,f21176,f21180,f21184,f21188,f21191,f21193,f21194,f21334])).
fof(f21334,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21333,f546])).
fof(f546,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f310,f307])).
fof(f307,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4',aHL_TRUTH)).
fof(f310,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f38])).
fof(f38,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.bool.NOT_CLAUSES|split|2')).
fof(f21333,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21332,f15525])).
fof(f15525,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f598,f12895,f780])).
fof(f780,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f774,f546])).
fof(f774,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f406,f321])).
fof(f321,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4',aHL_BOOL_CASES)).
fof(f406,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.bool.IMP_ANTISYM_AX')).
fof(f12895,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f12871,f2225])).
fof(f2225,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2220,f547])).
fof(f547,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f318,f314])).
fof(f314,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.arithmetic.NUMERAL_DEF')).
fof(f318,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_MOD_REDUCE|split|0')).
fof(f2220,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(trivial_inequality_removal,[],[f2214])).
fof(f2214,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(superposition,[],[f413,f547])).
fof(f413,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & (((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & (((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & ~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X16,X17] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X17),s('type.integer.int',X16))) & ~s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X17) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_DIVIDES_MOD|48|')).
fof(f12871,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f12863,f416])).
fof(f416,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X7,X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_INJ')).
fof(f12863,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | $spl3),
  inference(forward_demodulation,[],[f12862,f314])).
fof(f12862,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | $spl3),
  inference(forward_demodulation,[],[f545,f314])).
fof(f598,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f546,f332])).
fof(f332,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.bool.EQ_CLAUSES|split|0')).
fof(f21332,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21053,f555])).
fof(f555,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f554,f314])).
fof(f554,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f402,f314])).
fof(f402,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_MOD_REDUCE|split|2')).
fof(f21053,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(trivial_inequality_removal,[],[f21017])).
fof(f21017,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f2305,f13021])).
fof(f13021,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f13020,f2399])).
fof(f2399,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f2391,f416])).
fof(f2391,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(condensation,[],[f2390])).
fof(f2390,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f2389,f1640])).
fof(f1640,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f389,f420])).
fof(f420,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.integer_244_prove_ltb|split|1')).
fof(f389,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.arithmetic.SUC_NOT')).
fof(f2389,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(forward_demodulation,[],[f2340,f400])).
fof(f400,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f2340,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) )),
  inference(superposition,[],[f422,f551])).
fof(f551,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f549,f315])).
fof(f315,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.arithmetic.ADD_|48|')).
fof(f549,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(backward_demodulation,[],[f400,f317])).
fof(f317,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.arithmetic.BIT|49|')).
fof(f422,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f13020,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f12967,f307])).
fof(f12967,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f410,f12872])).
fof(f12872,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f5682,f5679])).
fof(f5679,plain,(
  ( ! [X4] : (sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl4),
  inference(global_subsumption,[],[f306,f305,f307,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f389,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1437,f1436,f1435,f1434,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2298,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2840,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4115,f4114,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4434,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4669,f4668,f4527,f4526,f392,f4989,f393,f4990,f2294,f2218,f2305,f2225,f4148,f2872,f4166,f2222,f4102,f2871,f4084,f2864,f2861,f5361,f4988,f4987,f4447,f4411,f4408,f2784,f2781,f2710,f2707,f1468,f1467,f1466,f1465,f1464,f1463,f1462,f1461,f1460,f1459,f1458,f1457,f1456,f1455,f1454,f1453,f1452,f1451,f1450,f1411,f1410,f1409,f1408,f1407,f1406,f1405,f1404,f1403,f1402,f1401,f1400,f1399,f1398,f1397,f1396,f1395,f1394,f1393,f415,f414,f413,f412,f410,f399,f398,f4525,f4667,f546,f337,f336,f4912,f5678])).
fof(f5678,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f5669,f332])).
fof(f5669,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X5)) ) | $spl4),
  inference(superposition,[],[f4912,f729])).
fof(f336,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.bool.EQ_CLAUSES|split|3')).
fof(f337,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f4667,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'T_0') != s(bool,f(s(X159,sK16(X160,X161,bool,X159)))) | s(fun(X159,bool),X160) = s(fun(X159,bool),X161) | s(bool,X162) = s(bool,g(s(X159,sK16(X160,X161,bool,X159)))) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f455,f729])).
fof(f4525,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,X170) != s(bool,f(s(X167,sK16(X168,X169,bool,X167)))) | s(fun(X167,bool),X168) = s(fun(X167,bool),X169) | s(bool,'T_0') = s(bool,X170) | s(bool,'T_0') = s(bool,g(s(X167,sK16(X168,X169,bool,X167))))) )),
  inference(superposition,[],[f455,f729])).
fof(f398,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X16,X17] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X17)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_DIVIDES_NEG|split|0')).
fof(f399,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X16,X17] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X16))),s('type.integer.int',X17)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_DIVIDES_NEG|split|1')).
fof(f412,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f250])).
fof(f414,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f415,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f1393,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int',X20))))) )),
  inference(superposition,[],[f431,f398])).
fof(f1394,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int',X23))))) )),
  inference(superposition,[],[f432,f398])).
fof(f1395,plain,(
  ( ! [X26,X27,X25] : (sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27) | p(s(bool,X27)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int',X26))))) )),
  inference(superposition,[],[f447,f398])).
fof(f1396,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29))))) | p(s(bool,X30)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int',X29))))) )),
  inference(superposition,[],[f450,f398])).
fof(f1397,plain,(
  ( ! [X33,X31,X34,X32] : (~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int',X32))))) )),
  inference(superposition,[],[f459,f398])).
fof(f1398,plain,(
  ( ! [X37,X35,X38,X36] : (~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int',X36))))) )),
  inference(superposition,[],[f463,f398])).
fof(f1399,plain,(
  ( ! [X39,X41,X42,X40] : (~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int',X40))))) )),
  inference(superposition,[],[f467,f398])).
fof(f1400,plain,(
  ( ! [X45,X43,X46,X44] : (~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int',X44))))) )),
  inference(superposition,[],[f486,f398])).
fof(f1401,plain,(
  ( ! [X47,X50,X48,X49] : (~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48))))) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int',X48))))) )),
  inference(superposition,[],[f495,f398])).
fof(f1402,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54) | p(s(bool,X53)) | p(s(bool,X54)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int',X52))))) )),
  inference(superposition,[],[f503,f398])).
fof(f1403,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57) | ~p(s(bool,X57)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int',X56))))) )),
  inference(superposition,[],[f508,f398])).
fof(f1404,plain,(
  ( ! [X61,X59,X62,X60] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62) | ~p(s(bool,X61)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))))) )),
  inference(superposition,[],[f509,f398])).
fof(f1405,plain,(
  ( ! [X66,X64,X65,X63] : (~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65) | p(s(bool,X65)) | p(s(bool,X66)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64))))) )),
  inference(superposition,[],[f510,f398])).
fof(f1406,plain,(
  ( ! [X70,X68,X69,X67] : (~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int',X68))))) )),
  inference(superposition,[],[f520,f398])).
fof(f1407,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73) | p(s(bool,X73)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int',X72))))) )),
  inference(superposition,[],[f525,f398])).
fof(f1408,plain,(
  ( ! [X78,X76,X77,X75] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78) | ~p(s(bool,X77)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int',X76))))) )),
  inference(superposition,[],[f526,f398])).
fof(f1409,plain,(
  ( ! [X80,X81,X79,X82] : (~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int',X80))))) )),
  inference(superposition,[],[f527,f398])).
fof(f1410,plain,(
  ( ! [X85,X83,X86,X84] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | ~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int',X84))))) )),
  inference(superposition,[],[f530,f398])).
fof(f1411,plain,(
  ( ! [X90,X88,X87,X89] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))))) )),
  inference(superposition,[],[f531,f398])).
fof(f1450,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int',X22))))) )),
  inference(superposition,[],[f431,f399])).
fof(f1451,plain,(
  ( ! [X26,X24,X25] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X24),s('type.integer.int',X25))))) )),
  inference(superposition,[],[f432,f399])).
fof(f1452,plain,(
  ( ! [X28,X29,X27] : (sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29) | p(s(bool,X29)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X27),s('type.integer.int',X28))))) )),
  inference(superposition,[],[f447,f399])).
fof(f1453,plain,(
  ( ! [X30,X31,X32] : (sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31))) | p(s(bool,X32)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X30),s('type.integer.int',X31))))) )),
  inference(superposition,[],[f450,f399])).
fof(f1454,plain,(
  ( ! [X35,X33,X36,X34] : (~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X33),s('type.integer.int',X34))))) )),
  inference(superposition,[],[f459,f399])).
fof(f1455,plain,(
  ( ! [X39,X37,X38,X40] : (~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X37),s('type.integer.int',X38))))) )),
  inference(superposition,[],[f463,f399])).
fof(f1456,plain,(
  ( ! [X43,X41,X44,X42] : (~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int',X42))))) )),
  inference(superposition,[],[f467,f399])).
fof(f1457,plain,(
  ( ! [X47,X45,X48,X46] : (~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X45),s('type.integer.int',X46))))) )),
  inference(superposition,[],[f486,f399])).
fof(f1458,plain,(
  ( ! [X52,X50,X51,X49] : (~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50))) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X49),s('type.integer.int',X50))))) )),
  inference(superposition,[],[f495,f399])).
fof(f1459,plain,(
  ( ! [X54,X56,X55,X53] : (~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56) | p(s(bool,X55)) | p(s(bool,X56)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X53),s('type.integer.int',X54))))) )),
  inference(superposition,[],[f503,f399])).
fof(f1460,plain,(
  ( ! [X59,X57,X60,X58] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59) | ~p(s(bool,X59)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X57),s('type.integer.int',X58))))) )),
  inference(superposition,[],[f508,f399])).
fof(f1461,plain,(
  ( ! [X61,X64,X62,X63] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64) | ~p(s(bool,X63)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X61),s('type.integer.int',X62))))) )),
  inference(superposition,[],[f509,f399])).
fof(f1462,plain,(
  ( ! [X68,X66,X67,X65] : (~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67) | p(s(bool,X67)) | p(s(bool,X68)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int',X66))))) )),
  inference(superposition,[],[f510,f399])).
fof(f1463,plain,(
  ( ! [X70,X72,X71,X69] : (~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70))) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X69),s('type.integer.int',X70))))) )),
  inference(superposition,[],[f520,f399])).
fof(f1464,plain,(
  ( ! [X76,X74,X75,X73] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75) | p(s(bool,X75)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X73),s('type.integer.int',X74))))) )),
  inference(superposition,[],[f525,f399])).
fof(f1465,plain,(
  ( ! [X80,X78,X79,X77] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80) | ~p(s(bool,X79)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X77),s('type.integer.int',X78))))) )),
  inference(superposition,[],[f526,f399])).
fof(f1466,plain,(
  ( ! [X83,X81,X84,X82] : (~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X81),s('type.integer.int',X82))))) )),
  inference(superposition,[],[f527,f399])).
fof(f1467,plain,(
  ( ! [X88,X87,X85,X86] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X85),s('type.integer.int',X86))))) )),
  inference(superposition,[],[f530,f399])).
fof(f1468,plain,(
  ( ! [X92,X90,X91,X89] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X89),s('type.integer.int',X90))))) )),
  inference(superposition,[],[f531,f399])).
fof(f2707,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X3),s('type.integer.int',X4)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f410,f398])).
fof(f2710,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) )),
  inference(superposition,[],[f410,f399])).
fof(f2781,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X3),s('type.integer.int',X4)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) )),
  inference(superposition,[],[f412,f398])).
fof(f2784,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f412,f399])).
fof(f4408,plain,(
  ( ! [X30,X28,X29] : (s(bool,X30) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) )),
  inference(superposition,[],[f729,f398])).
fof(f4411,plain,(
  ( ! [X39,X37,X38] : (s(bool,X39) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X37),s('type.integer.int',X38))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)))) )),
  inference(superposition,[],[f729,f399])).
fof(f4447,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int',X22))) | s(bool,X23) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f729,f398])).
fof(f4987,plain,(
  ( ! [X127,X125,X126] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X126),s('type.integer.int',X125))) | s('type.integer.int',X126) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X127) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X125),s('type.integer.int',X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(subsumption_resolution,[],[f4656,f307])).
fof(f4656,plain,(
  ( ! [X127,X125,X126] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X126),s('type.integer.int',X125))) | s('type.integer.int',X126) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X127) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X125),s('type.integer.int',X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f412,f729])).
fof(f4988,plain,(
  ( ! [X130,X128,X129] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X129),s('type.integer.int',X128))) | s('type.integer.int',X128) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X130) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X128),s('type.integer.int',X129))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(subsumption_resolution,[],[f4657,f307])).
fof(f4657,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X129),s('type.integer.int',X128))) | s('type.integer.int',X128) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X130) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X128),s('type.integer.int',X129))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f410,f729])).
fof(f5361,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2701,f2861])).
fof(f2861,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2860,f547])).
fof(f2860,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))))) )),
  inference(subsumption_resolution,[],[f2841,f331])).
fof(f2841,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))))) )),
  inference(superposition,[],[f414,f547])).
fof(f2864,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4)))))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2843,f331])).
fof(f2843,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4)))))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f414,f331])).
fof(f4084,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f547,f2399,f414])).
fof(f2871,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))))) )),
  inference(subsumption_resolution,[],[f2854,f2399])).
fof(f2854,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))))) )),
  inference(superposition,[],[f414,f555])).
fof(f4102,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) )),
  inference(forward_demodulation,[],[f4087,f399])).
fof(f4087,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f2347,f2399,f413])).
fof(f2222,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(forward_demodulation,[],[f2212,f548])).
fof(f2212,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(unit_resulting_resolution,[],[f1059,f548,f413])).
fof(f4166,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) )),
  inference(forward_demodulation,[],[f4151,f399])).
fof(f4151,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f2347,f2412,f413])).
fof(f2872,plain,(
  ( ! [X6,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) )),
  inference(subsumption_resolution,[],[f2855,f2412])).
fof(f2855,plain,(
  ( ! [X6,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) )),
  inference(superposition,[],[f414,f557])).
fof(f4148,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f548,f2412,f414])).
fof(f2218,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(trivial_inequality_removal,[],[f2216])).
fof(f2216,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f413,f331])).
fof(f2294,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(trivial_inequality_removal,[],[f2292])).
fof(f2292,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X5) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f415,f331])).
fof(f4990,plain,(
  ( ! [X154,X152,X155,X153] : (p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | s(bool,X155) = s(bool,happ(s(fun(X152,bool),X153),s(X152,sK15(X153,X152)))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f4666,f307])).
fof(f4666,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | s(bool,X155) = s(bool,happ(s(fun(X152,bool),X153),s(X152,sK15(X153,X152)))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f393,f729])).
fof(f393,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f244])).
fof(f244,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X12] : ! [X14] : (~! [X8] : p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X8)))) <=> ? [X8] : ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X8)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.bool.NOT_FORALL_THM')).
fof(f4989,plain,(
  ( ! [X151,X149,X150,X148] : (p(s(bool,happ(s(fun(X148,bool),X149),s(X148,X150)))) | s(bool,X151) = s(bool,happ(s(fun(X148,bool),X149),s(X148,sK14(X149,X148)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(subsumption_resolution,[],[f4665,f307])).
fof(f4665,plain,(
  ( ! [X151,X149,X150,X148] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X148,bool),X149),s(X148,X150)))) | s(bool,X151) = s(bool,happ(s(fun(X148,bool),X149),s(X148,sK14(X149,X148)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f392,f729])).
fof(f392,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f4526,plain,(
  ( ! [X173,X171,X172] : (~p(s(bool,happ(s(fun(bool,bool),X171),s(bool,X172)))) | p(s(bool,happ(s(fun(bool,bool),X171),s(bool,X173)))) | s(bool,'T_0') = s(bool,X172) | s(bool,'T_0') = s(bool,sK14(X171,bool))) )),
  inference(superposition,[],[f392,f729])).
fof(f4527,plain,(
  ( ! [X175,X176,X174] : (~p(s(bool,happ(s(fun(bool,bool),X174),s(bool,X175)))) | p(s(bool,happ(s(fun(bool,bool),X174),s(bool,X176)))) | s(bool,'T_0') = s(bool,X175) | s(bool,'T_0') = s(bool,sK15(X174,bool))) )),
  inference(superposition,[],[f393,f729])).
fof(f4668,plain,(
  ( ! [X165,X163,X164] : (~p(s(bool,happ(s(fun(bool,bool),X163),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X163),s(bool,X164)))) | s(bool,X165) = s(bool,sK14(X163,bool)) | s(bool,'T_0') = s(bool,X165)) )),
  inference(superposition,[],[f392,f729])).
fof(f4669,plain,(
  ( ! [X167,X166,X168] : (~p(s(bool,happ(s(fun(bool,bool),X166),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X166),s(bool,X167)))) | s(bool,X168) = s(bool,sK15(X166,bool)) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f393,f729])).
fof(f4670,plain,(
  ( ! [X171,X169,X172,X170] : (s(X171,f(s(bool,'T_0'))) != s(X171,g(s(bool,'T_0'))) | s(fun(bool,X171),X169) = s(fun(bool,X171),X170) | s(bool,X172) = s(bool,sK16(X169,X170,X171,bool)) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f455,f729])).
fof(f4985,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f4642,f307])).
fof(f4642,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f531,f729])).
fof(f4984,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4641,f307])).
fof(f4641,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f530,f729])).
fof(f4983,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4640,f307])).
fof(f4640,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f527,f729])).
fof(f4982,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4639,f307])).
fof(f4639,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f526,f729])).
fof(f4981,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4638,f307])).
fof(f4638,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f525,f729])).
fof(f4980,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4637,f307])).
fof(f4637,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f520,f729])).
fof(f4979,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4636,f307])).
fof(f4636,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f510,f729])).
fof(f4978,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4635,f307])).
fof(f4635,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f509,f729])).
fof(f4977,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4634,f307])).
fof(f4634,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f508,f729])).
fof(f4976,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4633,f307])).
fof(f4633,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f503,f729])).
fof(f4975,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4632,f307])).
fof(f4632,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f495,f729])).
fof(f4974,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4631,f307])).
fof(f4631,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f486,f729])).
fof(f4973,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4630,f307])).
fof(f4630,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f467,f729])).
fof(f4972,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4629,f307])).
fof(f4629,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f463,f729])).
fof(f4971,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4628,f307])).
fof(f4628,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f459,f729])).
fof(f4970,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4627,f307])).
fof(f4627,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f450,f729])).
fof(f4969,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4626,f307])).
fof(f4626,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f447,f729])).
fof(f4968,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f4625,f307])).
fof(f4625,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f432,f729])).
fof(f4967,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f4624,f307])).
fof(f4624,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f431,f729])).
fof(f4966,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4623,f307])).
fof(f4623,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f409,f729])).
fof(f4965,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4619,f307])).
fof(f4619,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f336,f729])).
fof(f4586,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f332,f729])).
fof(f4528,plain,(
  ( ! [X177,X180,X178,X179] : (s(X179,f(s(bool,X180))) != s(X179,g(s(bool,X180))) | s(fun(bool,X179),X177) = s(fun(bool,X179),X178) | s(bool,'T_0') = s(bool,X180) | s(bool,'T_0') = s(bool,sK16(X177,X178,X179,bool))) )),
  inference(superposition,[],[f455,f729])).
fof(f4434,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f729,f729])).
fof(f4386,plain,(
  ~sP5('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f3477,f457])).
fof(f4381,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f3472,f457])).
fof(f4373,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(forward_demodulation,[],[f4329,f1813])).
fof(f4329,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) )),
  inference(superposition,[],[f553,f1883])).
fof(f4328,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))))) )),
  inference(superposition,[],[f551,f1883])).
fof(f4327,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3))),s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.num.num',X4)))) )),
  inference(superposition,[],[f401,f1883])).
fof(f4326,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f1883])).
fof(f4305,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f570,f575,f3471,f460])).
fof(f4296,plain,(
  ~sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f3466,f458])).
fof(f4253,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X25))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24))))) | s('type.num.num',X25) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f416,f1786])).
fof(f4290,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f4241,f1813])).
fof(f4241,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))) )),
  inference(superposition,[],[f553,f1786])).
fof(f4240,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5)))))) )),
  inference(superposition,[],[f551,f1786])).
fof(f4239,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3))),s('type.num.num',X4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f401,f1786])).
fof(f4238,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f1786])).
fof(f4182,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1656,f418])).
fof(f4114,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1640,f331])).
fof(f4115,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f1640,f418])).
fof(f4055,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1437,f458])).
fof(f4002,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f1435,f2347])).
fof(f3999,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1435,f458])).
fof(f3970,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f1434,f2347])).
fof(f3942,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1382,f458])).
fof(f3895,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1380,f2347])).
fof(f3892,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f1380,f458])).
fof(f3849,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1379,f2347])).
fof(f3764,plain,(
  ( ! [X63] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X63))) | s('type.integer.int',X63) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f3686,f2347])).
fof(f3686,plain,(
  ( ! [X63] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X63))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X63) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f418,f2347])).
fof(f3650,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f418,f2347])).
fof(f3445,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2298,f409])).
fof(f3447,plain,(
  ( ! [X0] : (sP0(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2298,f431])).
fof(f3448,plain,(
  ( ! [X0] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2298,f432])).
fof(f3449,plain,(
  ~sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f2298,f446])).
fof(f3451,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2298,f446])).
fof(f3452,plain,(
  ~sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2298,f2298,f446])).
fof(f3453,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f546,f2298,f447])).
fof(f3454,plain,(
  sP1('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f2298,f450])).
fof(f3457,plain,(
  ~sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f459])).
fof(f3458,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2298,f546,f2298,f459])).
fof(f3460,plain,(
  sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f570,f567,f574,f2298,f460])).
fof(f3461,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f463])).
fof(f3466,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f463])).
fof(f3467,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2298,f546,f2298,f463])).
fof(f3471,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2298,f464])).
fof(f3472,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f467])).
fof(f3477,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f467])).
fof(f3478,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2298,f546,f2298,f467])).
fof(f3482,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2298,f468])).
fof(f3483,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2298,f472])).
fof(f3484,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2298,f473])).
fof(f3485,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2298,f474])).
fof(f3490,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f398,f2298,f476])).
fof(f3491,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f399,f2298,f476])).
fof(f3496,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f398,f2298,f476])).
fof(f3497,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f399,f2298,f476])).
fof(f3498,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f598,f2298,f477])).
fof(f3500,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f598,f2298,f477])).
fof(f3502,plain,(
  ( ! [X0] : (~sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f2298,f480])).
fof(f3503,plain,(
  ( ! [X0] : (~sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2298,f481])).
fof(f3510,plain,(
  ~sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f486])).
fof(f3515,plain,(
  ~sP6('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f486])).
fof(f3516,plain,(
  ~sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2298,f546,f2298,f486])).
fof(f3520,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2298,f487])).
fof(f3521,plain,(
  sP7('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f2298,f492])).
fof(f3526,plain,(
  sP7('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f307,f2298,f492])).
fof(f3528,plain,(
  sP7('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2298,f2298,f492])).
fof(f3531,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f2298,f492])).
fof(f3532,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f2298,f307,f2298,f492])).
fof(f3533,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2298,f2298,f492])).
fof(f3534,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2298,f2298,f2298,f492])).
fof(f3535,plain,(
  ~sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f495])).
fof(f3539,plain,(
  ~sP7('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f495])).
fof(f3540,plain,(
  ~sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2298,f546,f2298,f495])).
fof(f3543,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f546,f2298,f496])).
fof(f3545,plain,(
  ( ! [X0] : (~sP9('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2298,f497])).
fof(f3549,plain,(
  ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f2298,f503])).
fof(f3552,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2298,f504])).
fof(f3553,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2298,f505])).
fof(f3554,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f307,f2298,f508])).
fof(f3558,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2298,f2298,f508])).
fof(f3559,plain,(
  ( ! [X0] : (sP9('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f307,f2298,f509])).
fof(f3563,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2298,f2298,f509])).
fof(f3564,plain,(
  ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f2298,f510])).
fof(f3565,plain,(
  ( ! [X0] : (~sP11('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2298,f514])).
fof(f3580,plain,(
  ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f2298,f520])).
fof(f3585,plain,(
  ~sP10('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f307,f2298,f520])).
fof(f3586,plain,(
  ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f2298,f2298,f520])).
fof(f3590,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2298,f522])).
fof(f3591,plain,(
  sP11('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f546,f2298,f524])).
fof(f3594,plain,(
  ( ! [X0] : (sP11('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f307,f2298,f526])).
fof(f3598,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2298,f2298,f526])).
fof(f3599,plain,(
  ~sP11('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f527])).
fof(f3602,plain,(
  ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2298,f527])).
fof(f3603,plain,(
  ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2298,f546,f2298,f527])).
fof(f3604,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f2298,f530])).
fof(f3606,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f583,f2298,f530])).
fof(f3607,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f579,f2298,f531])).
fof(f3608,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f580,f2298,f531])).
fof(f3401,plain,(
  ( ! [X21,X20] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X20))))) | s('type.num.num',X21) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) )),
  inference(superposition,[],[f416,f1813])).
fof(f3391,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f551,f1813])).
fof(f3425,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) )),
  inference(forward_demodulation,[],[f3389,f1883])).
fof(f3389,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4)))))))) )),
  inference(superposition,[],[f553,f1813])).
fof(f3388,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f400,f1813])).
fof(f3387,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f401,f1813])).
fof(f3330,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f416,f1571])).
fof(f3320,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f551,f1571])).
fof(f3357,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f3356,f1044])).
fof(f3356,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f3318,f1883])).
fof(f3318,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f553,f1571])).
fof(f3355,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f3317,f1808])).
fof(f3317,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f400,f1571])).
fof(f3354,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f3316,f1879])).
fof(f3316,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f401,f1571])).
fof(f3133,plain,(
  ( ! [X39,X38,X40] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39))))))))) | s('type.integer.int',X40) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39)))))))))) )),
  inference(superposition,[],[f418,f559])).
fof(f3120,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(superposition,[],[f331,f559])).
fof(f2945,plain,(
  ( ! [X39,X38,X40] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))))))))) | s('type.integer.int',X40) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39)))))))))) )),
  inference(superposition,[],[f418,f558])).
fof(f2932,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(superposition,[],[f331,f558])).
fof(f2840,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f1060,f547,f414])).
fof(f2701,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f410,f331])).
fof(f2590,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f557,f548])).
fof(f2476,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f555,f547])).
fof(f2409,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f2394,f421])).
fof(f2410,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f2394,f420])).
fof(f2412,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f2394,f416])).
fof(f2396,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f2391,f421])).
fof(f2397,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f2391,f420])).
fof(f2394,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(condensation,[],[f2393])).
fof(f2393,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9)) )),
  inference(subsumption_resolution,[],[f2392,f1640])).
fof(f2392,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9)) )),
  inference(forward_demodulation,[],[f2341,f400])).
fof(f2341,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))))))) )),
  inference(superposition,[],[f422,f553])).
fof(f2352,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f2347,f2136])).
fof(f2351,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))))))) )),
  inference(backward_demodulation,[],[f2347,f1997])).
fof(f2350,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))))) )),
  inference(backward_demodulation,[],[f2347,f1930])).
fof(f2347,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f2335,f315])).
fof(f2335,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f316,f316,f422])).
fof(f2298,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f2297,f548])).
fof(f2297,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f2289,f548])).
fof(f2289,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))))) )),
  inference(unit_resulting_resolution,[],[f548,f548,f415])).
fof(f2221,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f2213,f548])).
fof(f2213,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f1060,f548,f413])).
fof(f2138,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14))))) | s('type.integer.int',X14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f419,f331])).
fof(f2136,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f418,f331])).
fof(f2127,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1059,f331])).
fof(f1997,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f418,f548])).
fof(f1930,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f418,f547])).
fof(f1883,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f1882,f1813])).
fof(f1882,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))))))) )),
  inference(forward_demodulation,[],[f1881,f1786])).
fof(f1881,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3)))))))))) )),
  inference(forward_demodulation,[],[f1880,f400])).
fof(f1880,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))))) )),
  inference(forward_demodulation,[],[f1854,f400])).
fof(f1854,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))))) )),
  inference(superposition,[],[f401,f551])).
fof(f1879,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f1843,f316])).
fof(f1843,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(superposition,[],[f401,f1044])).
fof(f1813,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f1812,f1786])).
fof(f1812,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f1787,f400])).
fof(f1787,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f553])).
fof(f1810,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f1786,f1568])).
fof(f1809,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(backward_demodulation,[],[f1786,f1575])).
fof(f1786,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f400,f551])).
fof(f1808,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f1785,f315])).
fof(f1785,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f400,f1044])).
fof(f1663,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f421,f316])).
fof(f1662,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f421,f553])).
fof(f1661,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f421,f551])).
fof(f1660,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f421,f315])).
fof(f1658,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f421,f314])).
fof(f1656,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f389,f421])).
fof(f1657,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1051,f421])).
fof(f1654,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f420,f316])).
fof(f1653,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f420,f553])).
fof(f1652,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f420,f551])).
fof(f1651,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f420,f315])).
fof(f1641,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f1051,f420])).
fof(f1575,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X8))))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(forward_demodulation,[],[f1574,f400])).
fof(f1574,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(forward_demodulation,[],[f1562,f400])).
fof(f1562,plain,(
  ( ! [X8,X9] : (s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))))) )),
  inference(superposition,[],[f417,f553])).
fof(f1560,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) )),
  inference(superposition,[],[f416,f553])).
fof(f1571,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1554,f1044])).
fof(f1554,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f553,f316])).
fof(f1553,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f553,f1044])).
fof(f1568,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f1567,f400])).
fof(f1567,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f1545,f400])).
fof(f1545,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f553,f417])).
fof(f1434,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f399,f478])).
fof(f1435,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f399,f476])).
fof(f1436,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f399,f478])).
fof(f1437,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f399,f476])).
fof(f1379,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f398,f478])).
fof(f1380,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f398,f476])).
fof(f1381,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f546,f398,f478])).
fof(f1382,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f307,f398,f476])).
fof(f1315,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f572,f568,f573,f546,f461])).
fof(f1281,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f570,f567,f574,f307,f460])).
fof(f1244,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1059,f418])).
fof(f1245,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1060,f418])).
fof(f1223,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))))) | s('type.num.num',X7) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) )),
  inference(forward_demodulation,[],[f1208,f400])).
fof(f1208,plain,(
  ( ! [X6,X7] : (s('type.num.num',X7) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(superposition,[],[f417,f551])).
fof(f1217,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f1202,f400])).
fof(f1202,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f551,f417])).
fof(f1066,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) )),
  inference(superposition,[],[f416,f316])).
fof(f1065,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) | s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f416,f551])).
fof(f1064,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f416,f315])).
fof(f1061,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f416,f314])).
fof(f1059,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f389,f416])).
fof(f1060,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f1051,f416])).
fof(f1051,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f389,f1044])).
fof(f1044,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f551,f316])).
fof(f1007,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f998,f515])).
fof(f998,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f307,f520])).
fof(f987,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f584,f307,f546,f519])).
fof(f956,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f948,f498])).
fof(f948,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f307,f503])).
fof(f942,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f581,f546,f546,f502])).
fof(f914,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f905,f479])).
fof(f905,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f486])).
fof(f894,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f576,f307,f307,f482])).
fof(f890,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f885,f456])).
fof(f885,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f546,f471])).
fof(f881,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f872,f457])).
fof(f872,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f467])).
fof(f855,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f463])).
fof(f811,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f598,f477])).
fof(f815,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f598,f477])).
fof(f795,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f598,f475])).
fof(f797,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f598,f475])).
fof(f729,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f321,f321])).
fof(f679,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f307,f526])).
fof(f671,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f307,f525])).
fof(f669,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f514])).
fof(f668,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f546,f513])).
fof(f660,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f307,f509])).
fof(f656,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f307,f508])).
fof(f653,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f497])).
fof(f649,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f546,f496])).
fof(f648,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f546,f546,f491])).
fof(f642,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f546,f490])).
fof(f628,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f481])).
fof(f625,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f307,f546,f480])).
fof(f620,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f560])).
fof(f617,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f450])).
fof(f612,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f447])).
fof(f609,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f446])).
fof(f608,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f445])).
fof(f584,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f523])).
fof(f583,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f522])).
fof(f582,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f521])).
fof(f581,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f506])).
fof(f580,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f505])).
fof(f579,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f504])).
fof(f578,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f489])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f488])).
fof(f576,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f487])).
fof(f575,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f307,f474])).
fof(f574,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f473])).
fof(f573,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f472])).
fof(f572,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f470])).
fof(f571,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f469])).
fof(f570,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f468])).
fof(f569,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f466])).
fof(f568,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f465])).
fof(f567,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f464])).
fof(f566,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f432])).
fof(f565,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f431])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f183,f200,f199])).
fof(f199,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f200,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) => p(s(bool,X18)))) <=> ((p(s(bool,X16)) | p(s(bool,X17))) & (p(s(bool,X16)) | ~p(s(bool,X18))) & (~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.sat.sat_27_prove_ltb')).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f301])).
fof(f301,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f199])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f302])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f302])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f298])).
fof(f298,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f200])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f299])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f299])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f516])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f157,f197,f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) | p(s(bool,X18)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X17))) & (p(s(bool,X16)) | ~p(s(bool,X18))) & (p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.sat.sat_25_prove_ltb')).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f293])).
fof(f293,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f196])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f294])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f294])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f294])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f290])).
fof(f290,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f291])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f499])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f155,f194,f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f194,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) & p(s(bool,X18)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X17)) | ~p(s(bool,X16))) & (p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.sat.sat_23_prove_ltb')).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f288])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f193])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f286])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f286])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f283])).
fof(f283,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f194])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f284])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f284])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f153,f191,f190,f189,f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f189,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f190,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f191,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> s(bool,X17) = s(bool,X18)) <=> ((p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X18))) & (p(s(bool,X16)) | ~p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X16))) & (p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.sat.sat_2|49|_prove_ltb')).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f188])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f280])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f277])).
fof(f277,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f189])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f278])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f278])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f278])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f274])).
fof(f274,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f190])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f275])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f275])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f275])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f271])).
fof(f271,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f191])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f272])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f272])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f455,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4',aHL_EXT)).
fof(f445,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f266])).
fof(f266,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f446,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f267])).
fof(f447,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f450,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f560,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f430,f546])).
fof(f430,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f431,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f432,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f421,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f254])).
fof(f418,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X8,X13] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X13))) <=> s('type.integer.int',X8) = s('type.integer.int',X13))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_EQ_NEG')).
fof(f419,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f252])).
fof(f417,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f251])).
fof(f409,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f559,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f405,f314])).
fof(f405,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X8,X6] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_MOD_REDUCE|split|4')).
fof(f558,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f404,f314])).
fof(f404,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X8,X6] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_MOD_REDUCE|split|5')).
fof(f557,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f556,f314])).
fof(f556,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f403,f314])).
fof(f403,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_MOD_REDUCE|split|3')).
fof(f401,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.arithmetic.ADD_CLAUSES|split|2')).
fof(f553,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f552,f315])).
fof(f552,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(forward_demodulation,[],[f550,f400])).
fof(f550,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(backward_demodulation,[],[f400,f319])).
fof(f333,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f331,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : (~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X16] : (~s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_MOD|48|')).
fof(f548,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f320,f314])).
fof(f320,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_MOD_REDUCE|split|1')).
fof(f319,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.arithmetic.BIT2')).
fof(f316,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f5682,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl41),
  inference(cnf_transformation,[],[f5682_D])).
fof(f5682_D,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f410,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f250])).
fof(f2305,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) )),
  inference(forward_demodulation,[],[f2296,f547])).
fof(f2296,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(trivial_inequality_removal,[],[f2290])).
fof(f2290,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(superposition,[],[f415,f547])).
fof(f21194,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f20898,f12871])).
fof(f20898,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f555,f13021])).
fof(f21193,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f20897,f12871])).
fof(f20897,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f13021,f555])).
fof(f21191,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21190,f12869])).
fof(f12869,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f12863,f420])).
fof(f21190,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21189,f2347])).
fof(f21189,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20855,f555])).
fof(f20855,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13021,f419])).
fof(f21188,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21187,f12869])).
fof(f21187,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21186,f2347])).
fof(f21186,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21185,f547])).
fof(f21185,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20856,f555])).
fof(f20856,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1060,f13021,f2138])).
fof(f21184,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21183,f12869])).
fof(f21183,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21182,f2347])).
fof(f21182,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21181,f555])).
fof(f21181,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20857,f547])).
fof(f20857,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2399,f13021,f2138])).
fof(f21180,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21179,f12869])).
fof(f21179,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21178,f2347])).
fof(f21178,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21177,f555])).
fof(f21177,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20858,f12899])).
fof(f12899,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f2347,f12871,f2352])).
fof(f20858,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12871,f13021,f2138])).
fof(f21176,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21175,f12869])).
fof(f21175,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21174,f2347])).
fof(f21174,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21173,f555])).
fof(f21173,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20859,f2127])).
fof(f20859,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1059,f13021,f2138])).
fof(f21172,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21171,f12869])).
fof(f21171,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21170,f2347])).
fof(f21170,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21169,f555])).
fof(f21169,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20860,f548])).
fof(f20860,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2412,f13021,f2138])).
fof(f21168,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21167,f12869])).
fof(f21167,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21166,f2347])).
fof(f21166,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21165,f555])).
fof(f21165,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20861,f18991])).
fof(f18991,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | $spl3),
  inference(forward_demodulation,[],[f18958,f13088])).
fof(f13088,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f2347,f12869,f2352])).
fof(f18958,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f12869,f2397,f5361])).
fof(f20861,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2397,f13021,f2138])).
fof(f21164,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21163,f12869])).
fof(f21163,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21162,f2347])).
fof(f21162,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21161,f555])).
fof(f21161,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20862,f13088])).
fof(f20862,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12869,f13021,f2138])).
fof(f21160,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21159,f12869])).
fof(f21159,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21158,f2347])).
fof(f21158,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21157,f555])).
fof(f21157,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20863,f4114])).
fof(f20863,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1640,f13021,f2138])).
fof(f21156,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21155,f12869])).
fof(f21155,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21154,f2347])).
fof(f21154,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21153,f555])).
fof(f21153,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20864,f19272])).
fof(f19272,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f19239,f4114])).
fof(f19239,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f1640,f2410,f5361])).
fof(f20864,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2410,f13021,f2138])).
fof(f21152,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21151,f546])).
fof(f21151,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21150,f15525])).
fof(f21150,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21149,f547])).
fof(f21149,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20865,f555])).
fof(f20865,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1060,f13021,f2294])).
fof(f21148,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21147,f546])).
fof(f21147,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21146,f15525])).
fof(f21146,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21145,f547])).
fof(f21145,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20866,f555])).
fof(f20866,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2399,f13021,f2294])).
fof(f21144,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21143,f546])).
fof(f21143,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21142,f15525])).
fof(f21142,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21141,f12899])).
fof(f21141,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20867,f555])).
fof(f20867,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12871,f13021,f2294])).
fof(f21140,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21139,f546])).
fof(f21139,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21138,f15525])).
fof(f21138,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21137,f2127])).
fof(f21137,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20868,f555])).
fof(f20868,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1059,f13021,f2294])).
fof(f21136,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21135,f546])).
fof(f21135,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21134,f15525])).
fof(f21134,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21133,f548])).
fof(f21133,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20869,f555])).
fof(f20869,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2412,f13021,f2294])).
fof(f21132,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21131,f546])).
fof(f21131,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21130,f15525])).
fof(f21130,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21129,f18991])).
fof(f21129,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20870,f555])).
fof(f20870,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2397,f13021,f2294])).
fof(f21128,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21127,f546])).
fof(f21127,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21126,f15525])).
fof(f21126,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21125,f13088])).
fof(f21125,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20871,f555])).
fof(f20871,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12869,f13021,f2294])).
fof(f21124,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21123,f546])).
fof(f21123,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21122,f15525])).
fof(f21122,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21121,f4114])).
fof(f21121,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20872,f555])).
fof(f20872,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1640,f13021,f2294])).
fof(f21120,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21119,f546])).
fof(f21119,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21118,f15525])).
fof(f21118,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21117,f19272])).
fof(f21117,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20873,f555])).
fof(f20873,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2410,f13021,f2294])).
fof(f21116,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21115,f546])).
fof(f21115,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21114,f15525])).
fof(f21114,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20874,f555])).
fof(f20874,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13021,f2305])).
fof(f21113,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21112,f546])).
fof(f21112,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21111,f15525])).
fof(f21111,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21110,f547])).
fof(f21110,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20875,f555])).
fof(f20875,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f547,f13021,f415])).
fof(f21109,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21108,f546])).
fof(f21108,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21107,f15525])).
fof(f21107,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21106,f548])).
fof(f21106,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20876,f555])).
fof(f20876,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f548,f13021,f415])).
fof(f21105,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21104,f546])).
fof(f21104,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21103,f15525])).
fof(f21103,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21102,f12899])).
fof(f21102,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20877,f555])).
fof(f20877,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12899,f13021,f415])).
fof(f21101,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21100,f21067])).
fof(f21067,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21066,f555])).
fof(f21066,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20893,f398])).
fof(f20893,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12869,f13021,f413])).
fof(f21100,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20878,f555])).
fof(f20878,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13021,f13021,f415])).
fof(f21099,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21098,f546])).
fof(f21098,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21097,f15525])).
fof(f21097,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21096,f13088])).
fof(f21096,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20879,f555])).
fof(f20879,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13088,f13021,f415])).
fof(f21095,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21094,f546])).
fof(f21094,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21093,f15525])).
fof(f21093,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21092,f555])).
fof(f21092,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20880,f399])).
fof(f20880,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2347,f13021,f415])).
fof(f21082,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21081,f21067])).
fof(f21081,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20884,f555])).
fof(f20884,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13021,f13021,f415])).
fof(f21061,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21060,f12869])).
fof(f21060,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21059,f2347])).
fof(f21059,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20896,f555])).
fof(f20896,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13021,f419])).
fof(f20509,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl78),
  inference(cnf_transformation,[],[f20509_D])).
fof(f20509_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5684,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl42),
  inference(cnf_transformation,[],[f5684_D])).
fof(f5684_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f20508,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f20493,f332])).
fof(f20493,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f4970,f321])).
fof(f20505,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl76),
  inference(cnf_transformation,[],[f20505_D])).
fof(f20505_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f20506,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f20489,f332])).
fof(f20489,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f4970,f729])).
fof(f20504,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f20488,f332])).
fof(f20488,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f4970,f729])).
fof(f20473,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl74),
  inference(cnf_transformation,[],[f20473_D])).
fof(f20473_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f20472,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f20457,f332])).
fof(f20457,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f4969,f321])).
fof(f20469,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl72),
  inference(cnf_transformation,[],[f20469_D])).
fof(f20469_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f20470,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f20453,f332])).
fof(f20453,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f4969,f729])).
fof(f20468,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f20452,f332])).
fof(f20452,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f4969,f729])).
fof(f13555,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) | $spl68),
  inference(cnf_transformation,[],[f13555_D])).
fof(f13555_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f13557,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl71),
  inference(cnf_transformation,[],[f13557_D])).
fof(f13557_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f13545,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK16(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2425,f729])).
fof(f2425,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(X3,sK16(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK16(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f455,f321])).
fof(f13544,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2425,f729])).
fof(f13534,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl62),
  inference(cnf_transformation,[],[f13534_D])).
fof(f13534_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13530,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool))) ) | $spl58),
  inference(cnf_transformation,[],[f13530_D])).
fof(f13530_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f13537,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl67),
  inference(cnf_transformation,[],[f13537_D])).
fof(f13537_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f13528,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK16(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2423,f321])).
fof(f2423,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'F_0'))) != s(X5,g(s(bool,'F_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,X5,bool))) )),
  inference(superposition,[],[f455,f321])).
fof(f13535,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl64),
  inference(cnf_transformation,[],[f13535_D])).
fof(f13535_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13527,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK16(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2423,f729])).
fof(f13526,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK16(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2423,f729])).
fof(f13529,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl56),
  inference(cnf_transformation,[],[f13529_D])).
fof(f13529_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13532,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl61),
  inference(cnf_transformation,[],[f13532_D])).
fof(f13532_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f13525,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2423,f729])).
fof(f13524,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2423,f729])).
fof(f13012,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | $spl50),
  inference(cnf_transformation,[],[f13012_D])).
fof(f13012_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f13019,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | $spl54),
  inference(cnf_transformation,[],[f13019_D])).
fof(f13019_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f13017,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f12966,f307])).
fof(f12966,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(superposition,[],[f412,f12872])).
fof(f13014,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | $spl52),
  inference(cnf_transformation,[],[f13014_D])).
fof(f13014_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f13010,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f12964,f307])).
fof(f12964,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(superposition,[],[f2784,f12872])).
fof(f12867,plain,(
  $false | ($spl0 | $spl40)),
  inference(subsumption_resolution,[],[f7014,f12865])).
fof(f12865,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f12864,f314])).
fof(f12864,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f542,f314])).
fof(f7014,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl40),
  inference(unit_resulting_resolution,[],[f307,f5683,f446])).
fof(f12866,plain,(
  $false | ($spl0 | $spl44)),
  inference(subsumption_resolution,[],[f12865,f7016])).
fof(f7016,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl44),
  inference(unit_resulting_resolution,[],[f307,f5718,f446])).
fof(f12861,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f12860,f546])).
fof(f12860,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f12822,f605])).
fof(f605,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f337])).
fof(f564,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f563,f314])).
fof(f563,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f543,f314])).
fof(f12822,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl2),
  inference(trivial_inequality_removal,[],[f12821])).
fof(f12821,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl2),
  inference(superposition,[],[f2871,f562])).
fof(f562,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | $spl2),
  inference(forward_demodulation,[],[f561,f314])).
fof(f561,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | $spl2),
  inference(forward_demodulation,[],[f544,f314])).
fof(f12662,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl48),
  inference(cnf_transformation,[],[f12662_D])).
fof(f12662_D,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12663,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s('type.integer.int',X15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))))) )),
  inference(subsumption_resolution,[],[f12588,f332])).
fof(f12588,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s('type.integer.int',X15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))))) )),
  inference(superposition,[],[f2864,f729])).
fof(f12661,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13)))))) )),
  inference(subsumption_resolution,[],[f12587,f332])).
fof(f12587,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2864,f729])).
fof(f8656,plain,(
  ( ! [X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl46),
  inference(cnf_transformation,[],[f8656_D])).
fof(f8656_D,plain,(
  ( ! [X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f8657,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f8646,f332])).
fof(f8646,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f2861,f729])).
fof(f8655,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f8645,f332])).
fof(f8645,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2861,f729])).
fof(f6814,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5722,f5734])).
fof(f5734,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f5684,f409])).
fof(f5722,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f336])).
fof(f6813,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5725,f307])).
fof(f5725,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f598,f5684,f409])).
fof(f6812,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5727,f307])).
fof(f5727,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f599,f5684,f409])).
fof(f599,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f332])).
fof(f6811,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5731,f5684])).
fof(f5731,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f599,f5684,f409])).
fof(f6810,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5735,f307])).
fof(f5735,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f598,f5684,f409])).
fof(f6809,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5737,f307])).
fof(f5737,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f599,f5684,f409])).
fof(f6808,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5741,f5684])).
fof(f5741,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f599,f5684,f409])).
fof(f6807,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5748,f307])).
fof(f5748,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f617,f5684,f446])).
fof(f6806,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5749,f5684])).
fof(f5749,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3454,f5684,f446])).
fof(f6805,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5755,f307])).
fof(f5755,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f612,f5684,f446])).
fof(f6804,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5758,f5684])).
fof(f5758,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3453,f5684,f446])).
fof(f6803,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5775,f307])).
fof(f5775,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f797,f546,f5684,f459])).
fof(f6802,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5776,f5684])).
fof(f5776,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f802,f546,f5684,f459])).
fof(f802,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f605,f476])).
fof(f6801,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5777,f5684])).
fof(f5777,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1722,f546,f5684,f459])).
fof(f1722,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f1671,f476])).
fof(f1671,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f599,f1519])).
fof(f1519,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X4),s('type.integer.int',X5)))) ) | $spl1),
  inference(forward_demodulation,[],[f1484,f398])).
fof(f1484,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X5)))) ) | $spl1),
  inference(superposition,[],[f1414,f399])).
fof(f1414,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) ) | $spl1),
  inference(forward_demodulation,[],[f1383,f605])).
fof(f1383,plain,(
  ( ! [X2,X3] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int',X3))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3)))))) ) | $spl1),
  inference(superposition,[],[f398,f1037])).
fof(f1037,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f1034,f546])).
fof(f1034,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f780,f605])).
fof(f6800,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5778,f5684])).
fof(f5778,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2231,f546,f5684,f459])).
fof(f2231,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2185,f476])).
fof(f2185,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1671,f2064])).
fof(f2064,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X8),s('type.integer.int',X9))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))),s('type.integer.int',X9)))) ) | $spl1),
  inference(forward_demodulation,[],[f2049,f398])).
fof(f2049,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))),s('type.integer.int',X9)))) ) | $spl1),
  inference(superposition,[],[f1522,f399])).
fof(f1522,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int',X11)))) ) | $spl1),
  inference(subsumption_resolution,[],[f1492,f307])).
fof(f1492,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int',X11)))) ) | $spl1),
  inference(superposition,[],[f336,f1414])).
fof(f6799,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5779,f5684])).
fof(f5779,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3253,f546,f5684,f459])).
fof(f3253,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2228,f476])).
fof(f2228,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2185,f2064])).
fof(f6798,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5781,f5684])).
fof(f5781,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3486,f546,f5684,f459])).
fof(f3486,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f605,f2298,f476])).
fof(f6797,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5813,f307])).
fof(f5813,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f802,f546,f5684,f459])).
fof(f6796,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5814,f307])).
fof(f5814,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1722,f546,f5684,f459])).
fof(f6795,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5815,f307])).
fof(f5815,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2231,f546,f5684,f459])).
fof(f6794,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5816,f307])).
fof(f5816,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3253,f546,f5684,f459])).
fof(f6793,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5817,f307])).
fof(f5817,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f1315,f546,f5684,f459])).
fof(f6792,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5818,f5684])).
fof(f5818,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3486,f546,f5684,f459])).
fof(f6791,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5860,f6078])).
fof(f6078,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f474])).
fof(f5860,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f570,f567,f5684,f460])).
fof(f6790,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5861,f6078])).
fof(f5861,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f570,f567,f5684,f460])).
fof(f816,plain,(
  ~sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f599,f477])).
fof(f6789,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5862,f6078])).
fof(f5862,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f571,f567,f5684,f460])).
fof(f6788,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5863,f6078])).
fof(f5863,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f587,f567,f5684,f460])).
fof(f587,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f469])).
fof(f6787,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5864,f6078])).
fof(f5864,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3499,f587,f567,f5684,f460])).
fof(f3499,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f599,f2298,f477])).
fof(f6786,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5865,f6078])).
fof(f5865,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f570,f3471,f5684,f460])).
fof(f6785,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5866,f6078])).
fof(f5866,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f571,f3471,f5684,f460])).
fof(f6784,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5867,f6078])).
fof(f5867,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f571,f568,f5684,f460])).
fof(f6783,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5868,f6078])).
fof(f5868,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f571,f568,f5684,f460])).
fof(f6782,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5874,f6078])).
fof(f5874,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f570,f568,f5684,f460])).
fof(f6781,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5875,f6078])).
fof(f5875,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f570,f568,f5684,f460])).
fof(f6780,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5878,f6078])).
fof(f5878,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f587,f585,f5684,f460])).
fof(f585,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f465])).
fof(f6779,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5882,f6078])).
fof(f5882,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3499,f587,f585,f5684,f460])).
fof(f6778,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5883,f6078])).
fof(f5883,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f570,f585,f5684,f460])).
fof(f6777,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5902,f6038])).
fof(f6038,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f464])).
fof(f5902,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f570,f573,f5684,f460])).
fof(f6776,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5903,f6038])).
fof(f5903,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f570,f573,f5684,f460])).
fof(f6775,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5905,f6038])).
fof(f5905,plain,(
  ~sP4('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f570,f573,f5684,f460])).
fof(f6774,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5906,f6038])).
fof(f5906,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f571,f573,f5684,f460])).
fof(f6773,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5907,f6038])).
fof(f5907,plain,(
  ~sP4('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f571,f573,f5684,f460])).
fof(f6772,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5908,f6038])).
fof(f5908,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f587,f573,f5684,f460])).
fof(f6771,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5914,f6072])).
fof(f6072,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f468])).
fof(f5914,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f567,f573,f5684,f460])).
fof(f6770,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5915,f6072])).
fof(f5915,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f567,f573,f5684,f460])).
fof(f6769,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5922,f6072])).
fof(f5922,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f3471,f573,f5684,f460])).
fof(f6768,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5928,f6072])).
fof(f5928,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f568,f573,f5684,f460])).
fof(f6767,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5929,f6072])).
fof(f5929,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f568,f573,f5684,f460])).
fof(f6766,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5936,f6072])).
fof(f5936,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f585,f573,f5684,f460])).
fof(f6765,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5943,f6038])).
fof(f5943,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f570,f574,f5684,f460])).
fof(f6764,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5944,f6038])).
fof(f5944,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f570,f574,f5684,f460])).
fof(f6763,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5945,f6038])).
fof(f5945,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f571,f574,f5684,f460])).
fof(f6762,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5946,f6038])).
fof(f5946,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f587,f574,f5684,f460])).
fof(f6761,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5947,f6038])).
fof(f5947,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3499,f587,f574,f5684,f460])).
fof(f6760,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5948,f6072])).
fof(f5948,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f567,f574,f5684,f460])).
fof(f6759,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5949,f6072])).
fof(f5949,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f567,f574,f5684,f460])).
fof(f6758,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5951,f6072])).
fof(f5951,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3499,f567,f574,f5684,f460])).
fof(f6757,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5954,f6123])).
fof(f6123,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f598,f5684,f477])).
fof(f5954,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f571,f567,f574,f5684,f460])).
fof(f6756,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5956,f6125])).
fof(f6125,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f599,f5684,f477])).
fof(f5956,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f567,f574,f5684,f460])).
fof(f6755,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5958,f6072])).
fof(f5958,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f815,f568,f574,f5684,f460])).
fof(f6754,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5959,f6123])).
fof(f5959,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f571,f568,f574,f5684,f460])).
fof(f6753,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5961,f815])).
fof(f5961,plain,(
  sP5('T_0','T_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f570,f568,f574,f5684,f460])).
fof(f6752,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5963,f6072])).
fof(f5963,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f585,f574,f5684,f460])).
fof(f6751,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5964,f6072])).
fof(f5964,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3499,f585,f574,f5684,f460])).
fof(f6750,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5965,f6125])).
fof(f5965,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f585,f574,f5684,f460])).
fof(f6749,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5967,f816])).
fof(f5967,plain,(
  sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f570,f585,f574,f5684,f460])).
fof(f6748,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5972,f6123])).
fof(f5972,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','F_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f571,f586,f574,f5684,f460])).
fof(f586,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f466])).
fof(f6747,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5973,f6125])).
fof(f5973,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f586,f574,f5684,f460])).
fof(f6746,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5988,f5684])).
fof(f5988,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f585,f546,f5684,f463])).
fof(f6745,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5989,f5684])).
fof(f5989,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1912,f546,f5684,f463])).
fof(f1912,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1724,f458])).
fof(f1724,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f1671,f476])).
fof(f6744,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5990,f5684])).
fof(f5990,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2453,f546,f5684,f463])).
fof(f2453,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f2233,f458])).
fof(f2233,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2185,f476])).
fof(f6743,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5991,f5684])).
fof(f5991,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3430,f546,f5684,f463])).
fof(f3430,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f3255,f458])).
fof(f3255,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2228,f476])).
fof(f6742,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5992,f5684])).
fof(f5992,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f569,f546,f5684,f463])).
fof(f6741,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5993,f5684])).
fof(f5993,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f586,f546,f5684,f463])).
fof(f6740,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5998,f6038])).
fof(f5998,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5684,f463])).
fof(f6739,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6009,f6038])).
fof(f6009,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f5684,f463])).
fof(f6738,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6011,f5684])).
fof(f6011,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f568,f546,f5684,f463])).
fof(f6737,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6012,f5684])).
fof(f6012,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f585,f546,f5684,f463])).
fof(f6736,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6013,f307])).
fof(f6013,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1912,f546,f5684,f463])).
fof(f6735,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6014,f307])).
fof(f6014,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2453,f546,f5684,f463])).
fof(f6734,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6015,f307])).
fof(f6015,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3430,f546,f5684,f463])).
fof(f6733,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6017,f5684])).
fof(f6017,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f586,f546,f5684,f463])).
fof(f6732,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6025,f6038])).
fof(f6025,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5684,f463])).
fof(f6731,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6036,f6038])).
fof(f6036,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f5684,f463])).
fof(f6730,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6041,f5684])).
fof(f6041,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f546,f5684,f467])).
fof(f6729,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6042,f5684])).
fof(f6042,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f572,f546,f5684,f467])).
fof(f6728,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6043,f5684])).
fof(f6043,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f546,f5684,f467])).
fof(f588,plain,(
  ( ! [X0,X1] : (sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f470])).
fof(f6727,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6047,f6072])).
fof(f6047,plain,(
  ( ! [X0] : (~sP3(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5684,f467])).
fof(f6726,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6054,f6072])).
fof(f6054,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f5684,f467])).
fof(f6725,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6056,f5684])).
fof(f6056,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f571,f546,f5684,f467])).
fof(f6724,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6057,f5684])).
fof(f6057,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f546,f5684,f467])).
fof(f6723,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6059,f5684])).
fof(f6059,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f546,f5684,f467])).
fof(f6722,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6063,f6072])).
fof(f6063,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5684,f467])).
fof(f6721,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6070,f6072])).
fof(f6070,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f5684,f467])).
fof(f6719,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6087,f605])).
fof(f6087,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f825,f5684,f476])).
fof(f825,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f605,f478])).
fof(f6718,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6088,f1671])).
fof(f6088,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1723,f5684,f476])).
fof(f1723,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f1671,f478])).
fof(f6717,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6089,f2185])).
fof(f6089,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2232,f5684,f476])).
fof(f2232,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2185,f478])).
fof(f6716,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6090,f2228])).
fof(f6090,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3254,f5684,f476])).
fof(f3254,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2228,f478])).
fof(f6715,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6091,f605])).
fof(f6091,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f824,f5684,f476])).
fof(f824,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f605,f478])).
fof(f6714,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6092,f1671])).
fof(f6092,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1721,f5684,f476])).
fof(f1721,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f1671,f478])).
fof(f6713,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6093,f2185])).
fof(f6093,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2230,f5684,f476])).
fof(f2230,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2185,f478])).
fof(f6712,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6094,f2228])).
fof(f6094,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3252,f5684,f476])).
fof(f3252,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2228,f478])).
fof(f6711,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6095,f398])).
fof(f6095,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1379,f5684,f476])).
fof(f6710,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6096,f398])).
fof(f6096,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1381,f5684,f476])).
fof(f6709,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6097,f399])).
fof(f6097,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1434,f5684,f476])).
fof(f6708,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6098,f399])).
fof(f6098,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1436,f5684,f476])).
fof(f6707,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6101,f598])).
fof(f6101,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f795,f5684,f477])).
fof(f6706,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6102,f598])).
fof(f6102,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f797,f5684,f477])).
fof(f6705,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6103,f599])).
fof(f6103,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f798,f5684,f477])).
fof(f798,plain,(
  sP5('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f599,f475])).
fof(f6704,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6104,f599])).
fof(f6104,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f796,f5684,f477])).
fof(f796,plain,(
  sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f599,f475])).
fof(f6703,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6120,f598])).
fof(f6120,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1273,f5684,f477])).
fof(f1273,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f572,f586,f575,f307,f460])).
fof(f6702,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6136,f6273])).
fof(f6273,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f5684,f5684,f492])).
fof(f6136,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f480])).
fof(f6701,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6143,f6273])).
fof(f6143,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f481])).
fof(f6700,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6146,f5684])).
fof(f6146,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f578,f628,f5684,f482])).
fof(f6699,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6147,f5684])).
fof(f6147,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f576,f628,f5684,f482])).
fof(f6698,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6149,f5684])).
fof(f6149,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f628,f5684,f482])).
fof(f589,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f488])).
fof(f6697,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6153,f6219])).
fof(f6219,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f487])).
fof(f6153,plain,(
  ~sP6('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f628,f307,f5684,f482])).
fof(f6696,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6161,f6219])).
fof(f6161,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f482])).
fof(f6695,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6163,f6219])).
fof(f6163,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f628,f5684,f5684,f482])).
fof(f6694,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6166,f5684])).
fof(f6166,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f577,f625,f5684,f482])).
fof(f6693,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6167,f5684])).
fof(f6167,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f576,f625,f5684,f482])).
fof(f6692,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6169,f5684])).
fof(f6169,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f590,f625,f5684,f482])).
fof(f590,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f489])).
fof(f6691,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6174,f6219])).
fof(f6174,plain,(
  ~sP6('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f625,f307,f5684,f482])).
fof(f6690,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6182,f6219])).
fof(f6182,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f482])).
fof(f6689,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6184,f6219])).
fof(f6184,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f628,f5684,f5684,f482])).
fof(f6688,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6188,f5684])).
fof(f6188,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f546,f5684,f486])).
fof(f6687,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6189,f5684])).
fof(f6189,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f578,f546,f5684,f486])).
fof(f6686,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6190,f5684])).
fof(f6190,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f590,f546,f5684,f486])).
fof(f6685,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6194,f6219])).
fof(f6194,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5684,f486])).
fof(f6684,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6201,f6219])).
fof(f6201,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f5684,f486])).
fof(f6683,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6203,f5684])).
fof(f6203,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f577,f546,f5684,f486])).
fof(f6682,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6204,f5684])).
fof(f6204,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f546,f5684,f486])).
fof(f6681,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6206,f5684])).
fof(f6206,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f590,f546,f5684,f486])).
fof(f6680,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6210,f6219])).
fof(f6210,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5684,f486])).
fof(f6679,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6217,f6219])).
fof(f6217,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f5684,f486])).
fof(f6678,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6222,f5684])).
fof(f6222,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f628,f307,f5684,f492])).
fof(f6677,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6224,f5684])).
fof(f6224,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f628,f5684,f5684,f492])).
fof(f6676,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6225,f307])).
fof(f6225,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f914,f5684,f5684,f492])).
fof(f6675,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6226,f307])).
fof(f6226,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f628,f307,f5684,f492])).
fof(f6674,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6232,f307])).
fof(f6232,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f492])).
fof(f6673,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6233,f307])).
fof(f6233,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f628,f5684,f5684,f492])).
fof(f6672,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6239,f5684])).
fof(f6239,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f625,f307,f5684,f492])).
fof(f6671,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6241,f5684])).
fof(f6241,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f492])).
fof(f6670,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6243,f307])).
fof(f6243,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f914,f5684,f5684,f492])).
fof(f6669,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6244,f307])).
fof(f6244,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f625,f307,f5684,f492])).
fof(f6668,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6250,f307])).
fof(f6250,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f492])).
fof(f6667,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6251,f307])).
fof(f6251,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f628,f5684,f5684,f492])).
fof(f6666,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6258,f307])).
fof(f6258,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f914,f307,f5684,f492])).
fof(f6665,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6259,f5684])).
fof(f6259,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f492])).
fof(f6664,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6261,f307])).
fof(f6261,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f914,f5684,f5684,f492])).
fof(f6663,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6263,f307])).
fof(f6263,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f914,f307,f5684,f492])).
fof(f6662,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6269,f5684])).
fof(f6269,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f628,f5684,f5684,f492])).
fof(f6661,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6270,f307])).
fof(f6270,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f914,f5684,f5684,f492])).
fof(f6660,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6276,f5684])).
fof(f6276,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f707,f546,f5684,f495])).
fof(f707,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f589,f546,f532])).
fof(f6659,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6277,f5684])).
fof(f6277,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f648,f546,f5684,f495])).
fof(f6658,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6278,f5684])).
fof(f6278,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f709,f546,f5684,f495])).
fof(f709,plain,(
  ( ! [X0] : (sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f590,f546,f532])).
fof(f6657,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6291,f5684])).
fof(f6291,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f642,f546,f5684,f495])).
fof(f6656,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6292,f5684])).
fof(f6292,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f707,f546,f5684,f495])).
fof(f6655,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6294,f5684])).
fof(f6294,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f709,f546,f5684,f495])).
fof(f6654,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6311,f6587])).
fof(f6587,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f581,f5684,f531])).
fof(f6311,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f496])).
fof(f6653,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6317,f6587])).
fof(f6317,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f497])).
fof(f6652,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6326,f6351])).
fof(f6351,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f505])).
fof(f6326,plain,(
  ~sP8('F_0','T_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f653,f546,f5684,f500])).
fof(f6651,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6334,f6351])).
fof(f6334,plain,(
  ~sP8('F_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f649,f546,f5684,f501])).
fof(f6650,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6347,f6351])).
fof(f6347,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f546,f5684,f503])).
fof(f6649,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6355,f307])).
fof(f6355,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f956,f5684,f508])).
fof(f6648,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6360,f307])).
fof(f6360,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f649,f5684,f508])).
fof(f6647,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6361,f5684])).
fof(f6361,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f653,f5684,f508])).
fof(f6646,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6369,f307])).
fof(f6369,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f956,f5684,f509])).
fof(f6645,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6374,f5684])).
fof(f6374,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f649,f5684,f509])).
fof(f6644,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6375,f307])).
fof(f6375,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f653,f5684,f509])).
fof(f6643,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6396,f6577])).
fof(f6577,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f584,f5684,f530])).
fof(f6396,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684,f514])).
fof(f6642,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6399,f5684])).
fof(f6399,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f584,f669,f5684,f517])).
fof(f6641,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6400,f5684])).
fof(f6400,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f583,f669,f5684,f517])).
fof(f6640,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6402,f5684])).
fof(f6402,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f592,f669,f5684,f517])).
fof(f592,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f521])).
fof(f6639,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6406,f6501])).
fof(f6501,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f522])).
fof(f6406,plain,(
  ~sP10('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f669,f307,f5684,f517])).
fof(f6638,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6414,f6501])).
fof(f6414,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f668,f5684,f5684,f517])).
fof(f6637,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6416,f6501])).
fof(f6416,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5684,f5684,f517])).
fof(f6636,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6429,f6501])).
fof(f6429,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f668,f5684,f5684,f517])).
fof(f6635,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6430,f6501])).
fof(f6430,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5684,f5684,f517])).
fof(f6634,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6447,f6501])).
fof(f6447,plain,(
  ~sP10('F_0','F_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f668,f546,f5684,f518])).
fof(f6633,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6461,f6501])).
fof(f6461,plain,(
  ~sP10('F_0','F_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f668,f546,f5684,f519])).
fof(f6632,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6478,f5684])).
fof(f6478,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f592,f546,f5684,f520])).
fof(f6631,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6479,f5684])).
fof(f6479,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f584,f546,f5684,f520])).
fof(f6630,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6480,f5684])).
fof(f6480,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f593,f546,f5684,f520])).
fof(f593,plain,(
  ( ! [X0,X1] : (sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f564,f523])).
fof(f6629,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6481,f6501])).
fof(f6481,plain,(
  ( ! [X0] : (~sP10(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f307,f546,f5684,f520])).
fof(f6628,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6483,f6501])).
fof(f6483,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f546,f5684,f520])).
fof(f6627,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6493,f5684])).
fof(f6493,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f582,f546,f5684,f520])).
fof(f6626,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6494,f5684])).
fof(f6494,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f592,f546,f5684,f520])).
fof(f6625,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6496,f5684])).
fof(f6496,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f593,f546,f5684,f520])).
fof(f6624,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6497,f6501])).
fof(f6497,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f307,f546,f5684,f520])).
fof(f6623,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6499,f6501])).
fof(f6499,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f546,f5684,f520])).
fof(f6622,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6516,f307])).
fof(f6516,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f1007,f5684,f526])).
fof(f6621,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6521,f5684])).
fof(f6521,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f668,f5684,f526])).
fof(f6620,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6522,f307])).
fof(f6522,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5684,f526])).
fof(f6619,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6528,f307])).
fof(f6528,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f671,f546,f5684,f527])).
fof(f6618,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6529,f307])).
fof(f6529,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f691,f546,f5684,f527])).
fof(f691,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f592,f307,f530])).
fof(f6617,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6531,f5684])).
fof(f6531,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f988,f546,f5684,f527])).
fof(f988,plain,(
  sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f593,f307,f546,f519])).
fof(f6616,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6532,f5684])).
fof(f6532,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3579,f546,f5684,f527])).
fof(f3579,plain,(
  sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f593,f2298,f519])).
fof(f6615,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6533,f5684])).
fof(f6533,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3605,f546,f5684,f527])).
fof(f3605,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f592,f2298,f530])).
fof(f6614,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6549,f5684])).
fof(f6549,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f691,f546,f5684,f527])).
fof(f6613,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6550,f307])).
fof(f6550,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f987,f546,f5684,f527])).
fof(f6612,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6551,f307])).
fof(f6551,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f988,f546,f5684,f527])).
fof(f6611,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6552,f5684])).
fof(f6552,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3579,f546,f5684,f527])).
fof(f6610,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6553,f5684])).
fof(f6553,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3605,f546,f5684,f527])).
fof(f6609,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6568,f582])).
fof(f6568,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f668,f5684,f530])).
fof(f6608,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6569,f583])).
fof(f6569,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5684,f530])).
fof(f6607,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6580,f579])).
fof(f6580,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f649,f5684,f531])).
fof(f6606,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6581,f580])).
fof(f6581,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f653,f5684,f531])).
fof(f6588,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f581,f5684,f531])).
fof(f6586,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f653,f580,f5684,f531])).
fof(f6584,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f579,f5684,f531])).
fof(f6578,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f584,f5684,f530])).
fof(f6576,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f583,f5684,f530])).
fof(f6574,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f669,f592,f5684,f530])).
fof(f6572,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f582,f5684,f530])).
fof(f6567,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f671,f546,f5684,f5684,f527])).
fof(f6558,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f987,f546,f307,f5684,f527])).
fof(f6547,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f671,f546,f5684,f5684,f527])).
fof(f6538,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f671,f546,f307,f5684,f527])).
fof(f6527,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f5684,f5684,f526])).
fof(f6525,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f307,f5684,f526])).
fof(f6520,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f5684,f5684,f526])).
fof(f6518,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f1007,f307,f5684,f526])).
fof(f6513,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f546,f5684,f525])).
fof(f6508,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f546,f546,f5684,f524])).
fof(f6502,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f998,f5684,f522])).
fof(f6500,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f582,f5684,f546,f5684,f520])).
fof(f6498,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f582,f307,f546,f5684,f520])).
fof(f6484,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f582,f5684,f546,f5684,f520])).
fof(f6482,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f584,f307,f546,f5684,f520])).
fof(f6467,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f584,f546,f5684,f519])).
fof(f6464,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f582,f546,f5684,f519])).
fof(f6453,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f584,f546,f5684,f518])).
fof(f6450,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f582,f546,f5684,f518])).
fof(f6439,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f584,f5684,f5684,f517])).
fof(f6437,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f583,f5684,f5684,f517])).
fof(f6435,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f669,f592,f5684,f5684,f517])).
fof(f6433,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f668,f582,f5684,f5684,f517])).
fof(f6417,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f584,f669,f5684,f5684,f517])).
fof(f6415,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f584,f668,f5684,f5684,f517])).
fof(f6407,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f584,f669,f307,f5684,f517])).
fof(f6397,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f987,f546,f5684,f514])).
fof(f6388,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f942,f546,f546,f5684,f510])).
fof(f6380,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f5684,f5684,f509])).
fof(f6378,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f653,f307,f5684,f509])).
fof(f6373,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f5684,f5684,f509])).
fof(f6371,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f956,f307,f5684,f509])).
fof(f6366,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f5684,f5684,f508])).
fof(f6364,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f307,f5684,f508])).
fof(f6359,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f5684,f5684,f508])).
fof(f6357,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f956,f307,f5684,f508])).
fof(f6352,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f948,f5684,f505])).
fof(f6350,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f948,f5684,f504])).
fof(f6348,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f581,f546,f546,f5684,f503])).
fof(f6339,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f581,f546,f5684,f501])).
fof(f6337,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f579,f546,f5684,f501])).
fof(f6327,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f581,f653,f546,f5684,f500])).
fof(f6318,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f942,f546,f5684,f497])).
fof(f6312,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f942,f546,f5684,f496])).
fof(f6306,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f642,f546,f5684,f5684,f495])).
fof(f6299,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f642,f546,f307,f5684,f495])).
fof(f6290,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f642,f546,f5684,f5684,f495])).
fof(f6283,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f648,f546,f307,f5684,f495])).
fof(f6274,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f5684,f492])).
fof(f6272,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f628,f307,f5684,f5684,f492])).
fof(f6267,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f307,f5684,f492])).
fof(f6265,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f914,f307,f307,f5684,f492])).
fof(f6256,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f5684,f492])).
fof(f6254,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f625,f307,f5684,f5684,f492])).
fof(f6249,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f307,f5684,f492])).
fof(f6247,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f625,f307,f307,f5684,f492])).
fof(f6238,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f625,f5684,f5684,f5684,f492])).
fof(f6236,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f625,f307,f5684,f5684,f492])).
fof(f6231,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f628,f5684,f307,f5684,f492])).
fof(f6229,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f628,f307,f307,f5684,f492])).
fof(f6220,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f905,f5684,f487])).
fof(f6218,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f577,f546,f5684,f5684,f486])).
fof(f6211,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f577,f546,f307,f5684,f486])).
fof(f6202,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f577,f546,f5684,f5684,f486])).
fof(f6195,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f546,f307,f5684,f486])).
fof(f6185,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f628,f5684,f5684,f482])).
fof(f6183,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f577,f625,f5684,f5684,f482])).
fof(f6175,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f577,f625,f307,f5684,f482])).
fof(f6164,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f628,f5684,f5684,f482])).
fof(f6162,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f577,f625,f5684,f5684,f482])).
fof(f6154,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f628,f307,f5684,f482])).
fof(f6144,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f648,f546,f5684,f481])).
fof(f6137,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f642,f546,f5684,f480])).
fof(f6130,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f796,f599,f5684,f477])).
fof(f6128,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f795,f598,f5684,f477])).
fof(f6126,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f798,f599,f5684,f477])).
fof(f6124,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f797,f598,f5684,f477])).
fof(f6605,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f6086])).
fof(f6086,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f890,f5684,f476])).
fof(f6604,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f6085])).
fof(f6085,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f881,f5684,f476])).
fof(f6079,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f885,f5684,f474])).
fof(f6077,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f885,f5684,f473])).
fof(f6075,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f885,f5684,f472])).
fof(f6073,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f872,f5684,f468])).
fof(f6071,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f571,f546,f5684,f5684,f467])).
fof(f6064,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f571,f546,f307,f5684,f467])).
fof(f6055,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f571,f546,f5684,f5684,f467])).
fof(f6048,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f572,f546,f307,f5684,f467])).
fof(f6039,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f855,f5684,f464])).
fof(f6037,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f568,f546,f5684,f5684,f463])).
fof(f6026,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1912,f546,f307,f5684,f463])).
fof(f6010,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f568,f546,f5684,f5684,f463])).
fof(f5999,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f569,f546,f307,f5684,f463])).
fof(f5968,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f570,f585,f574,f5684,f460])).
fof(f5966,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f587,f585,f574,f5684,f460])).
fof(f5962,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f815,f570,f568,f574,f5684,f460])).
fof(f5960,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f815,f571,f568,f574,f5684,f460])).
fof(f5957,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f587,f567,f574,f5684,f460])).
fof(f5955,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f815,f571,f567,f574,f5684,f460])).
fof(f5953,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f815,f570,f567,f574,f5684,f460])).
fof(f5940,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f570,f585,f573,f5684,f460])).
fof(f5938,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f587,f585,f573,f5684,f460])).
fof(f5933,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f815,f570,f568,f573,f5684,f460])).
fof(f5931,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f815,f571,f568,f573,f5684,f460])).
fof(f5926,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f571,f3471,f573,f5684,f460])).
fof(f5924,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f3457,f570,f3471,f573,f5684,f460])).
fof(f5921,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f816,f587,f567,f573,f5684,f460])).
fof(f5919,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f815,f571,f567,f573,f5684,f460])).
fof(f5917,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f815,f570,f567,f573,f5684,f460])).
fof(f5859,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f797,f546,f5684,f5684,f459])).
fof(f5835,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f802,f546,f307,f5684,f459])).
fof(f5811,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f797,f546,f5684,f5684,f459])).
fof(f5787,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f797,f546,f307,f5684,f459])).
fof(f5774,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f608,f546,f5684,f450])).
fof(f5768,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f608,f546,f5684,f447])).
fof(f5762,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f612,f5684,f5684,f446])).
fof(f5760,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f612,f307,f5684,f446])).
fof(f5754,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f612,f5684,f5684,f446])).
fof(f5752,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f617,f307,f5684,f446])).
fof(f5746,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f620,f5684,f432])).
fof(f5744,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f620,f5684,f431])).
fof(f5742,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f5684,f599,f5684,f409])).
fof(f5740,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f598,f5684,f409])).
fof(f5738,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f307,f599,f5684,f409])).
fof(f5736,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f307,f598,f5684,f409])).
fof(f5732,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f5684,f599,f5684,f409])).
fof(f5730,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f5684,f598,f5684,f409])).
fof(f5728,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f307,f599,f5684,f409])).
fof(f5726,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f307,f598,f5684,f409])).
fof(f5721,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f598,f5684,f333])).
fof(f5719,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5684])).
fof(f5718,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | $spl44),
  inference(cnf_transformation,[],[f5718_D])).
fof(f5718_D,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5716,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0')) ) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f5711,f332])).
fof(f5711,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl8)),
  inference(superposition,[],[f4916,f1037])).
fof(f5683,plain,(
  sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl40),
  inference(cnf_transformation,[],[f5683_D])).
fof(f5683_D,plain,(
  sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5681,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5676,f332])).
fof(f5676,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl4)),
  inference(superposition,[],[f4912,f1037])).
fof(f5358,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5355,f5357])).
fof(f5357,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5237,f307])).
fof(f5237,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'T_0'))) ) | $spl6),
  inference(condensation,[],[f5224])).
fof(f5224,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f2218])).
fof(f5355,plain,(
  ( ! [X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5222,f307])).
fof(f5222,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f2225])).
fof(f5349,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5210,f307])).
fof(f5210,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f2701])).
fof(f5347,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5205,f307])).
fof(f5205,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f2784])).
fof(f5345,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5204,f307])).
fof(f5204,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f2781])).
fof(f5343,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5203,f307])).
fof(f5203,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f2710])).
fof(f5341,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5202,f307])).
fof(f5202,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f2707])).
fof(f5339,plain,(
  ( ! [X92,X90,X91,X89] : (~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92)) ) | $spl6),
  inference(subsumption_resolution,[],[f5201,f307])).
fof(f5201,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1468])).
fof(f5337,plain,(
  ( ! [X88,X87,X85,X86] : (~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88)) ) | $spl6),
  inference(subsumption_resolution,[],[f5200,f307])).
fof(f5200,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1467])).
fof(f5335,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X83)) | p(s(bool,X84)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83)) ) | $spl6),
  inference(subsumption_resolution,[],[f5199,f307])).
fof(f5199,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1466])).
fof(f5333,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X79)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80)) ) | $spl6),
  inference(subsumption_resolution,[],[f5198,f307])).
fof(f5198,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X79)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1465])).
fof(f5331,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X75)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75)) ) | $spl6),
  inference(subsumption_resolution,[],[f5197,f307])).
fof(f5197,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X75)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1464])).
fof(f5329,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5196,f307])).
fof(f5196,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70)))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1463])).
fof(f5327,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X67)) | p(s(bool,X68)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67)) ) | $spl6),
  inference(subsumption_resolution,[],[f5195,f307])).
fof(f5195,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1462])).
fof(f5325,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X63)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64)) ) | $spl6),
  inference(subsumption_resolution,[],[f5194,f307])).
fof(f5194,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1461])).
fof(f5323,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X59)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59)) ) | $spl6),
  inference(subsumption_resolution,[],[f5193,f307])).
fof(f5193,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X59)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1460])).
fof(f5321,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,X55)) | p(s(bool,X56)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56)) ) | $spl6),
  inference(subsumption_resolution,[],[f5192,f307])).
fof(f5192,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1459])).
fof(f5319,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5191,f307])).
fof(f5191,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50)))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1458])).
fof(f5317,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48)) ) | $spl6),
  inference(subsumption_resolution,[],[f5190,f307])).
fof(f5190,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1457])).
fof(f5315,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44)) ) | $spl6),
  inference(subsumption_resolution,[],[f5189,f307])).
fof(f5189,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1456])).
fof(f5313,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40)) ) | $spl6),
  inference(subsumption_resolution,[],[f5188,f307])).
fof(f5188,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1455])).
fof(f5311,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36)) ) | $spl6),
  inference(subsumption_resolution,[],[f5187,f307])).
fof(f5187,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1454])).
fof(f5309,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X32)) | sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5186,f307])).
fof(f5186,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31)))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1453])).
fof(f5307,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X29)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29)) ) | $spl6),
  inference(subsumption_resolution,[],[f5185,f307])).
fof(f5185,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1452])).
fof(f5305,plain,(
  ( ! [X26,X24,X25] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26)) ) | $spl6),
  inference(subsumption_resolution,[],[f5184,f307])).
fof(f5184,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1451])).
fof(f5303,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5183,f307])).
fof(f5183,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22)))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1450])).
fof(f5301,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90)) ) | $spl6),
  inference(subsumption_resolution,[],[f5182,f307])).
fof(f5182,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1411])).
fof(f5299,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86)) ) | $spl6),
  inference(subsumption_resolution,[],[f5181,f307])).
fof(f5181,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1410])).
fof(f5297,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81)) ) | $spl6),
  inference(subsumption_resolution,[],[f5180,f307])).
fof(f5180,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1409])).
fof(f5295,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78)) ) | $spl6),
  inference(subsumption_resolution,[],[f5179,f307])).
fof(f5179,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1408])).
fof(f5293,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X73)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73)) ) | $spl6),
  inference(subsumption_resolution,[],[f5178,f307])).
fof(f5178,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X73)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1407])).
fof(f5291,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5177,f307])).
fof(f5177,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1406])).
fof(f5289,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | p(s(bool,X66)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65)) ) | $spl6),
  inference(subsumption_resolution,[],[f5176,f307])).
fof(f5176,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1405])).
fof(f5287,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62)) ) | $spl6),
  inference(subsumption_resolution,[],[f5175,f307])).
fof(f5175,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1404])).
fof(f5285,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57)) ) | $spl6),
  inference(subsumption_resolution,[],[f5174,f307])).
fof(f5174,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1403])).
fof(f5283,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X53)) | p(s(bool,X54)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54)) ) | $spl6),
  inference(subsumption_resolution,[],[f5173,f307])).
fof(f5173,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | p(s(bool,X54)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1402])).
fof(f5281,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5172,f307])).
fof(f5172,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1401])).
fof(f5279,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46)) ) | $spl6),
  inference(subsumption_resolution,[],[f5171,f307])).
fof(f5171,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1400])).
fof(f5277,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42)) ) | $spl6),
  inference(subsumption_resolution,[],[f5170,f307])).
fof(f5170,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1399])).
fof(f5275,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38)) ) | $spl6),
  inference(subsumption_resolution,[],[f5169,f307])).
fof(f5169,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1398])).
fof(f5273,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34)) ) | $spl6),
  inference(subsumption_resolution,[],[f5168,f307])).
fof(f5168,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1397])).
fof(f5271,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5167,f307])).
fof(f5167,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1396])).
fof(f5269,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27)) ) | $spl6),
  inference(subsumption_resolution,[],[f5166,f307])).
fof(f5166,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1395])).
fof(f5267,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24)) ) | $spl6),
  inference(subsumption_resolution,[],[f5165,f307])).
fof(f5165,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24)) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1394])).
fof(f5265,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5164,f307])).
fof(f5164,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f1393])).
fof(f5263,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5161,f307])).
fof(f5161,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f413])).
fof(f5261,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5160,f307])).
fof(f5160,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f412])).
fof(f5259,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5159,f307])).
fof(f5159,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f410])).
fof(f5250,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f5032,f4913])).
fof(f5032,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f4913,f336])).
fof(f5356,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5355])).
fof(f5354,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5219,f307])).
fof(f5219,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4913,f4166])).
fof(f5353,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5218,f307])).
fof(f5218,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4913,f2222])).
fof(f5352,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5217,f307])).
fof(f5217,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4913,f4102])).
fof(f5351,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5216,f307])).
fof(f5216,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4913,f2221])).
fof(f5350,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349])).
fof(f5348,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347])).
fof(f5346,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345])).
fof(f5344,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343])).
fof(f5342,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341])).
fof(f5340,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339])).
fof(f5338,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337])).
fof(f5336,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335])).
fof(f5334,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333])).
fof(f5332,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331])).
fof(f5330,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329])).
fof(f5328,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327])).
fof(f5326,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325])).
fof(f5324,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323])).
fof(f5322,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321])).
fof(f5320,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319])).
fof(f5318,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317])).
fof(f5316,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315])).
fof(f5314,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313])).
fof(f5312,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311])).
fof(f5310,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309])).
fof(f5308,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307])).
fof(f5306,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305])).
fof(f5304,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303])).
fof(f5302,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301])).
fof(f5300,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299])).
fof(f5298,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297])).
fof(f5296,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295])).
fof(f5294,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293])).
fof(f5292,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291])).
fof(f5290,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289])).
fof(f5288,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287])).
fof(f5286,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285])).
fof(f5284,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283])).
fof(f5282,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281])).
fof(f5280,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279])).
fof(f5278,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277])).
fof(f5276,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275])).
fof(f5274,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273])).
fof(f5272,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269,f5271])).
fof(f5270,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267,f5269])).
fof(f5268,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265,f5267])).
fof(f5266,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263,f5265])).
fof(f5264,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261,f5263])).
fof(f5262,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259,f5261])).
fof(f5260,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250,f5259])).
fof(f5252,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5034,f307])).
fof(f5034,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4913,f546])).
fof(f5251,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f565,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f598,f608,f609,f612,f617,f620,f625,f628,f642,f648,f649,f653,f656,f660,f668,f669,f671,f679,f729,f780,f797,f795,f815,f811,f855,f872,f881,f885,f890,f894,f905,f914,f942,f948,f956,f987,f998,f1007,f1044,f1051,f1060,f1059,f1061,f1064,f1065,f1066,f1217,f1223,f1245,f1244,f1281,f1315,f1382,f1381,f1380,f1379,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1437,f1436,f1435,f1434,f1450,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1568,f1553,f1571,f1560,f1575,f1641,f1640,f1651,f1652,f1653,f1654,f1657,f1656,f1658,f1660,f1661,f1662,f1663,f1808,f1786,f1809,f1810,f1813,f1879,f1883,f1930,f1997,f2127,f2136,f2138,f2221,f2222,f2225,f2218,f2298,f2305,f2294,f2347,f2350,f2351,f2352,f2391,f2394,f2399,f2397,f2396,f2412,f2410,f2409,f2476,f2590,f2701,f2707,f2710,f2781,f2784,f2840,f2861,f2864,f2871,f2872,f2932,f2945,f3120,f3133,f3354,f3355,f3357,f3320,f3330,f3387,f3388,f3425,f3391,f3401,f3608,f3607,f3606,f3604,f3603,f3602,f3599,f3598,f3594,f3591,f3590,f3586,f3585,f3580,f3565,f3564,f3563,f3559,f3558,f3554,f3553,f3552,f3549,f3545,f3543,f3540,f3539,f3535,f3534,f3533,f3532,f3531,f3528,f3526,f3521,f3520,f3516,f3515,f3510,f3503,f3502,f3500,f3498,f3497,f3496,f3491,f3490,f3485,f3484,f3483,f3482,f3478,f3477,f3472,f3471,f3467,f3466,f3461,f3460,f3458,f3457,f3454,f3453,f3452,f3451,f3449,f3448,f3447,f3445,f3650,f3764,f3849,f3892,f3895,f3942,f3970,f3999,f4002,f4055,f4102,f4084,f4115,f4114,f4166,f4148,f4182,f4238,f4239,f4240,f4290,f4253,f4296,f4305,f4326,f4327,f4328,f4373,f4381,f4386,f4408,f4411,f4434,f4447,f4525,f4526,f4527,f4528,f4586,f4911,f4915,f4918,f4921,f4924,f4927,f4930,f4933,f4936,f4939,f4942,f4945,f4948,f4951,f4954,f4957,f4960,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4987,f4988,f4989,f4990,f4667,f4668,f4669,f4670,f4996,f4998,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f4913,f5250])).
fof(f4961,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f4961_D])).
fof(f4961_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4913,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f4913_D])).
fof(f4913_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f5028,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f5027])).
fof(f5027,plain,(
  ( ! [X94,X92,X95,X93] : (~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f4698,f780])).
fof(f4698,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f531,f729])).
fof(f4958,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f4958_D])).
fof(f4958_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5026,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f5025])).
fof(f5025,plain,(
  ( ! [X90,X88,X91,X89] : (~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f4697,f780])).
fof(f4697,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f530,f729])).
fof(f4955,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f4955_D])).
fof(f4955_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5024,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f5023])).
fof(f5023,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f4696,f780])).
fof(f4696,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f527,f729])).
fof(f4952,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f4952_D])).
fof(f4952_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5022,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f5021])).
fof(f5021,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f4695,f780])).
fof(f4695,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f526,f729])).
fof(f4949,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f4949_D])).
fof(f4949_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f5020,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f5019])).
fof(f5019,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f4694,f780])).
fof(f4694,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f525,f729])).
fof(f4946,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f4946_D])).
fof(f4946_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5018,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f5017])).
fof(f5017,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f4693,f780])).
fof(f4693,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f520,f729])).
fof(f4943,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f4943_D])).
fof(f4943_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f5016,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f5015])).
fof(f5015,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f4692,f780])).
fof(f4692,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f510,f729])).
fof(f4940,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f4940_D])).
fof(f4940_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f5014,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f5013])).
fof(f5013,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f4691,f780])).
fof(f4691,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f509,f729])).
fof(f4937,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f4937_D])).
fof(f4937_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5012,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f5011])).
fof(f5011,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f4690,f780])).
fof(f4690,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f508,f729])).
fof(f4934,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f4934_D])).
fof(f4934_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f5010,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f5009])).
fof(f5009,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f4689,f780])).
fof(f4689,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f503,f729])).
fof(f4931,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f4931_D])).
fof(f4931_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5008,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f5007])).
fof(f5007,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f4688,f780])).
fof(f4688,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f495,f729])).
fof(f4928,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f4928_D])).
fof(f4928_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5006,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f5005])).
fof(f5005,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f4687,f780])).
fof(f4687,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f486,f729])).
fof(f4925,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f4925_D])).
fof(f4925_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5004,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f5003])).
fof(f5003,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f4686,f780])).
fof(f4686,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f467,f729])).
fof(f4922,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f4922_D])).
fof(f4922_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5002,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f5001])).
fof(f5001,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f4685,f780])).
fof(f4685,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f463,f729])).
fof(f4919,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl10),
  inference(cnf_transformation,[],[f4919_D])).
fof(f4919_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5000,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4999])).
fof(f4999,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f4684,f780])).
fof(f4684,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f459,f729])).
fof(f4916,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl8),
  inference(cnf_transformation,[],[f4916_D])).
fof(f4916_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4998,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4997])).
fof(f4997,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4683,f780])).
fof(f4683,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f450,f729])).
fof(f4912,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl4),
  inference(cnf_transformation,[],[f4912_D])).
fof(f4912_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4996,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4995])).
fof(f4995,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f4682,f780])).
fof(f4682,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f447,f729])).
fof(f4960,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4959])).
fof(f4959,plain,(
  ( ! [X94,X92,X95,X93] : (~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f4613,f780])).
fof(f4613,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f531,f729])).
fof(f4957,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4956])).
fof(f4956,plain,(
  ( ! [X90,X88,X91,X89] : (~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f4612,f780])).
fof(f4612,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f530,f729])).
fof(f4954,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4953])).
fof(f4953,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f4611,f780])).
fof(f4611,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f527,f729])).
fof(f4951,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4950])).
fof(f4950,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f4610,f780])).
fof(f4610,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f526,f729])).
fof(f4948,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4947])).
fof(f4947,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f4609,f780])).
fof(f4609,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f525,f729])).
fof(f4945,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f4944])).
fof(f4944,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f4608,f780])).
fof(f4608,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f520,f729])).
fof(f4942,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4941])).
fof(f4941,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f4607,f780])).
fof(f4607,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f510,f729])).
fof(f4939,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4938])).
fof(f4938,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f4606,f780])).
fof(f4606,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f509,f729])).
fof(f4936,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4935])).
fof(f4935,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f4605,f780])).
fof(f4605,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f508,f729])).
fof(f4933,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4932])).
fof(f4932,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f4604,f780])).
fof(f4604,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f503,f729])).
fof(f4930,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4929])).
fof(f4929,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f4603,f780])).
fof(f4603,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f495,f729])).
fof(f4927,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4926])).
fof(f4926,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f4602,f780])).
fof(f4602,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f486,f729])).
fof(f4924,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4923])).
fof(f4923,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f4601,f780])).
fof(f4601,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f467,f729])).
fof(f4921,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4920])).
fof(f4920,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f4600,f780])).
fof(f4600,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f463,f729])).
fof(f4918,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4917])).
fof(f4917,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f4599,f780])).
fof(f4599,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f459,f729])).
fof(f4915,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4914])).
fof(f4914,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4598,f780])).
fof(f4598,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f450,f729])).
fof(f4911,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4910])).
fof(f4910,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f4597,f780])).
fof(f4597,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f447,f729])).
fof(f542,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | $spl0),
  inference(cnf_transformation,[],[f542_D])).
fof(f542_D,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f544,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | $spl2),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f305,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))))))))),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))) & (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f202])).
fof(f202,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) & (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ? [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) <~> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|4','thm.integer.INT_DIVIDES_REDUCE|split|4')).
fof(f543,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | $spl1),
  inference(cnf_transformation,[],[f543_D])).
fof(f543_D,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f545,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | $spl3),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f306,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))))))))),
  inference(cnf_transformation,[],[f203])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 787
% Passive clauses: 2397
% Generated clauses: 21221
% Final active clauses: 542
% Final passive clauses: 1156
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 196
% Trivial inequalities: 136
% Fw subsumption resolutions: 841
% Fw demodulations: 5430
% Bw demodulations: 245
% Condensations: 220
% Global subsumptions: 51
% 
% Simple tautologies: 109
% Equational tautologies: 72
% Forward subsumptions: 11132
% Fw demodulations to eq. taut.: 147
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 5389
% Factoring: 1
% Forward superposition: 4664
% Backward superposition: 3422
% Self superposition: 49
% Equality factoring: 28
% Equality resolution: 24
% 
% Split clauses: 57
% Split components: 119
% SAT solver clauses: 2997
% SAT solver unit clauses: 714
% SAT solver binary clauses: 846
% 
% Sat splits: 57
% Sat splitting refutations: 433
% 
% Memory used [KB]: 13432
% Time elapsed: 0.660 s
% ------------------------------
% ------------------------------
% Success in time 1.987 s
