Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (54110)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 5696
% Generated clauses: 15324
% Final active clauses: 129
% Final passive clauses: 5450
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 490
% Trivial inequalities: 183
% Fw subsumption resolutions: 1080
% Fw demodulations: 449
% Bw demodulations: 3
% 
% Simple tautologies: 270
% Equational tautologies: 25
% Forward subsumptions: 5803
% 
% Binary resolution: 78
% Forward superposition: 7222
% Backward superposition: 5380
% Self superposition: 186
% Equality factoring: 76
% Equality resolution: 23
% 
% Split clauses: 343
% Split components: 694
% SAT solver clauses: 462
% SAT solver unit clauses: 5
% SAT solver binary clauses: 45
% 
% Sat splits: 343
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 461
% TWLsolver calls for satisfiability: 87
% 
% Memory used [KB]: 14839
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (54154)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 82
% Passive clauses: 1744
% Generated clauses: 9183
% Final active clauses: 45
% Final passive clauses: 754
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 120
% Trivial inequalities: 34
% Fw subsumption resolutions: 989
% Fw demodulations: 349
% Bw demodulations: 878
% Condensations: 216
% 
% Simple tautologies: 60
% Equational tautologies: 17
% Forward subsumptions: 3273
% Bw demodulations to eq. taut.: 57
% 
% Binary resolution: 17
% Forward superposition: 1946
% Backward superposition: 1665
% Self superposition: 105
% Equality factoring: 39
% Equality resolution: 8
% 
% Split clauses: 72
% Split components: 144
% SAT solver clauses: 1967
% SAT solver unit clauses: 45
% SAT solver binary clauses: 61
% 
% Sat splits: 72
% Sat splitting refutations: 102
% 
% TWLsolver clauses: 1967
% TWLsolver calls for satisfiability: 128
% 
% Memory used [KB]: 17526
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% (54208)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 268
% Passive clauses: 13336
% Generated clauses: 67263
% Final active clauses: 58
% Final passive clauses: 117
% Input formulas: 21
% Initial clauses: 26
% Discarded non-redundant clauses: 3024
% 
% Duplicate literals: 1793
% Trivial inequalities: 478
% Fw subsumption resolutions: 2179
% Fw demodulations: 5539
% Bw demodulations: 7803
% Condensations: 2604
% Global subsumptions: 1503
% 
% Simple tautologies: 188
% Equational tautologies: 139
% Forward subsumptions: 36874
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 882
% 
% Unit resulting resolution: 708
% Forward superposition: 14745
% Backward superposition: 15101
% Self superposition: 412
% Equality factoring: 99
% Equality resolution: 8
% 
% Split clauses: 221
% Split components: 447
% SAT solver clauses: 16875
% SAT solver unit clauses: 558
% SAT solver binary clauses: 1023
% 
% Sat splits: 221
% Sat splitting refutations: 1591
% 
% Memory used [KB]: 115392
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.ringNorm
% (54423)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 5925
% Passive clauses: 28177
% Generated clauses: 29079
% Final active clauses: 5925
% Final passive clauses: 22101
% Input formulas: 21
% Initial clauses: 12
% 
% Selected by SInE selection: 7
% SInE iterations: 1
% 
% Fw demodulations: 262
% 
% Simple tautologies: 2
% Forward subsumptions: 800
% 
% Binary resolution: 4970
% Forward superposition: 23724
% Equality resolution: 124
% 
% SAT solver clauses: 1
% 
% Memory used [KB]: 30319
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f11794,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f99,f114,f112,f985,f1038,f1071,f1076,f1078,f1086,f1088,f1092,f1094,f1096,f1327,f1339,f1340,f1341,f1342,f1343,f523,f1641,f1864,f1875,f1883,f1890,f1891,f1896,f1897,f527,f1640,f2625,f2837,f3320,f3559,f3588,f3871,f3873,f3877,f3885,f3887,f3888,f3889,f3890,f3891,f3565,f2156,f4906,f3568,f5883,f6678,f7943,f1037,f8529,f8961,f3558,f11789,f11791])).
fof(f11791,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f11790,f69])).
fof(f69,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1',aHL_TRUTH)).
fof(f11790,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(forward_demodulation,[],[f11787,f10813])).
fof(f10813,plain,(
  ( ! [X6,X4,X8,X7,X5,X3,X9] : (s(bool,'T_0') = s(bool,sK5(X3,X4,X5,'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X5),s(X6,X7),s(fun(X6,fun(X6,X6)),X8),s(fun(X6,fun(X6,X6)),X9)),X6))) ) | $spl6),
  inference(forward_demodulation,[],[f10588,f85])).
fof(f85,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X16,X17,X18,X19] : s(X6,X17) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X16),s(X6,X17),s(fun(X6,fun(X6,X6)),X18),s(fun(X6,fun(X6,X6)),X19)))))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1','thm.semi_ring.semi_ring_SR|49|')).
fof(f10588,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.recordtype|46|semi_ring_4'(s(bool,X0),s(bool,'T_0'),s(fun(bool,fun(bool,bool)),X1),s(fun(bool,fun(bool,bool)),X2))))) = s(bool,sK5(X3,X4,X5,'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X5),s(X6,X7),s(fun(X6,fun(X6,X6)),X8),s(fun(X6,fun(X6,X6)),X9)),X6))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f85,f391,f3944])).
fof(f3944,plain,(
  ( ! [X6,X5] : (s(bool,X5) = s(bool,X6) | s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X6))) ) | $spl6),
  inference(subsumption_resolution,[],[f3937,f70])).
fof(f70,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1',aHL_FALSITY)).
fof(f3937,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | p(s(bool,'F_0')) | ~p(s(bool,X6))) ) | $spl6),
  inference(superposition,[],[f2112,f478])).
fof(f478,plain,(
  ( ! [X10,X8,X9] : (s(bool,X9) = s(bool,X10) | s(bool,X8) = s(bool,X9) | p(s(bool,X10)) | ~p(s(bool,X8))) )),
  inference(superposition,[],[f354,f96])).
fof(f96,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f70,f71])).
fof(f71,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1',aHL_BOOL_CASES)).
fof(f354,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X4) = s(bool,X5) | p(s(bool,X4))) )),
  inference(superposition,[],[f69,f94])).
fof(f94,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f71,f71])).
fof(f2112,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f69,f1038])).
fof(f391,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,sK5(X0,X1,X2,'const.semi_ring.recordtype|46|semi_ring_4'(s(X3,X2),s(X3,X4),s(fun(X3,fun(X3,X3)),X5),s(fun(X3,fun(X3,X3)),X6)),X3)))) )),
  inference(unit_resulting_resolution,[],[f84,f91])).
fof(f91,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | p(s(bool,sK5(X4,X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,sK4(X4,X3,X2,X1,X0))) | s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & (s(X0,X2) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | p(s(bool,sK4(X4,X3,X2,X1,X0)))) & ((~p(s(bool,sK5(X4,X3,X2,X1,X0))) | s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & (s(X0,X2) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | p(s(bool,sK5(X4,X3,X2,X1,X0)))) & s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))))) = s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,sK5(X4,X3,X2,X1,X0)),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))))),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4)))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((~p(s(bool,X5)) | s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & (s(X0,X2) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | p(s(bool,X5))) & ? [X6] : ((~p(s(bool,X6)) | s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & (s(X0,X2) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | p(s(bool,X6))) & s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))))) = s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,X6),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,X5),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))))),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4)))))))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : (((~p(s(bool,X5)) | s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & (s(X0,X2) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | p(s(bool,X5)))) & ? [X6] : (((~p(s(bool,X6)) | s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & (s(X0,X2) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | p(s(bool,X6)))) & s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))))) = s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,X6),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,X5),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))))),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4)))))))))))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((p(s(bool,X5)) <=> s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & ? [X6] : ((p(s(bool,X6)) <=> s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))))) = s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,X6),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,X5),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))))),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4)))))))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ? [X5] : ((p(s(bool,X5)) <=> s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & ? [X6] : ((p(s(bool,X6)) <=> s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1)))) & s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X4))))) = s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,X6),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.bool.COND_3'(s(bool,X5),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4))))),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X2),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4)))))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X12,X13,X8,X9] : ? [X14] : ((p(s(bool,X14)) <=> s(X6,X13) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X12)))) & ? [X15] : ((p(s(bool,X15)) <=> s(X6,X13) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X12)))) & s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X13),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X9))))) = s('type.canonical.canonical_sum'(X6),'const.bool.COND_3'(s(bool,X15),s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.canonical_sum'(X6),X9))),s('type.canonical.canonical_sum'(X6),'const.bool.COND_3'(s(bool,X14),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.canonical_sum'(X6),X9))))),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X13),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.canonical_sum'(X6),X9)))))))))))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1','thm.canonical.canonical_sum_simplify_def|split|0')).
fof(f84,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X16,X17,X18,X19] : s(X6,X16) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X16),s(X6,X17),s(fun(X6,fun(X6,X6)),X18),s(fun(X6,fun(X6,X6)),X19)))))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1','thm.semi_ring.semi_ring_SR|48|')).
fof(f11787,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,sK5(X0,X1,'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3)))),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3))))),s('type.canonical.canonical_sum'(sK0),X2),s(fun('type.canonical.canonical_sum'(sK0),fun('type.canonical.canonical_sum'(sK0),'type.canonical.canonical_sum'(sK0))),X3),s(fun('type.canonical.canonical_sum'(sK0),fun('type.canonical.canonical_sum'(sK0),'type.canonical.canonical_sum'(sK0))),X4)),'type.canonical.canonical_sum'(sK0))))) )),
  inference(backward_demodulation,[],[f11785,f10540])).
fof(f10540,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,sK5(X0,X1,'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3)))),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK3))))),s('type.canonical.canonical_sum'(sK0),X2),s(fun('type.canonical.canonical_sum'(sK0),fun('type.canonical.canonical_sum'(sK0),'type.canonical.canonical_sum'(sK0))),X3),s(fun('type.canonical.canonical_sum'(sK0),fun('type.canonical.canonical_sum'(sK0),'type.canonical.canonical_sum'(sK0))),X4)),'type.canonical.canonical_sum'(sK0))))) )),
  inference(unit_resulting_resolution,[],[f68,f391,f228])).
fof(f228,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,sK5(X24,X25,X20,X19,X18))) | ~p(s(bool,sK5(X22,X23,X21,X19,X18))) | s(X18,X20) = s(X18,X21)) )),
  inference(superposition,[],[f90,f90])).
fof(f90,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | ~p(s(bool,sK5(X4,X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f68,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK3)))))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3))))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),X9))))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),X9))))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X9)))))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1','thm.ringNorm.canonical_sum_simplify_def|split|1')).
fof(f11785,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.canonical.canonical_sum'(X5),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X7),s('type.canonical.canonical_sum'(X5),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X5),X6),s('type.canonical.canonical_sum'(X5),X8))))) = s('type.canonical.canonical_sum'(X5),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X5),X6),s('type.canonical.canonical_sum'(X5),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X7),s('type.canonical.canonical_sum'(X5),X8)))))) )),
  inference(forward_demodulation,[],[f11754,f779])).
fof(f779,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(forward_demodulation,[],[f758,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X20,X21] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X6),X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1','ARITY_1_2_const.ringNorm.r_canonical_sum_simplify')).
fof(f758,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(superposition,[],[f75,f73])).
fof(f73,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1','thm.ringNorm.r_canonical_sum_simplify_def')).
fof(f75,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X20,X21] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X6),X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1','ARITY_1_2_const.canonical.canonical_sum_simplify')).
fof(f11754,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.canonical.canonical_sum'(X5),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X5),X6),s('type.canonical.canonical_sum'(X5),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X7),s('type.canonical.canonical_sum'(X5),X8))))) = s('type.canonical.canonical_sum'(X5),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X7),s('type.canonical.canonical_sum'(X5),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X5),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X5),X6))),s('type.canonical.canonical_sum'(X5),X8)))))) )),
  inference(superposition,[],[f779,f83])).
fof(f83,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X12,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.canonical_sum'(X6),X9))))) = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X9)))))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|1','thm.canonical.canonical_sum_simplify_def|split|1')).
fof(f11789,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f11788,f69])).
fof(f11788,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(forward_demodulation,[],[f11786,f11461])).
fof(f11461,plain,(
  ( ! [X6,X4,X8,X7,X5,X3,X9] : (s(bool,'T_0') = s(bool,sK4(X3,X4,X5,'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X7),s(X6,X5),s(fun(X6,fun(X6,X6)),X8),s(fun(X6,fun(X6,X6)),X9)),X6))) ) | $spl6),
  inference(forward_demodulation,[],[f11236,f85])).
fof(f11236,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.recordtype|46|semi_ring_4'(s(bool,X0),s(bool,'T_0'),s(fun(bool,fun(bool,bool)),X1),s(fun(bool,fun(bool,bool)),X2))))) = s(bool,sK4(X3,X4,X5,'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X7),s(X6,X5),s(fun(X6,fun(X6,X6)),X8),s(fun(X6,fun(X6,X6)),X9)),X6))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f85,f595,f3944])).
fof(f595,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,sK4(X0,X1,X2,'const.semi_ring.recordtype|46|semi_ring_4'(s(X3,X4),s(X3,X2),s(fun(X3,fun(X3,X3)),X5),s(fun(X3,fun(X3,X3)),X6)),X3)))) )),
  inference(unit_resulting_resolution,[],[f85,f89])).
fof(f89,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | p(s(bool,sK4(X4,X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f11786,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,sK4(X0,X1,'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3)))),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.canonical.canonical_sum'(sK0),X2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3))))),s(fun('type.canonical.canonical_sum'(sK0),fun('type.canonical.canonical_sum'(sK0),'type.canonical.canonical_sum'(sK0))),X3),s(fun('type.canonical.canonical_sum'(sK0),fun('type.canonical.canonical_sum'(sK0),'type.canonical.canonical_sum'(sK0))),X4)),'type.canonical.canonical_sum'(sK0))))) )),
  inference(backward_demodulation,[],[f11785,f11188])).
fof(f11188,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,sK4(X0,X1,'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3)))),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.canonical.canonical_sum'(sK0),X2),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK3))))),s(fun('type.canonical.canonical_sum'(sK0),fun('type.canonical.canonical_sum'(sK0),'type.canonical.canonical_sum'(sK0))),X3),s(fun('type.canonical.canonical_sum'(sK0),fun('type.canonical.canonical_sum'(sK0),'type.canonical.canonical_sum'(sK0))),X4)),'type.canonical.canonical_sum'(sK0))))) )),
  inference(unit_resulting_resolution,[],[f68,f595,f124])).
fof(f124,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,sK4(X14,X15,X10,X9,X8))) | ~p(s(bool,sK4(X12,X13,X11,X9,X8))) | s(X8,X10) = s(X8,X11)) )),
  inference(superposition,[],[f88,f88])).
fof(f88,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) | ~p(s(bool,sK4(X4,X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f3558,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) | $spl16),
  inference(cnf_transformation,[],[f3558_D])).
fof(f3558_D,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f8961,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') != s(bool,X41) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42))) ) | $spl26),
  inference(cnf_transformation,[],[f8961_D])).
fof(f8961_D,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') != s(bool,X41) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8529,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') != s(bool,X41) | ~p(s(bool,X40)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42)) ) | ($spl6 | $spl14)),
  inference(duplicate_literal_removal,[],[f8447])).
fof(f8447,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') != s(bool,X41) | ~p(s(bool,X40)) | ~p(s(bool,X42)) | ~p(s(bool,X40)) | s(bool,X41) != s(bool,X42)) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f1038,f3034])).
fof(f3034,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f1038,f2837])).
fof(f1037,plain,(
  ( ! [X22] : (p(s(bool,X22))) ) | $spl4),
  inference(cnf_transformation,[],[f1037_D])).
fof(f1037_D,plain,(
  ( ! [X22] : (p(s(bool,X22))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f7943,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') = s(bool,X45) | s(bool,X45) != s(bool,X46) | p(s(bool,X46))) ) | $spl24),
  inference(cnf_transformation,[],[f7943_D])).
fof(f7943_D,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') = s(bool,X45) | s(bool,X45) != s(bool,X46) | p(s(bool,X46))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6678,plain,(
  ( ! [X45,X46,X44] : (s(bool,'F_0') = s(bool,X45) | p(s(bool,X46)) | s(bool,X45) != s(bool,X46) | p(s(bool,X44))) ) | ($spl10 | $spl12)),
  inference(duplicate_literal_removal,[],[f6493])).
fof(f6493,plain,(
  ( ! [X45,X46,X44] : (s(bool,'F_0') = s(bool,X45) | p(s(bool,X46)) | s(bool,X45) != s(bool,X46) | p(s(bool,X44)) | p(s(bool,X44))) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f3950,f2156])).
fof(f3950,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) ) | $spl10),
  inference(superposition,[],[f1641,f354])).
fof(f5883,plain,(
  ( ! [X72,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X72)) | ~p(s(bool,X73))) ) | $spl22),
  inference(cnf_transformation,[],[f5883_D])).
fof(f5883_D,plain,(
  ( ! [X72,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X72)) | ~p(s(bool,X73))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3568,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) != s(bool,X73) | ~p(s(bool,X73)) | p(s(bool,X72)) | ~p(s(bool,X74))) ) | $spl6),
  inference(subsumption_resolution,[],[f2549,f1038])).
fof(f2549,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) != s(bool,X73) | ~p(s(bool,X73)) | s(bool,'F_0') = s(bool,X74) | p(s(bool,X72)) | ~p(s(bool,X74))) ) | $spl6),
  inference(superposition,[],[f1038,f478])).
fof(f4906,plain,(
  ( ! [X28,X26] : (s(bool,X26) = s(bool,X28)) ) | $spl20),
  inference(cnf_transformation,[],[f4906_D])).
fof(f4906_D,plain,(
  ( ! [X28,X26] : (s(bool,X26) = s(bool,X28)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2156,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X15))) ) | $spl12),
  inference(cnf_transformation,[],[f2156_D])).
fof(f2156_D,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X15))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3565,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | p(s(bool,X27))) ) | $spl14),
  inference(subsumption_resolution,[],[f3117,f3132])).
fof(f3132,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') = s(bool,X27) | p(s(bool,X27)) | p(s(bool,X26)) | s(bool,X26) = s(bool,X28)) ) | $spl14),
  inference(subsumption_resolution,[],[f3069,f495])).
fof(f495,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X8)) | p(s(bool,X9))) )),
  inference(superposition,[],[f69,f354])).
fof(f3069,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') = s(bool,X27) | p(s(bool,X27)) | p(s(bool,X26)) | ~p(s(bool,X28)) | s(bool,X26) = s(bool,X28)) ) | $spl14),
  inference(superposition,[],[f495,f2837])).
fof(f3117,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') = s(bool,X27) | ~p(s(bool,X28)) | s(bool,X26) = s(bool,X28) | p(s(bool,X27))) ) | $spl14),
  inference(subsumption_resolution,[],[f3022,f101])).
fof(f101,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f96,f96])).
fof(f3022,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') = s(bool,X27) | ~p(s(bool,X28)) | s(bool,X26) = s(bool,X28) | p(s(bool,X27)) | p(s(bool,X26))) ) | $spl14),
  inference(superposition,[],[f2837,f495])).
fof(f3891,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3805,f3870])).
fof(f3870,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | ($spl0 | $spl6)),
  inference(trivial_inequality_removal,[],[f3861])).
fof(f3861,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,X3))) ) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f3855,f2148])).
fof(f2148,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X3))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2128])).
fof(f2128,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X3)) | ~p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f1038,f96])).
fof(f3855,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f3854,f70])).
fof(f3854,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(forward_demodulation,[],[f3846,f112])).
fof(f3846,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,sK4(X8,X9,X6,X7,bool))) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(trivial_inequality_removal,[],[f3602])).
fof(f3602,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,sK4(X8,X9,X6,X7,bool))) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(backward_demodulation,[],[f112,f141])).
fof(f141,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X7))) | p(s(bool,sK4(X8,X9,X6,X7,bool))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f89,f71])).
fof(f3805,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(backward_demodulation,[],[f112,f393])).
fof(f393,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,sK5(X0,X1,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X2),'const.semi_ring.recordtype|46|semi_ring_4'(s(X2,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X2),X3))),s(X2,X4),s(fun(X2,fun(X2,X2)),X5),s(fun(X2,fun(X2,X2)),X6)))),X3,X2)))) )),
  inference(unit_resulting_resolution,[],[f84,f91])).
fof(f3890,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3802,f3870])).
fof(f3802,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(backward_demodulation,[],[f112,f597])).
fof(f597,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,sK5(X0,X1,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X2),'const.semi_ring.recordtype|46|semi_ring_4'(s(X2,X3),s(X2,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X2),X4))),s(fun(X2,fun(X2,X2)),X5),s(fun(X2,fun(X2,X2)),X6)))),X4,X2)))) )),
  inference(unit_resulting_resolution,[],[f85,f91])).
fof(f3889,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3792,f3870])).
fof(f3792,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(backward_demodulation,[],[f112,f391])).
fof(f3888,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3742,f3870])).
fof(f3742,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(backward_demodulation,[],[f112,f392])).
fof(f392,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,sK4(X0,X1,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X2),'const.semi_ring.recordtype|46|semi_ring_4'(s(X2,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X2),X3))),s(X2,X4),s(fun(X2,fun(X2,X2)),X5),s(fun(X2,fun(X2,X2)),X6)))),X3,X2)))) )),
  inference(unit_resulting_resolution,[],[f84,f89])).
fof(f3887,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3736,f3870])).
fof(f3736,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(backward_demodulation,[],[f112,f596])).
fof(f596,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,sK4(X0,X1,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X2),'const.semi_ring.recordtype|46|semi_ring_4'(s(X2,X3),s(X2,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X2),X4))),s(fun(X2,fun(X2,X2)),X5),s(fun(X2,fun(X2,X2)),X6)))),X4,X2)))) )),
  inference(unit_resulting_resolution,[],[f85,f89])).
fof(f3885,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3724,f3870])).
fof(f3724,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(backward_demodulation,[],[f112,f595])).
fof(f3877,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3876,f3870])).
fof(f3876,plain,(
  ( ! [X74] : (p(s(bool,X74))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3875,f3870])).
fof(f3875,plain,(
  ( ! [X74,X73] : (p(s(bool,X73)) | p(s(bool,X74))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3874,f3870])).
fof(f3874,plain,(
  ( ! [X76,X74,X77,X75,X73] : (p(s(bool,sK4(X76,X77,X73,X75,bool))) | p(s(bool,X73)) | p(s(bool,X74))) ) | $spl0),
  inference(subsumption_resolution,[],[f3625,f112])).
fof(f3625,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'F_0') != s(bool,X74) | p(s(bool,sK4(X76,X77,X73,X75,bool))) | p(s(bool,X73)) | p(s(bool,X74))) ) | $spl0),
  inference(backward_demodulation,[],[f112,f851])).
fof(f851,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,X74) != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X75))) | p(s(bool,sK4(X76,X77,X73,X75,bool))) | p(s(bool,X73)) | p(s(bool,X74))) )),
  inference(superposition,[],[f89,f495])).
fof(f3873,plain,(
  $false | ($spl0 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f3872,f3855])).
fof(f3872,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | ($spl0 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f3859,f1641])).
fof(f3859,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f3855,f1038])).
fof(f3871,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3856,f69])).
fof(f3856,plain,(
  ~p(s(bool,'T_0')) | $spl0),
  inference(backward_demodulation,[],[f3855,f70])).
fof(f3588,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f113,f2112])).
fof(f113,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl2),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3559,plain,(
  ( ! [X44] : (s(bool,'F_0') != s(bool,X44) | s(bool,'T_0') != s(bool,X44)) ) | $spl18),
  inference(cnf_transformation,[],[f3559_D])).
fof(f3559_D,plain,(
  ( ! [X44] : (s(bool,'F_0') != s(bool,X44) | s(bool,'T_0') != s(bool,X44)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3320,plain,(
  ( ! [X43,X44] : (s(bool,'F_0') != s(bool,X44) | ~p(s(bool,X43)) | s(bool,'T_0') != s(bool,X44)) ) | ($spl3 | $spl6 | $spl14)),
  inference(duplicate_literal_removal,[],[f3281])).
fof(f3281,plain,(
  ( ! [X43,X44] : (s(bool,'F_0') != s(bool,X44) | ~p(s(bool,X43)) | ~p(s(bool,X43)) | s(bool,'T_0') != s(bool,X44)) ) | ($spl3 | $spl6 | $spl14)),
  inference(superposition,[],[f1038,f3037])).
fof(f3037,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X8)) ) | ($spl3 | $spl14)),
  inference(superposition,[],[f114,f2837])).
fof(f2837,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,X41) | ~p(s(bool,X42)) | s(bool,X41) = s(bool,X42)) ) | $spl14),
  inference(cnf_transformation,[],[f2837_D])).
fof(f2837_D,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,X41) | ~p(s(bool,X42)) | s(bool,X41) = s(bool,X42)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2625,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,X42) | p(s(bool,X40)) | ~p(s(bool,X42))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2513])).
fof(f2513,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,X42) | p(s(bool,X40)) | ~p(s(bool,X42)) | p(s(bool,X40))) ) | $spl3),
  inference(superposition,[],[f478,f526])).
fof(f526,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X16))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f506])).
fof(f506,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,X16) | p(s(bool,X16))) ) | $spl3),
  inference(superposition,[],[f114,f354])).
fof(f1640,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f1640_D])).
fof(f1640_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f527,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X15))) )),
  inference(duplicate_literal_removal,[],[f487])).
fof(f487,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f354,f71])).
fof(f1897,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1850,f1886])).
fof(f1886,plain,(
  ( ! [X40] : (~p(s(bool,X40))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1885,f69])).
fof(f1885,plain,(
  ( ! [X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40))) ) | ($spl6 | $spl8)),
  inference(forward_demodulation,[],[f1884,f1640])).
fof(f1884,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,sK4(X42,X43,X40,X41,bool)))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1670,f1640])).
fof(f1670,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X41))) | ~p(s(bool,X40)) | ~p(s(bool,sK4(X42,X43,X40,X41,bool)))) ) | ($spl6 | $spl8)),
  inference(backward_demodulation,[],[f1640,f1626])).
fof(f1626,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X41))) | ~p(s(bool,X40)) | ~p(s(bool,sK4(X42,X43,X40,X41,bool)))) ) | $spl6),
  inference(superposition,[],[f1038,f88])).
fof(f1850,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (p(s(bool,sK5(X23,X24,sK5(X18,X19,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X20),X21)),X21,X20),X22,bool)))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f1712])).
fof(f1712,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,sK5(X23,X24,sK5(X18,X19,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X20),X21)),X21,X20),X22,bool)))) ) | $spl8),
  inference(backward_demodulation,[],[f1640,f320])).
fof(f320,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X22))) | p(s(bool,sK5(X23,X24,sK5(X18,X19,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X20),X21)),X21,X20),X22,bool)))) )),
  inference(superposition,[],[f91,f294])).
fof(f294,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK5(X0,X1,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X2),X3)),X3,X2))) )),
  inference(unit_resulting_resolution,[],[f288,f96])).
fof(f288,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK5(X0,X1,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X2),X3)),X3,X2)))) )),
  inference(equality_resolution,[],[f91])).
fof(f1896,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1851,f1886])).
fof(f1851,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (p(s(bool,sK5(X27,X28,sK4(X22,X23,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X24),X25)),X25,X24),X26,bool)))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f1708])).
fof(f1708,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,sK5(X27,X28,sK4(X22,X23,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X24),X25)),X25,X24),X26,bool)))) ) | $spl8),
  inference(backward_demodulation,[],[f1640,f271])).
fof(f271,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X26))) | p(s(bool,sK5(X27,X28,sK4(X22,X23,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X24),X25)),X25,X24),X26,bool)))) )),
  inference(superposition,[],[f91,f153])).
fof(f153,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK4(X0,X1,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X2),X3)),X3,X2))) )),
  inference(unit_resulting_resolution,[],[f151,f96])).
fof(f151,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK4(X0,X1,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X2),X3)),X3,X2)))) )),
  inference(equality_resolution,[],[f89])).
fof(f1891,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1858,f1886])).
fof(f1858,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (p(s(bool,sK4(X37,X38,sK5(X32,X33,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X34),X35)),X35,X34),X36,bool)))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f1686])).
fof(f1686,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,sK4(X37,X38,sK5(X32,X33,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X34),X35)),X35,X34),X36,bool)))) ) | $spl8),
  inference(backward_demodulation,[],[f1640,f322])).
fof(f322,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X36))) | p(s(bool,sK4(X37,X38,sK5(X32,X33,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X34),X35)),X35,X34),X36,bool)))) )),
  inference(superposition,[],[f89,f294])).
fof(f1890,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1859,f1886])).
fof(f1859,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (p(s(bool,sK4(X30,X31,sK4(X25,X26,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X27),X28)),X28,X27),X29,bool)))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f1684])).
fof(f1684,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,sK4(X30,X31,sK4(X25,X26,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X27),X28)),X28,X27),X29,bool)))) ) | $spl8),
  inference(backward_demodulation,[],[f1640,f265])).
fof(f265,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X29))) | p(s(bool,sK4(X30,X31,sK4(X25,X26,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X27),X28)),X28,X27),X29,bool)))) )),
  inference(superposition,[],[f89,f153])).
fof(f1883,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1882,f69])).
fof(f1882,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl8)),
  inference(forward_demodulation,[],[f1881,f1640])).
fof(f1881,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,sK4(X38,X39,X37,X36,bool)))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1880,f69])).
fof(f1880,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,sK4(X38,X39,X37,X36,bool)))) ) | ($spl6 | $spl8)),
  inference(forward_demodulation,[],[f1879,f1640])).
fof(f1879,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X36)))) | ~p(s(bool,sK4(X38,X39,X37,X36,bool)))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1669,f1640])).
fof(f1669,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X36)))) | ~p(s(bool,sK4(X38,X39,X37,X36,bool)))) ) | ($spl6 | $spl8)),
  inference(backward_demodulation,[],[f1640,f1625])).
fof(f1625,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | ~p(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X36)))) | ~p(s(bool,sK4(X38,X39,X37,X36,bool)))) ) | $spl6),
  inference(superposition,[],[f1038,f88])).
fof(f1875,plain,(
  $false | ($spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1874,f1640])).
fof(f1874,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | ($spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1664,f554])).
fof(f554,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) | $spl3),
  inference(superposition,[],[f114,f526])).
fof(f1664,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | ($spl6 | $spl8)),
  inference(backward_demodulation,[],[f1640,f1038])).
fof(f1864,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1642,f69])).
fof(f1642,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f1640,f70])).
fof(f1641,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) ) | $spl10),
  inference(cnf_transformation,[],[f1641_D])).
fof(f1641_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f523,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) )),
  inference(equality_factoring,[],[f354])).
fof(f1343,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f1072,f1331])).
fof(f1331,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,X21)) ) | $spl4),
  inference(subsumption_resolution,[],[f1330,f1037])).
fof(f1330,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,X21) | ~p(s(bool,X21))) ) | $spl4),
  inference(subsumption_resolution,[],[f1329,f1037])).
fof(f1329,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X20) = s(bool,X21) | ~p(s(bool,X21)) | ~p(s(bool,sK4(X22,X23,X20,X19,bool)))) ) | $spl4),
  inference(subsumption_resolution,[],[f1154,f69])).
fof(f1154,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'T_0')) | s(bool,X20) = s(bool,X21) | ~p(s(bool,X21)) | ~p(s(bool,sK4(X22,X23,X20,X19,bool)))) ) | $spl4),
  inference(backward_demodulation,[],[f1095,f170])).
fof(f170,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X20) = s(bool,X21) | ~p(s(bool,X21)) | ~p(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X19)))) | ~p(s(bool,sK4(X22,X23,X20,X19,bool)))) )),
  inference(superposition,[],[f101,f88])).
fof(f1095,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1037,f96])).
fof(f1072,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f70,f1037,f708])).
fof(f708,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl3),
  inference(superposition,[],[f554,f96])).
fof(f1342,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f1085,f69])).
fof(f1085,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f114,f1037,f101])).
fof(f1341,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f1087,f1037])).
fof(f1087,plain,(
  ~p(s(bool,'F_0')) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f114,f1037,f101])).
fof(f1340,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f1091,f69])).
fof(f1091,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f114,f1037,f101])).
fof(f1339,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f1093,f1037])).
fof(f1093,plain,(
  ~p(s(bool,'F_0')) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f114,f1037,f101])).
fof(f1327,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1124,f69])).
fof(f1124,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f1095,f70])).
fof(f1096,plain,(
  $false | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f114,f1037,f96])).
fof(f1094,plain,(
  $false | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f1037,f114,f1037,f101])).
fof(f1092,plain,(
  $false | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f69,f114,f1037,f101])).
fof(f1088,plain,(
  $false | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f1037,f114,f1037,f101])).
fof(f1086,plain,(
  $false | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f69,f114,f1037,f101])).
fof(f1078,plain,(
  $false | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f70,f85,f1037,f708])).
fof(f1076,plain,(
  $false | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f70,f84,f1037,f708])).
fof(f1071,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f70,f1037])).
fof(f1038,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X23))) ) | $spl6),
  inference(cnf_transformation,[],[f1038_D])).
fof(f1038_D,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X23))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f985,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f970])).
fof(f970,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X22))) ) | $spl3),
  inference(superposition,[],[f708,f526])).
fof(f112,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f114,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f99,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f71])).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 114
% Passive clauses: 2151
% Generated clauses: 12484
% Final active clauses: 83
% Final passive clauses: 1266
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 293
% Trivial inequalities: 182
% Fw subsumption resolutions: 878
% Fw demodulations: 1277
% Bw demodulations: 668
% 
% Simple tautologies: 320
% Equational tautologies: 65
% Forward subsumptions: 7285
% Bw demodulations to eq. taut.: 31
% 
% Unit resulting resolution: 628
% Forward superposition: 4371
% Backward superposition: 3202
% Self superposition: 138
% Equality factoring: 28
% Equality resolution: 18
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 2196
% SAT solver unit clauses: 205
% SAT solver binary clauses: 616
% 
% Sat splits: 10
% Sat splitting refutations: 33
% 
% Memory used [KB]: 13816
% Time elapsed: 0.580 s
% ------------------------------
% ------------------------------
% Success in time 16.32 s
