Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real_sigma
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real_sigma
% SZS output start Proof for thm.real_sigma
fof(f7084,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7083,f214,f322,f324,f193,f325])).
fof(f325,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f325_D])).
fof(f325_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f193,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_DISJOINT_UNION',aHL_FALSITY)).
fof(f324,plain,(
  p(s(bool,'F_0')) | $spl2),
  inference(cnf_transformation,[],[f324_D])).
fof(f324_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f322,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f322_D])).
fof(f322_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f214,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.real_sigma.REAL_SUM_IMAGE_DISJOINT_UNION','thm.bool.IMP_CLAUSES|split|2')).
fof(f7083,plain,(
  $false | $spl0),
  inference(trivial_inequality_removal,[],[f7080])).
fof(f7080,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),sK13),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),sK13),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) | $spl0),
  inference(superposition,[],[f186,f3766])).
fof(f3766,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK11))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK12)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f3765,f322])).
fof(f3765,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK11))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK12)))))) ) | $spl0),
  inference(forward_demodulation,[],[f3764,f339])).
fof(f339,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK11)))),
  inference(resolution,[],[f201,f183])).
fof(f183,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK11))))),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK11)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK12)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12)))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),sK13),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),sK13),s(fun(sK10,bool),sK11))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),sK13),s(fun(sK10,bool),sK12)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f83])).
fof(f83,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & ? [X3] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) & ? [X3] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) => ! [X3] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) => ! [X3] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))) => ! [X9] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X7))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))) => ! [X9] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X7))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X8))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_DISJOINT_UNION','thm.real_sigma.REAL_SUM_IMAGE_DISJOINT_UNION')).
fof(f201,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_DISJOINT_UNION','thm.bool.EQ_CLAUSES|split|1')).
fof(f3764,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK11))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK12))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK11))))) ) | $spl0),
  inference(subsumption_resolution,[],[f3763,f322])).
fof(f3763,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK11))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK12))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK11))))) )),
  inference(forward_demodulation,[],[f3738,f340])).
fof(f340,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK12)))),
  inference(resolution,[],[f201,f184])).
fof(f184,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK12))))),
  inference(cnf_transformation,[],[f120])).
fof(f3738,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK11))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),X0),s(fun(sK10,bool),sK12))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK12)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK10,bool),sK11))))) )),
  inference(resolution,[],[f227,f185])).
fof(f185,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))),
  inference(cnf_transformation,[],[f120])).
fof(f227,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2))))))))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) => ! [X3] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) => ! [X3] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X1))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X3),s(fun(X0,bool),X2))))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) => ! [X9] : s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X7))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X9),s(fun(X6,bool),X8))))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_DISJOINT_UNION','thm.real_sigma.real_sigma_|49|8_prove_ltb')).
fof(f186,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),sK13),s(fun(sK10,bool),'const.pred_set.UNION_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),sK12))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),sK13),s(fun(sK10,bool),sK11))),s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK10,'type.realax.real'),sK13),s(fun(sK10,bool),sK12)))))),
  inference(cnf_transformation,[],[f120])).
% SZS output end Proof for thm.real_sigma
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 409
% Passive clauses: 2195
% Generated clauses: 6828
% Final active clauses: 404
% Final passive clauses: 1641
% Input formulas: 32
% Initial clauses: 136
% 
% Duplicate literals: 228
% Trivial inequalities: 17
% Fw subsumption resolutions: 761
% Fw demodulations: 108
% Bw demodulations: 5
% 
% Simple tautologies: 151
% Equational tautologies: 32
% Forward subsumptions: 3105
% 
% Binary resolution: 3296
% Factoring: 775
% Forward superposition: 503
% Backward superposition: 860
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 251
% Split components: 508
% SAT solver clauses: 381
% SAT solver unit clauses: 14
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 1
% 
% Sat splits: 251
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 380
% TWLsolver calls for satisfiability: 202
% 
% Memory used [KB]: 3454
% Time elapsed: 0.148 s
% ------------------------------
% ------------------------------
% Success in time 0.149 s
