Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bool
% (97446)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 288
% Passive clauses: 7028
% Generated clauses: 28066
% Final active clauses: 231
% Final passive clauses: 6362
% Input formulas: 69
% Initial clauses: 161
% 
% Duplicate literals: 121
% Trivial inequalities: 243
% Fw subsumption resolutions: 2166
% Fw demodulations: 6061
% Bw demodulations: 39
% 
% Simple tautologies: 276
% Equational tautologies: 16
% Forward subsumptions: 12061
% 
% Binary resolution: 1839
% Factoring: 1
% Forward superposition: 8172
% Backward superposition: 9298
% Self superposition: 61
% Equality factoring: 18
% Equality resolution: 11
% 
% Split clauses: 169
% Split components: 339
% SAT solver clauses: 280
% SAT solver unit clauses: 8
% SAT solver binary clauses: 43
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 2
% 
% Sat splits: 169
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 279
% TWLsolver calls for satisfiability: 89
% 
% Memory used [KB]: 15863
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bool
% (97511)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 162
% Passive clauses: 2004
% Generated clauses: 19251
% Final active clauses: 144
% Final passive clauses: 1571
% Input formulas: 69
% Initial clauses: 161
% 
% Duplicate literals: 478
% Trivial inequalities: 22
% Fw subsumption resolutions: 1235
% Fw demodulations: 3035
% Bw demodulations: 282
% Condensations: 814
% Global subsumptions: 1
% 
% Simple tautologies: 79
% Equational tautologies: 27
% Forward subsumptions: 8566
% 
% Binary resolution: 751
% Factoring: 1
% Forward superposition: 5768
% Backward superposition: 4523
% Self superposition: 80
% Equality factoring: 42
% Equality resolution: 1
% 
% Split clauses: 171
% Split components: 343
% SAT solver clauses: 3076
% SAT solver unit clauses: 166
% SAT solver binary clauses: 65
% SAT solver learnt clauses: 2
% 
% Sat splits: 171
% Sat splitting refutations: 16
% 
% TWLsolver clauses: 3072
% TWLsolver calls for satisfiability: 1462
% 
% Memory used [KB]: 14839
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bool
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bool
% SZS output start Proof for thm.bool
fof(f12238,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f351,f515,f513,f511,f350,f516,f349,f512,f510,f544,f552,f550,f548,f558,f556,f690,f695,f692,f833,f2029,f2056,f2055,f2068,f2386,f2400,f3775,f3804,f3802,f3952,f6498,f6499,f6577,f6578,f691,f9121,f9125,f9123,f11340,f11631,f11632,f11633,f11634,f11635,f11636,f11637,f11638,f11639,f11640,f11641,f11642,f11643,f11825,f11826,f11827,f11828,f11829,f11830,f11831,f11832,f11833,f11834,f11835,f11836,f11837,f11838,f11839,f11840,f11841,f11842,f11843,f11844,f11845,f11846,f11847,f11848,f11849,f11850,f11851,f11852,f11853,f11854,f11855,f11856,f11857,f11858,f11859,f11860,f11861,f11862,f11863,f11864,f11865,f11866,f11867,f11868,f11869,f11870,f11871,f11872,f11873,f11874,f11875,f11876,f11877,f11878,f11879,f11880,f11881,f11882,f11883,f11884,f11885,f11886,f11887,f11888,f11889,f11890,f11891,f11892,f11893,f11894,f11895,f11896,f11897,f11898,f11899,f11900,f11901,f11902,f11903,f11904,f11905,f11906,f11907,f11908,f11909,f11910,f11911,f11912,f11938,f11939,f11940,f11941,f11942,f11943,f11944,f11945,f11946,f11947,f11948,f11949,f11950,f12235,f12237])).
fof(f12237,plain,(
  $false | ($spl1 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f12236,f352])).
fof(f352,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bool.UEXISTS_SIMP',aHL_TRUTH)).
fof(f12236,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f12201,f481])).
fof(f481,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : ! [X12,X13] : s(X11,X12) = s(X11,happ(s(fun(X10,X11),'_dst_x0_1'(s(X11,X12))),s(X10,X13)))),
  file('thm.bool.UEXISTS_SIMP','_dst_x0_ax')).
fof(f12201,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(sK0,X0))))) ) | ($spl1 | $spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f516,f11988,f370])).
fof(f370,plain,(
  ( ! [X4,X0,X1] : (s(X0,sK7(X1,X0)) != s(X0,sK8(X1,X0)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,X0))))) & s(X0,sK7(X1,X0)) != s(X0,sK8(X1,X0))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f238])).
fof(f238,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s(X0,X2) = s(X0,X3))))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X14] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X6,bool),X14)))) <=> (? [X8] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X8)))) & ! [X8,X9] : ((p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X8)))) & p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X9))))) => s(X6,X8) = s(X6,X9))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.EXISTS_UNIQUE_THM')).
fof(f11988,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f511,f2406])).
fof(f2406,plain,(
  s(bool,'T_0') = s(bool,sK1) | $spl2),
  inference(unit_resulting_resolution,[],[f512,f541])).
fof(f541,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f353,f355])).
fof(f355,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bool.UEXISTS_SIMP',aHL_BOOL_CASES)).
fof(f353,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bool.UEXISTS_SIMP',aHL_FALSITY)).
fof(f12235,plain,(
  $false | ($spl1 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f12234,f352])).
fof(f12234,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f12202,f3149])).
fof(f3149,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(superposition,[],[f360,f481])).
fof(f360,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.EXISTS_DEF')).
fof(f12202,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f516,f11988,f365])).
fof(f365,plain,(
  ( ! [X0,X1] : (s(X0,sK4(X1,X0)) != s(X0,sK5(X1,X0)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X4) = s(X0,X5)))) & (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0))))) & s(X0,sK4(X1,X0)) != s(X0,sK5(X1,X0))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f234])).
fof(f234,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X4,X5] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X4) = s(X0,X5)))) & (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s(X0,X2) = s(X0,X3))))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X5] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X6,bool),X5)))) <=> (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5)))) & ! [X8,X9] : ((p(s(bool,happ(s(fun(X6,bool),X5),s(X6,X8)))) & p(s(bool,happ(s(fun(X6,bool),X5),s(X6,X9))))) => s(X6,X8) = s(X6,X9))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.EXISTS_UNIQUE_DEF')).
fof(f11950,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11548,f11338])).
fof(f11338,plain,(
  ( ! [X0,X1] : (s(sK0,X0) = s(sK0,X1)) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f11337,f352])).
fof(f11337,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(sK0,X0) = s(sK0,X1)) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f11336,f481])).
fof(f11336,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(sK0,X1)))) | s(sK0,X0) = s(sK0,X1)) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f11335,f352])).
fof(f11335,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(sK0,X1)))) | s(sK0,X0) = s(sK0,X1)) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f11334,f481])).
fof(f11334,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(sK0,X1)))) | s(sK0,X0) = s(sK0,X1)) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f11321,f352])).
fof(f11321,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(sK0,X1)))) | s(sK0,X0) = s(sK0,X1)) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f362,f2473])).
fof(f2473,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f2409,f541])).
fof(f2409,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f2406,f510])).
fof(f362,plain,(
  ( ! [X4,X0,X5,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | s(X0,X4) = s(X0,X5)) )),
  inference(cnf_transformation,[],[f235])).
fof(f11548,plain,(
  ( ! [X51] : (s(sK0,sK3) != s(sK0,X51)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6802,f11338])).
fof(f6802,plain,(
  s(sK0,sK3) != s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)) | $spl5),
  inference(unit_resulting_resolution,[],[f6208,f501])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X17,X18,X19] : (p(s(bool,happ(s(fun(X17,bool),'=_1'(s(X17,X18))),s(X17,X19)))) <=> s(X17,X18) = s(X17,X19))),
  file('thm.bool.UEXISTS_SIMP',arity_1_HL_EQ)).
fof(f6208,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK3))),s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2732,f380])).
fof(f380,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6] : ! [X14] : (~! [X8] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X8)))) <=> ? [X8] : ~p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X8)))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.NOT_FORALL_THM')).
fof(f2732,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK3))),s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2698,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,X0))))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,X0))))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f241])).
fof(f241,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6] : ! [X34] : (p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X34)))) <=> ! [X8] : p(s(bool,happ(s(fun(X6,bool),X34),s(X6,X8)))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.FORALL_THM')).
fof(f2698,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK3)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2412,f371])).
fof(f371,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f2412,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK3))),s(sK0,sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f515,f500])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f347])).
fof(f11949,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11547,f11338])).
fof(f11547,plain,(
  ( ! [X50] : (s(sK0,sK2) != s(sK0,X50)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6387,f11338])).
fof(f6387,plain,(
  s(sK0,sK2) != s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)) | $spl5),
  inference(unit_resulting_resolution,[],[f6156,f501])).
fof(f6156,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK2))),s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2575,f380])).
fof(f2575,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK2))),s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2556,f372])).
fof(f2556,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK2)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2411,f371])).
fof(f2411,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK2))),s(sK0,sK3)))) | $spl5),
  inference(unit_resulting_resolution,[],[f515,f500])).
fof(f11948,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11544,f11338])).
fof(f11544,plain,(
  ( ! [X47] : (s(sK0,sK3) != s(sK0,X47)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6662,f11338])).
fof(f6662,plain,(
  s(sK0,sK3) != s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)) | $spl5),
  inference(unit_resulting_resolution,[],[f6207,f501])).
fof(f6207,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK3))),s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2732,f379])).
fof(f379,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f11947,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11541,f11338])).
fof(f11541,plain,(
  ( ! [X44] : (s(sK0,sK2) != s(sK0,X44)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6259,f11338])).
fof(f6259,plain,(
  s(sK0,sK2) != s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)) | $spl5),
  inference(unit_resulting_resolution,[],[f6155,f501])).
fof(f6155,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK2))),s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2575,f379])).
fof(f11946,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11538,f11338])).
fof(f11538,plain,(
  ( ! [X41] : (s(sK0,sK3) != s(sK0,X41)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6205,f11338])).
fof(f6205,plain,(
  s(sK0,sK3) != s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)) | $spl5),
  inference(unit_resulting_resolution,[],[f2732,f501])).
fof(f11945,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11535,f11338])).
fof(f11535,plain,(
  ( ! [X38] : (s(sK0,sK2) != s(sK0,X38)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6153,f11338])).
fof(f6153,plain,(
  s(sK0,sK2) != s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)) | $spl5),
  inference(unit_resulting_resolution,[],[f2575,f501])).
fof(f11944,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11528,f11338])).
fof(f11528,plain,(
  ( ! [X31] : (s(sK0,X31) != s(sK0,sK14('=_1'(s(sK0,X31)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6802,f11338])).
fof(f11943,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11527,f11338])).
fof(f11527,plain,(
  ( ! [X30] : (s(sK0,X30) != s(sK0,sK13('=_1'(s(sK0,X30)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6662,f11338])).
fof(f11942,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11526,f11338])).
fof(f11526,plain,(
  ( ! [X29] : (s(sK0,X29) != s(sK0,sK9('=_1'(s(sK0,X29)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6205,f11338])).
fof(f11941,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11519,f11338])).
fof(f11519,plain,(
  ( ! [X19] : (s(sK0,sK2) != s(sK0,X19)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f515,f11338])).
fof(f11940,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11512,f11338])).
fof(f11512,plain,(
  ( ! [X12] : (s(sK0,X12) != s(sK0,sK14('=_1'(s(sK0,X12)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6387,f11338])).
fof(f11939,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11511,f11338])).
fof(f11511,plain,(
  ( ! [X11] : (s(sK0,X11) != s(sK0,sK13('=_1'(s(sK0,X11)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6259,f11338])).
fof(f11938,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11510,f11338])).
fof(f11510,plain,(
  ( ! [X10] : (s(sK0,X10) != s(sK0,sK9('=_1'(s(sK0,X10)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6153,f11338])).
fof(f11912,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11794,f2272])).
fof(f2272,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'=_1'(s(X0,X1))))))) )),
  inference(forward_demodulation,[],[f2211,f360])).
fof(f2211,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1))))))))) )),
  inference(unit_resulting_resolution,[],[f357,f501])).
fof(f357,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1)))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.min.@_1'(s(fun(X0,bool),'=_1'(s(X0,X1)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X8] : s(X6,X8) = s(X6,'const.min.@_1'(s(fun(X6,bool),'=_1'(s(X6,X8)))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.SELECT_REFL_2')).
fof(f11794,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11220])).
fof(f11220,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK45('F_0','=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6868,f447])).
fof(f447,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK45(X2,X1,X0))))) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((! [X6] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X2,X1,X0))))) | p(s(bool,X2))) & ((! [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK45(X2,X1,X0))))) | p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK44,sK45])],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((! [X6] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2))) & ((! [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))) & ((! [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2] : ((! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & ((! [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : (? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : ! [X14,X15] : (? [X8] : (p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X8)))) => p(s(bool,X15))) <=> (! [X8] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X8)))) => p(s(bool,X15))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.LEFT_EXISTS_IMP_THM')).
fof(f6868,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6802,f500])).
fof(f11480,plain,(
  ( ! [X171,X172] : (s(bool,'|dollar|exists_1'(s(fun(sK0,bool),X171))) = s(bool,happ(s(fun(sK0,bool),X171),s(sK0,X172)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f360,f11338])).
fof(f11911,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11793,f2272])).
fof(f11793,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11225])).
fof(f11225,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK45(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6868,f6868,f447])).
fof(f11910,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11792,f2272])).
fof(f11792,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10375])).
fof(f10375,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK45('F_0','=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6659,f447])).
fof(f6659,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6387,f500])).
fof(f11909,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11791,f2272])).
fof(f11791,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10380])).
fof(f10380,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK45(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6659,f6659,f447])).
fof(f11908,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11790,f2272])).
fof(f11790,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10634])).
fof(f10634,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK45('F_0','=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6799,f447])).
fof(f6799,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6662,f500])).
fof(f11907,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11789,f2272])).
fof(f11789,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10639])).
fof(f10639,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK45(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6799,f6799,f447])).
fof(f11906,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11788,f2272])).
fof(f11788,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10055])).
fof(f10055,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK45('F_0','=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6384,f447])).
fof(f6384,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6259,f500])).
fof(f11905,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11787,f2272])).
fof(f11787,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10060])).
fof(f10060,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK45(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6384,f6384,f447])).
fof(f11904,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11786,f2272])).
fof(f11786,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9046])).
fof(f9046,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK45('F_0','=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6256,f447])).
fof(f6256,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6205,f500])).
fof(f11903,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11785,f2272])).
fof(f11785,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9051])).
fof(f9051,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK45(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6256,f6256,f447])).
fof(f11902,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11784,f2272])).
fof(f11784,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11218])).
fof(f11218,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK44('F_0','=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6868,f445])).
fof(f445,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X2,X1,X0))))) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f315])).
fof(f11901,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11783,f2272])).
fof(f11783,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11224])).
fof(f11224,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK44(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6868,f6868,f445])).
fof(f11900,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11782,f2272])).
fof(f11782,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10373])).
fof(f10373,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK44('F_0','=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6659,f445])).
fof(f11899,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11781,f2272])).
fof(f11781,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10379])).
fof(f10379,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK44(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6659,f6659,f445])).
fof(f11898,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11780,f2272])).
fof(f11780,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10632])).
fof(f10632,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK44('F_0','=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6799,f445])).
fof(f11897,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11779,f2272])).
fof(f11779,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10638])).
fof(f10638,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK44(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6799,f6799,f445])).
fof(f11896,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11778,f2272])).
fof(f11778,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10053])).
fof(f10053,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK44('F_0','=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6384,f445])).
fof(f11895,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11777,f2272])).
fof(f11777,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10059])).
fof(f10059,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK44(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6384,f6384,f445])).
fof(f11894,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11776,f2272])).
fof(f11776,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9044])).
fof(f9044,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK44('F_0','=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6256,f445])).
fof(f11893,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11775,f2272])).
fof(f11775,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9050])).
fof(f9050,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK44(happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6256,f6256,f445])).
fof(f11892,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11774,f2272])).
fof(f11774,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11202])).
fof(f11202,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6868,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK32(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK31(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK32(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31,sK32])],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X15,X14] : (! [X8] : (p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X8)))) | p(s(bool,X15))) <=> (! [X8] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X8)))) | p(s(bool,X15))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.LEFT_FORALL_OR_THM')).
fof(f11891,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11773,f2272])).
fof(f11773,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11223])).
fof(f11223,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6868,f6868,f419])).
fof(f11890,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11772,f2272])).
fof(f11772,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10357])).
fof(f10357,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6659,f419])).
fof(f11889,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11771,f2272])).
fof(f11771,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10378])).
fof(f10378,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6659,f6659,f419])).
fof(f11888,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11770,f2272])).
fof(f11770,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10618])).
fof(f10618,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6799,f419])).
fof(f11887,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11769,f2272])).
fof(f11769,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10637])).
fof(f10637,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6799,f6799,f419])).
fof(f11886,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11768,f2272])).
fof(f11768,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10039])).
fof(f10039,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6384,f419])).
fof(f11885,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11767,f2272])).
fof(f11767,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10058])).
fof(f10058,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6384,f6384,f419])).
fof(f11884,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11766,f2272])).
fof(f11766,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9032])).
fof(f9032,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6256,f419])).
fof(f11883,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11765,f2272])).
fof(f11765,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9049])).
fof(f9049,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6256,f6256,f419])).
fof(f11882,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11764,f2272])).
fof(f11764,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8798])).
fof(f8798,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6202,f419])).
fof(f6202,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6153,f500])).
fof(f11881,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11763,f2272])).
fof(f11763,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8811])).
fof(f8811,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK32('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6202,f6202,f419])).
fof(f11880,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11762,f2272])).
fof(f11762,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11200])).
fof(f11200,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6868,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK31(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f287])).
fof(f11879,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11761,f2272])).
fof(f11761,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11222])).
fof(f11222,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6868,f6868,f417])).
fof(f11878,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11760,f2272])).
fof(f11760,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10355])).
fof(f10355,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6659,f417])).
fof(f11877,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11759,f2272])).
fof(f11759,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10377])).
fof(f10377,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6659,f6659,f417])).
fof(f11876,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11758,f2272])).
fof(f11758,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10616])).
fof(f10616,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6799,f417])).
fof(f11875,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11757,f2272])).
fof(f11757,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10636])).
fof(f10636,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6799,f6799,f417])).
fof(f11874,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11756,f2272])).
fof(f11756,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10037])).
fof(f10037,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6384,f417])).
fof(f11873,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11755,f2272])).
fof(f11755,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10057])).
fof(f10057,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6384,f6384,f417])).
fof(f11872,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11754,f2272])).
fof(f11754,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9030])).
fof(f9030,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6256,f417])).
fof(f11871,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11753,f2272])).
fof(f11753,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9048])).
fof(f9048,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6256,f6256,f417])).
fof(f11870,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11752,f2272])).
fof(f11752,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8796])).
fof(f8796,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),'F_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f353,f6202,f417])).
fof(f11869,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11751,f2272])).
fof(f11751,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8810])).
fof(f8810,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK31('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6202,f6202,f417])).
fof(f11868,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11750,f2272])).
fof(f11750,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11199])).
fof(f11199,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK25('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6868,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24,sK25])],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6] : ! [X14,X15] : ((p(s(bool,X14)) & ! [X8] : p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X8))))) <=> ! [X8] : (p(s(bool,X14)) & p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X8))))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f11867,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11749,f2272])).
fof(f11749,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10354])).
fof(f10354,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK25('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6659,f408])).
fof(f11866,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11748,f2272])).
fof(f11748,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10615])).
fof(f10615,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK25('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6799,f408])).
fof(f11865,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11747,f2272])).
fof(f11747,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10036])).
fof(f10036,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK25('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6384,f408])).
fof(f11864,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11746,f2272])).
fof(f11746,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9029])).
fof(f9029,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK25('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6256,f408])).
fof(f11863,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11745,f2272])).
fof(f11745,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8795])).
fof(f8795,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK25('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6202,f408])).
fof(f11862,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11744,f2272])).
fof(f11744,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11198])).
fof(f11198,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK24('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6868,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f275])).
fof(f11861,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11743,f2272])).
fof(f11743,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10353])).
fof(f10353,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK24('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6659,f406])).
fof(f11860,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11742,f2272])).
fof(f11742,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10614])).
fof(f10614,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK24('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6799,f406])).
fof(f11859,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11741,f2272])).
fof(f11741,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10035])).
fof(f10035,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK24('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6384,f406])).
fof(f11858,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11740,f2272])).
fof(f11740,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9028])).
fof(f9028,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK24('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6256,f406])).
fof(f11857,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11739,f2272])).
fof(f11739,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8794])).
fof(f8794,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK24('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),'T_0',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f352,f6202,f406])).
fof(f11856,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11738,f2272])).
fof(f11738,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11197])).
fof(f11197,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK14('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6868,f380])).
fof(f11855,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11737,f2272])).
fof(f11737,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10352])).
fof(f10352,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK14('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6659,f380])).
fof(f11854,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11736,f2272])).
fof(f11736,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10613])).
fof(f10613,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK14('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6799,f380])).
fof(f11853,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11735,f2272])).
fof(f11735,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10034])).
fof(f10034,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK14('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6384,f380])).
fof(f11852,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11734,f2272])).
fof(f11734,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9027])).
fof(f9027,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK14('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6256,f380])).
fof(f11851,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11733,f2272])).
fof(f11733,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8793])).
fof(f8793,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK14('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6202,f380])).
fof(f11850,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11732,f2272])).
fof(f11732,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11196])).
fof(f11196,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK13('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6868,f379])).
fof(f11849,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11731,f2272])).
fof(f11731,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10351])).
fof(f10351,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK13('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6659,f379])).
fof(f11848,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11730,f2272])).
fof(f11730,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10612])).
fof(f10612,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK13('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6799,f379])).
fof(f11847,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11729,f2272])).
fof(f11729,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10033])).
fof(f10033,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK13('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6384,f379])).
fof(f11846,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11728,f2272])).
fof(f11728,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9026])).
fof(f9026,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK13('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6256,f379])).
fof(f11845,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11727,f2272])).
fof(f11727,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8792])).
fof(f8792,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK13('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6202,f379])).
fof(f11844,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11726,f2272])).
fof(f11726,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f11284])).
fof(f11284,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK9('=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f11195,f372])).
fof(f11195,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6868,f371])).
fof(f11843,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11725,f2272])).
fof(f11725,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10439])).
fof(f10439,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK9('=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f10350,f372])).
fof(f10350,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6659,f371])).
fof(f11842,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11724,f2272])).
fof(f11724,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10694])).
fof(f10694,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK9('=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f10611,f372])).
fof(f10611,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6799,f371])).
fof(f11841,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11723,f2272])).
fof(f11723,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f10115])).
fof(f10115,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK9('=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f10032,f372])).
fof(f10032,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6384,f371])).
fof(f11840,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11722,f2272])).
fof(f11722,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f9337])).
fof(f9337,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK9('=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f9025,f372])).
fof(f9025,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6256,f371])).
fof(f11839,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11721,f2272])).
fof(f11721,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f8877])).
fof(f8877,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK9('=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f8791,f372])).
fof(f8791,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),'=_1'(s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6202,f371])).
fof(f11838,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11719,f1155])).
fof(f1155,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f358,f361])).
fof(f361,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f358,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X16] : p(s(bool,'const.bool.?|33|_1'(s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(X6,X16))))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.EXISTS_UNIQUE_REFL')).
fof(f11719,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6869])).
fof(f6869,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3))),s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(forward_demodulation,[],[f6867,f505])).
fof(f505,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11,X10,X20] : ! [X12,X21,X13] : s(X11,happ(s(fun(X10,X11),'_dst_x1x0_2'(s(fun(X10,fun(X20,X11)),X12),s(X20,X21))),s(X10,X13))) = s(X11,happ(s(fun(X20,X11),happ(s(fun(X10,fun(X20,X11)),X12),s(X10,X13))),s(X20,X21)))),
  file('thm.bool.UEXISTS_SIMP','_dst_x1x0_ax')).
fof(f6867,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK14('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6802,f502])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X17,X18,X19] : (p(s(bool,happ(s(fun(X17,bool),happ(s(fun(X17,fun(X17,bool)),'=_0'),s(X17,X18))),s(X17,X19)))) <=> s(X17,X18) = s(X17,X19))),
  file('thm.bool.UEXISTS_SIMP',arity_0_HL_EQ)).
fof(f11837,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11718,f1155])).
fof(f11718,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6660])).
fof(f6660,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2))),s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(forward_demodulation,[],[f6658,f505])).
fof(f6658,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK14('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6387,f502])).
fof(f11836,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11717,f1155])).
fof(f11717,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6800])).
fof(f6800,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3))),s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(forward_demodulation,[],[f6798,f505])).
fof(f6798,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK13('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6662,f502])).
fof(f11835,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11716,f1155])).
fof(f11716,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6385])).
fof(f6385,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2))),s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(forward_demodulation,[],[f6383,f505])).
fof(f6383,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK13('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6259,f502])).
fof(f11834,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11715,f1155])).
fof(f11715,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6257])).
fof(f6257,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3))),s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(forward_demodulation,[],[f6255,f505])).
fof(f6255,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK9('=_1'(s(sK0,sK3)),sK0)))),s(sK0,sK3)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6205,f502])).
fof(f11833,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11714,f1155])).
fof(f11714,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6203])).
fof(f6203,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),'_dst_x1x0_2'(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2))),s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(forward_demodulation,[],[f6201,f505])).
fof(f6201,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK9('=_1'(s(sK0,sK2)),sK0)))),s(sK0,sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f6153,f502])).
fof(f11832,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11713,f11703])).
fof(f11703,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,X0))))))) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f11480,f11344])).
fof(f11344,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,X0))),s(sK0,X1))))) ) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f11338,f503])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f348])).
fof(f11713,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6865])).
fof(f6865,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3))),s(sK0,sK14('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6802,f502])).
fof(f11831,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11712,f11703])).
fof(f11712,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6656])).
fof(f6656,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2))),s(sK0,sK14('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6387,f502])).
fof(f11830,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11711,f11703])).
fof(f11711,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6796])).
fof(f6796,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3))),s(sK0,sK13('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6662,f502])).
fof(f11829,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11710,f11703])).
fof(f11710,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6381])).
fof(f6381,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2))),s(sK0,sK13('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6259,f502])).
fof(f11828,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11709,f11703])).
fof(f11709,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6253])).
fof(f6253,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3))),s(sK0,sK9('=_1'(s(sK0,sK3)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6205,f502])).
fof(f11827,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11708,f11703])).
fof(f11708,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f6199])).
fof(f6199,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2))),s(sK0,sK9('=_1'(s(sK0,sK2)),sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f6153,f502])).
fof(f11826,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11706,f11703])).
fof(f11706,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f5429])).
fof(f5429,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK2))),s(sK0,sK3)))) | $spl5),
  inference(unit_resulting_resolution,[],[f515,f502])).
fof(f11825,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11704,f11703])).
fof(f11704,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f11480,f5431])).
fof(f5431,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'=_0'),s(sK0,sK3))),s(sK0,sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f515,f502])).
fof(f11643,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11433,f11338])).
fof(f11433,plain,(
  ( ! [X51] : (s(sK0,sK3) != s(sK0,X51)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6802,f11338])).
fof(f11642,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11432,f11338])).
fof(f11432,plain,(
  ( ! [X50] : (s(sK0,sK2) != s(sK0,X50)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6387,f11338])).
fof(f11641,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11429,f11338])).
fof(f11429,plain,(
  ( ! [X47] : (s(sK0,sK3) != s(sK0,X47)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6662,f11338])).
fof(f11640,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11426,f11338])).
fof(f11426,plain,(
  ( ! [X44] : (s(sK0,sK2) != s(sK0,X44)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6259,f11338])).
fof(f11639,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11423,f11338])).
fof(f11423,plain,(
  ( ! [X41] : (s(sK0,sK3) != s(sK0,X41)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6205,f11338])).
fof(f11638,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11420,f11338])).
fof(f11420,plain,(
  ( ! [X38] : (s(sK0,sK2) != s(sK0,X38)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6153,f11338])).
fof(f11637,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11413,f11338])).
fof(f11413,plain,(
  ( ! [X31] : (s(sK0,X31) != s(sK0,sK14('=_1'(s(sK0,X31)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6802,f11338])).
fof(f11636,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11412,f11338])).
fof(f11412,plain,(
  ( ! [X30] : (s(sK0,X30) != s(sK0,sK13('=_1'(s(sK0,X30)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6662,f11338])).
fof(f11635,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11411,f11338])).
fof(f11411,plain,(
  ( ! [X29] : (s(sK0,X29) != s(sK0,sK9('=_1'(s(sK0,X29)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6205,f11338])).
fof(f11634,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11404,f11338])).
fof(f11404,plain,(
  ( ! [X19] : (s(sK0,sK2) != s(sK0,X19)) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f515,f11338])).
fof(f11633,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11397,f11338])).
fof(f11397,plain,(
  ( ! [X12] : (s(sK0,X12) != s(sK0,sK14('=_1'(s(sK0,X12)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6387,f11338])).
fof(f11632,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11396,f11338])).
fof(f11396,plain,(
  ( ! [X11] : (s(sK0,X11) != s(sK0,sK13('=_1'(s(sK0,X11)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6259,f11338])).
fof(f11631,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11395,f11338])).
fof(f11395,plain,(
  ( ! [X10] : (s(sK0,X10) != s(sK0,sK9('=_1'(s(sK0,X10)),sK0))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f6153,f11338])).
fof(f11340,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(unit_resulting_resolution,[],[f515,f11338])).
fof(f9123,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) | $spl32),
  inference(cnf_transformation,[],[f9123_D])).
fof(f9123_D,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f9125,plain,(
  p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,sK4('==>_1'(s(bool,'T_0')),bool))))) | $spl34),
  inference(cnf_transformation,[],[f9125_D])).
fof(f9125_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,sK4('==>_1'(s(bool,'T_0')),bool))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f9121,plain,(
  p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,sK4('==>_1'(s(bool,'T_0')),bool))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) | $spl26),
  inference(subsumption_resolution,[],[f9107,f463])).
fof(f463,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X18,X19] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X18))),s(bool,X19)))) <=> (p(s(bool,X18)) => p(s(bool,X19))))),
  file('thm.bool.UEXISTS_SIMP','==>_HL_1')).
fof(f9107,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,sK4('==>_1'(s(bool,'T_0')),bool))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) | $spl26),
  inference(superposition,[],[f363,f4086])).
fof(f4086,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))))) | $spl26),
  inference(unit_resulting_resolution,[],[f3803,f541])).
fof(f3803,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) | $spl26),
  inference(cnf_transformation,[],[f3803_D])).
fof(f3803_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f363,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f691,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(cnf_transformation,[],[f691_D])).
fof(f691_D,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f6578,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f6578_D])).
fof(f6578_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f6577,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6459,f541])).
fof(f6459,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f468,f355])).
fof(f468,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X18,X19] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X18))),s(bool,X19)))) <=> (p(s(bool,X18)) | p(s(bool,X19))))),
  file('thm.bool.UEXISTS_SIMP','|bslash||slash|_HL_0')).
fof(f6499,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f6499_D])).
fof(f6499_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f6498,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6422,f541])).
fof(f6422,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f468,f355])).
fof(f3952,plain,(
  $false | $spl27),
  inference(unit_resulting_resolution,[],[f352,f3870,f464])).
fof(f464,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f322])).
fof(f3870,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X0))))) ) | $spl27),
  inference(unit_resulting_resolution,[],[f3804,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f247])).
fof(f247,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : ! [X34] : (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X34)))) <=> ? [X8] : p(s(bool,happ(s(fun(X6,bool),X34),s(X6,X8)))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.EXISTS_THM')).
fof(f3802,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) | $spl24),
  inference(cnf_transformation,[],[f3802_D])).
fof(f3802_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3804,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) | $spl27),
  inference(cnf_transformation,[],[f3804_D])).
fof(f3804_D,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f3775,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) | p(s(bool,'const.min.@_1'(s(fun(bool,bool),'==>_1'(s(bool,'T_0')))))) | $spl2),
  inference(superposition,[],[f2723,f360])).
fof(f2723,plain,(
  ( ! [X8] : (~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X8)))) | p(s(bool,X8))) ) | $spl2),
  inference(forward_demodulation,[],[f2722,f2406])).
fof(f2722,plain,(
  ( ! [X8] : (p(s(bool,X8)) | ~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,sK1))),s(bool,X8))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2712,f352])).
fof(f2712,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | p(s(bool,X8)) | ~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,sK1))),s(bool,X8))))) ) | $spl2),
  inference(superposition,[],[f462,f2406])).
fof(f462,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f322])).
fof(f2400,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f515,f516])).
fof(f2386,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2385,f353])).
fof(f2385,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2384,f481])).
fof(f2384,plain,(
  p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(sK0,sK6('_dst_x0_1'(s(bool,'F_0')),sK0))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2359,f352])).
fof(f2359,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(sK0,sK6('_dst_x0_1'(s(bool,'F_0')),sK0))))) | ($spl0 | $spl3)),
  inference(superposition,[],[f366,f2076])).
fof(f2076,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0'))))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f2065,f541])).
fof(f2065,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0')))))) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f2062,f510])).
fof(f2062,plain,(
  s(bool,'F_0') = s(bool,sK1) | $spl3),
  inference(unit_resulting_resolution,[],[f353,f513,f458])).
fof(f458,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X24,X25] : ((p(s(bool,X24)) => p(s(bool,X25))) => ((p(s(bool,X25)) => p(s(bool,X24))) => s(bool,X24) = s(bool,X25)))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.IMP_ANTISYM_AX')).
fof(f366,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0)))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f2068,plain,(
  $false | ($spl0 | $spl3 | $spl11 | $spl15)),
  inference(subsumption_resolution,[],[f2067,f353])).
fof(f2067,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3 | $spl11 | $spl15)),
  inference(forward_demodulation,[],[f2065,f672])).
fof(f672,plain,(
  s(bool,'F_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0'))))) | ($spl11 | $spl15)),
  inference(backward_demodulation,[],[f651,f653])).
fof(f653,plain,(
  ( ! [X2,X0,X1] : (s(bool,sK47(X0,X1,'F_0',X2)) = s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0')))))) ) | ($spl11 | $spl15)),
  inference(unit_resulting_resolution,[],[f552,f621,f458])).
fof(f621,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK47(X0,X1,'F_0',X2)))) ) | $spl15),
  inference(unit_resulting_resolution,[],[f558,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK47(X3,X2,X1,X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,sK46(X3,X2,X1,X0))) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,sK46(X3,X2,X1,X0)))) & ((~p(s(bool,sK47(X3,X2,X1,X0))) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,sK47(X3,X2,X1,X0)))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,sK47(X3,X2,X1,X0)))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,sK46(X3,X2,X1,X0)))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK46,sK47])],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,X4))) & ? [X5] : ((~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,X5))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (((~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,X4)))) & ? [X5] : (((~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,X5)))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s(bool,'F_0') = s(bool,X1)) & ? [X5] : ((p(s(bool,X5)) <=> s(bool,'T_0') = s(bool,X1)) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s(bool,'F_0') = s(bool,X1)) & ? [X5] : ((p(s(bool,X5)) <=> s(bool,'T_0') = s(bool,X1)) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X5,X35,X32] : ? [X37] : ((p(s(bool,X37)) <=> s(bool,'F_0') = s(bool,X5)) & ? [X38] : ((p(s(bool,X38)) <=> s(bool,'T_0') = s(bool,X5)) & s(X6,'const.bool.COND_3'(s(bool,X5),s(X6,X35),s(X6,X32))) = s(X6,'const.min.@_1'(s(fun(X6,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X38))),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(X6,X35))))),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X37))),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(X6,X32)))))))))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.COND_DEF')).
fof(f651,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,sK47(X0,X1,'F_0',X2))) ) | $spl15),
  inference(unit_resulting_resolution,[],[f353,f621,f458])).
fof(f2055,plain,(
  ( ! [X7] : (p(s(bool,sK33(sK1,'_dst_x0_1'(s(bool,X7)),bool))) | s(bool,'T_0') = s(bool,X7)) ) | $spl20),
  inference(cnf_transformation,[],[f2055_D])).
fof(f2055_D,plain,(
  ( ! [X7] : (p(s(bool,sK33(sK1,'_dst_x0_1'(s(bool,X7)),bool))) | s(bool,'T_0') = s(bool,X7)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2056,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X8))))) ) | $spl22),
  inference(cnf_transformation,[],[f2056_D])).
fof(f2056_D,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X8))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2029,plain,(
  ( ! [X8,X7] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X8)))) | p(s(bool,sK33(sK1,'_dst_x0_1'(s(bool,X7)),bool))) | s(bool,'T_0') = s(bool,X7)) ) | $spl2),
  inference(superposition,[],[f463,f1735])).
fof(f1735,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK33(sK1,'_dst_x0_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f1713,f353])).
fof(f1713,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK33(sK1,'_dst_x0_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(superposition,[],[f1710,f355])).
fof(f1710,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK33(sK1,'_dst_x0_1'(s(bool,X1)),bool))) ) | $spl2),
  inference(superposition,[],[f1692,f481])).
fof(f1692,plain,(
  ( ! [X30] : (p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK33(sK1,X30,bool))) ) | $spl2),
  inference(forward_demodulation,[],[f1672,f553])).
fof(f553,plain,(
  s(bool,'T_0') = s(bool,sK1) | $spl2),
  inference(subsumption_resolution,[],[f545,f353])).
fof(f545,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK1) | $spl2),
  inference(superposition,[],[f512,f355])).
fof(f1672,plain,(
  ( ! [X30] : (p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | s(bool,sK1) = s(bool,sK33(sK1,X30,bool))) ) | $spl2),
  inference(superposition,[],[f421,f553])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK33(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : (((s(X0,X2) = s(X0,sK33(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK33(X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33])],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((? [X4] : (s(X0,X2) = s(X0,X4) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X34,X36] : (! [X8] : (s(X6,X8) = s(X6,X36) => p(s(bool,happ(s(fun(X6,bool),X34),s(X6,X8))))) <=> p(s(bool,happ(s(fun(X6,bool),X34),s(X6,X36)))))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.UNWIND_FORALL_THM|49|')).
fof(f833,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f698,f352])).
fof(f698,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f691,f353])).
fof(f692,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'F_0') = s(bool,X3)) ) | $spl18),
  inference(cnf_transformation,[],[f692_D])).
fof(f692_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'F_0') = s(bool,X3)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f695,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f666,f541])).
fof(f666,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f458,f355])).
fof(f690,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f660,f541])).
fof(f660,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f458,f355])).
fof(f556,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f556_D])).
fof(f556_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f558,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl15),
  inference(cnf_transformation,[],[f558_D])).
fof(f558_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f548,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f355])).
fof(f550,plain,(
  s(bool,'T_0') = s(bool,sK1) | $spl8),
  inference(cnf_transformation,[],[f550_D])).
fof(f550_D,plain,(
  s(bool,'T_0') = s(bool,sK1) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f552,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0')))))) | $spl11),
  inference(cnf_transformation,[],[f552_D])).
fof(f552_D,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0')))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f544,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,sK1) | $spl1),
  inference(superposition,[],[f511,f355])).
fof(f510,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1)))))) | $spl0),
  inference(cnf_transformation,[],[f510_D])).
fof(f510_D,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f512,plain,(
  p(s(bool,sK1)) | $spl2),
  inference(cnf_transformation,[],[f512_D])).
fof(f512_D,plain,(
  p(s(bool,sK1)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f349,plain,(
  p(s(bool,sK1)) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1))))))),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  (p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1)))))) | (p(s(bool,sK1)) & ! [X4,X5] : s(sK0,X4) = s(sK0,X5))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1)))))) | ~p(s(bool,sK1)) | s(sK0,sK2) != s(sK0,sK3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f230])).
fof(f230,plain,(
  ? [X0,X1] : ((p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) | (p(s(bool,X1)) & ! [X4,X5] : s(X0,X4) = s(X0,X5))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) | ~p(s(bool,X1)) | ? [X2,X3] : s(X0,X2) != s(X0,X3)))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ? [X0,X1] : ((p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) | (p(s(bool,X1)) & ! [X2,X3] : s(X0,X2) = s(X0,X3))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) | ~p(s(bool,X1)) | ? [X2,X3] : s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ? [X0,X1] : ((p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) | (p(s(bool,X1)) & ! [X2,X3] : s(X0,X2) = s(X0,X3))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) | (~p(s(bool,X1)) | ? [X2,X3] : s(X0,X2) != s(X0,X3))))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ? [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) <~> (p(s(bool,X1)) & ! [X2,X3] : s(X0,X2) = s(X0,X3)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) <=> (p(s(bool,X1)) & ! [X2,X3] : s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) <=> (p(s(bool,X1)) & ! [X2,X3] : s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X6,bool),'_dst_x0_1'(s(bool,X7)))))) <=> (p(s(bool,X7)) & ! [X8,X9] : s(X6,X8) = s(X6,X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X6,bool),'_dst_x0_1'(s(bool,X7)))))) <=> (p(s(bool,X7)) & ! [X8,X9] : s(X6,X8) = s(X6,X9)))),
  file('thm.bool.UEXISTS_SIMP','thm.bool.UEXISTS_SIMP')).
fof(f516,plain,(
  ( ! [X4,X5] : (s(sK0,X4) = s(sK0,X5)) ) | $spl6),
  inference(cnf_transformation,[],[f516_D])).
fof(f516_D,plain,(
  ( ! [X4,X5] : (s(sK0,X4) = s(sK0,X5)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f350,plain,(
  ( ! [X4,X5] : (s(sK0,X4) = s(sK0,X5) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1))))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f511,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1)))))) | $spl1),
  inference(cnf_transformation,[],[f511_D])).
fof(f511_D,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f513,plain,(
  ~p(s(bool,sK1)) | $spl3),
  inference(cnf_transformation,[],[f513_D])).
fof(f513_D,plain,(
  ~p(s(bool,sK1)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f515,plain,(
  s(sK0,sK2) != s(sK0,sK3) | $spl5),
  inference(cnf_transformation,[],[f515_D])).
fof(f515_D,plain,(
  s(sK0,sK2) != s(sK0,sK3) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f351,plain,(
  s(sK0,sK2) != s(sK0,sK3) | ~p(s(bool,sK1)) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK0,bool),'_dst_x0_1'(s(bool,sK1))))))),
  inference(cnf_transformation,[],[f231])).
% SZS output end Proof for thm.bool
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 275
% Passive clauses: 1768
% Generated clauses: 12075
% Final active clauses: 120
% Final passive clauses: 842
% Input formulas: 69
% Initial clauses: 161
% 
% Duplicate literals: 15
% Trivial inequalities: 1
% Fw subsumption resolutions: 785
% Fw demodulations: 3525
% Bw demodulations: 476
% Condensations: 110
% 
% Simple tautologies: 82
% Equational tautologies: 87
% Forward subsumptions: 5568
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 1632
% Factoring: 1
% Forward superposition: 2435
% Backward superposition: 2687
% Self superposition: 26
% Equality factoring: 2
% 
% Split clauses: 12
% Split components: 25
% SAT solver clauses: 1916
% SAT solver unit clauses: 75
% SAT solver binary clauses: 718
% 
% Sat splits: 12
% Sat splitting refutations: 122
% 
% Memory used [KB]: 9850
% Time elapsed: 0.325 s
% ------------------------------
% ------------------------------
% Success in time 1.609 s
