Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% (46728)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1165
% Passive clauses: 6689
% Generated clauses: 35182
% Final active clauses: 903
% Final passive clauses: 4845
% Input formulas: 107
% Initial clauses: 251
% 
% Duplicate literals: 249
% Trivial inequalities: 62
% Fw subsumption resolutions: 2331
% Fw demodulations: 3354
% Bw demodulations: 231
% 
% Simple tautologies: 219
% Equational tautologies: 33
% Forward subsumptions: 22043
% Fw demodulations to eq. taut.: 9
% 
% Binary resolution: 7317
% Factoring: 1396
% Forward superposition: 9927
% Backward superposition: 9847
% Self superposition: 31
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 281
% Split components: 563
% SAT solver clauses: 532
% SAT solver unit clauses: 17
% SAT solver binary clauses: 57
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 281
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 531
% TWLsolver calls for satisfiability: 489
% 
% Memory used [KB]: 18038
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f11490,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f11489,f365,f625,f407,f669,f461,f690,f2404,f2425,f395,f657,f418,f674,f10269,f10292,f819,f364,f621,f363,f623,f619,f408,f670,f632,f634,f388,f650,f537,f736,f10320,f10381,f1006,f1017,f951,f960,f508,f725,f10214,f513,f722,f9879,f803,f804,f460,f689,f518,f727,f10206,f362,f620,f622,f618,f9965,f9967])).
fof(f9967,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl3291),
  inference(cnf_transformation,[],[f9967_D])).
fof(f9967_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) <=> ~$spl3291),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3291])])).
fof(f9965,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ($spl1 | $spl2 | $spl78 | $spl104)),
  inference(forward_demodulation,[],[f9964,f690])).
fof(f9964,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10)))))) | ($spl1 | $spl2 | $spl78)),
  inference(forward_demodulation,[],[f9911,f670])).
fof(f9911,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f621,f619])).
fof(f618,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) | $spl0),
  inference(cnf_transformation,[],[f618_D])).
fof(f618_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f622,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | $spl5),
  inference(cnf_transformation,[],[f622_D])).
fof(f622_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f620,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) | $spl3),
  inference(cnf_transformation,[],[f620_D])).
fof(f620_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f362,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11))))))),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11)) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) & s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f274])).
fof(f274,plain,(
  ? [X0,X1] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ? [X0,X1] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))))),
  inference(nnf_transformation,[],[f234])).
fof(f234,plain,(
  ? [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) <~> (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ~! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) <=> (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ~! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.real.lt_int|split|1','thm.real.lt_int|split|1')).
fof(f10206,plain,(
  $false | ($spl5 | $spl15 | $spl38 | $spl302 | $spl3291)),
  inference(subsumption_resolution,[],[f10205,f634])).
fof(f10205,plain,(
  p(s(bool,'F_0')) | ($spl5 | $spl38 | $spl302 | $spl3291)),
  inference(forward_demodulation,[],[f10204,f650])).
fof(f10204,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | ($spl5 | $spl302 | $spl3291)),
  inference(subsumption_resolution,[],[f10180,f622])).
fof(f10180,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | ($spl302 | $spl3291)),
  inference(resolution,[],[f9967,f804])).
fof(f727,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl168),
  inference(cnf_transformation,[],[f727_D])).
fof(f727_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f518,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : (~s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X7,X6] : (~s('type.num.num',X6) = s('type.num.num',X7) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7))))))),
  file('thm.real.lt_int|split|1','thm.arithmetic.NOT_NUM_EQ')).
fof(f689,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) ) | $spl102),
  inference(cnf_transformation,[],[f689_D])).
fof(f689_D,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f460,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X6] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6)))),
  file('thm.real.lt_int|split|1','thm.arithmetic.LESS_EQ')).
fof(f804,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl302),
  inference(cnf_transformation,[],[f804_D])).
fof(f804_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f803,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl102 | $spl168)),
  inference(forward_demodulation,[],[f801,f689])).
fof(f801,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl102 | $spl168)),
  inference(backward_demodulation,[],[f689,f727])).
fof(f9879,plain,(
  $false | ($spl3 | $spl15 | $spl38 | $spl302 | $spl3251)),
  inference(subsumption_resolution,[],[f9878,f634])).
fof(f9878,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl38 | $spl302 | $spl3251)),
  inference(forward_demodulation,[],[f9877,f650])).
fof(f9877,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num','const.num.|48|_0')))) | ($spl3 | $spl302 | $spl3251)),
  inference(subsumption_resolution,[],[f9853,f620])).
fof(f9853,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | ($spl302 | $spl3251)),
  inference(resolution,[],[f9852,f804])).
fof(f9852,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) | $spl3251),
  inference(cnf_transformation,[],[f9852_D])).
fof(f9852_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) <=> ~$spl3251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3251])])).
fof(f722,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl158),
  inference(cnf_transformation,[],[f722_D])).
fof(f722_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f317])).
fof(f317,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f10214,plain,(
  $false | ($spl15 | $spl158 | $spl3250 | $spl3271)),
  inference(subsumption_resolution,[],[f10213,f9851])).
fof(f9851,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) | $spl3250),
  inference(cnf_transformation,[],[f9851_D])).
fof(f9851_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) <=> ~$spl3250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3250])])).
fof(f10213,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) | ($spl15 | $spl158 | $spl3271)),
  inference(subsumption_resolution,[],[f10212,f634])).
fof(f10212,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) | ($spl158 | $spl3271)),
  inference(resolution,[],[f9895,f722])).
fof(f9895,plain,(
  ~sP1('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))) | $spl3271),
  inference(cnf_transformation,[],[f9895_D])).
fof(f9895_D,plain,(
  ~sP1('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))) <=> ~$spl3271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3271])])).
fof(f725,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) ) | $spl164),
  inference(cnf_transformation,[],[f725_D])).
fof(f725_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f508,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f318])).
fof(f960,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~sP1('F_0',X12)) ) | $spl382),
  inference(cnf_transformation,[],[f960_D])).
fof(f960_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~sP1('F_0',X12)) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f951,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~sP1('F_0',X12)) ) | ($spl15 | $spl164)),
  inference(resolution,[],[f725,f634])).
fof(f1017,plain,(
  ( ! [X0] : (~sP1('F_0',X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl406),
  inference(cnf_transformation,[],[f1017_D])).
fof(f1017_D,plain,(
  ( ! [X0] : (~sP1('F_0',X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f1006,plain,(
  ( ! [X0] : (~sP1('F_0',X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl84 | $spl382)),
  inference(resolution,[],[f960,f674])).
fof(f10381,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))) | $spl3440),
  inference(cnf_transformation,[],[f10381_D])).
fof(f10381_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))) <=> ~$spl3440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3440])])).
fof(f10320,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))) | ($spl406 | $spl3270)),
  inference(resolution,[],[f9894,f1017])).
fof(f9894,plain,(
  sP1('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))) | $spl3270),
  inference(cnf_transformation,[],[f9894_D])).
fof(f9894_D,plain,(
  sP1('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))) <=> ~$spl3270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3270])])).
fof(f736,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl186),
  inference(cnf_transformation,[],[f736_D])).
fof(f736_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X13,X15,X18] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X15)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X15),s('type.realax.real',X18))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X18)))))),
  file('thm.real.lt_int|split|1','thm.real.REAL_LET_TRANS')).
fof(f650,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl38),
  inference(cnf_transformation,[],[f650_D])).
fof(f650_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f388,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.real.lt_int|split|1','thm.numeral.numeral_distrib|split|20')).
fof(f634,plain,(
  ~p(s(bool,'F_0')) | $spl15),
  inference(cnf_transformation,[],[f634_D])).
fof(f634_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f632,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f371,f365])).
fof(f371,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f34])).
fof(f34,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.real.lt_int|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f670,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl78),
  inference(cnf_transformation,[],[f670_D])).
fof(f670_D,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f408,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X13] : s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.real.lt_int|split|1','thm.real.REAL_NEG_LT|48|')).
fof(f619,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) | $spl1),
  inference(cnf_transformation,[],[f619_D])).
fof(f619_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f623,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | $spl4),
  inference(cnf_transformation,[],[f623_D])).
fof(f623_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f363,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11))))))),
  inference(cnf_transformation,[],[f275])).
fof(f621,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | $spl2),
  inference(cnf_transformation,[],[f621_D])).
fof(f621_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f364,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11))))))),
  inference(cnf_transformation,[],[f275])).
fof(f819,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl15 | $spl38 | $spl78 | $spl104)),
  inference(subsumption_resolution,[],[f818,f634])).
fof(f818,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl2 | $spl4 | $spl38 | $spl78 | $spl104)),
  inference(forward_demodulation,[],[f817,f650])).
fof(f817,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl0 | $spl2 | $spl4 | $spl78 | $spl104)),
  inference(forward_demodulation,[],[f816,f623])).
fof(f816,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ($spl0 | $spl2 | $spl78 | $spl104)),
  inference(forward_demodulation,[],[f815,f690])).
fof(f815,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10)))))) | ($spl0 | $spl2 | $spl78)),
  inference(forward_demodulation,[],[f814,f670])).
fof(f814,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f618,f621])).
fof(f10292,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11))))))) ) | $spl3402),
  inference(cnf_transformation,[],[f10292_D])).
fof(f10292_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11))))))) ) <=> ~$spl3402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3402])])).
fof(f10269,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10))))),s('type.realax.real',X0))))) ) | ($spl1 | $spl186)),
  inference(resolution,[],[f619,f736])).
fof(f674,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl84),
  inference(cnf_transformation,[],[f674_D])).
fof(f674_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f418,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.real.lt_int|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f657,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))))) ) | $spl52),
  inference(cnf_transformation,[],[f657_D])).
fof(f657_D,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f395,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))))),
  file('thm.real.lt_int|split|1','thm.real.REAL_POS')).
fof(f2425,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) ) | $spl834),
  inference(cnf_transformation,[],[f2425_D])).
fof(f2425_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) ) <=> ~$spl834),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl834])])).
fof(f2404,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) ) | ($spl52 | $spl84)),
  inference(resolution,[],[f657,f674])).
fof(f690,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) ) | $spl104),
  inference(cnf_transformation,[],[f690_D])).
fof(f690_D,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f461,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X7,X6] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6)))))),
  file('thm.real.lt_int|split|1','thm.real.REAL_LT')).
fof(f669,plain,(
  ( ! [X0] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl76),
  inference(cnf_transformation,[],[f669_D])).
fof(f669_D,plain,(
  ( ! [X0] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f407,plain,(
  ( ! [X0] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X13] : s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.real.lt_int|split|1','thm.real.REAL_NEG_LE|48|')).
fof(f625,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f625_D])).
fof(f625_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f365,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.lt_int|split|1',aHL_TRUTH)).
fof(f11489,plain,(
  $false | ($spl6 | $spl76 | $spl104 | $spl834 | $spl3402 | $spl3440)),
  inference(subsumption_resolution,[],[f11488,f625])).
fof(f11488,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl76 | $spl104 | $spl834 | $spl3402 | $spl3440)),
  inference(forward_demodulation,[],[f11487,f10381])).
fof(f11487,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))) | ($spl6 | $spl76 | $spl104 | $spl834 | $spl3402)),
  inference(forward_demodulation,[],[f11486,f690])).
fof(f11486,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) | ($spl6 | $spl76 | $spl834 | $spl3402)),
  inference(subsumption_resolution,[],[f11485,f625])).
fof(f11485,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) | ($spl76 | $spl834 | $spl3402)),
  inference(forward_demodulation,[],[f11463,f2425])).
fof(f11463,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK11)))))) | ($spl76 | $spl3402)),
  inference(superposition,[],[f10292,f669])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 404
% Passive clauses: 1875
% Generated clauses: 11144
% Final active clauses: 360
% Final passive clauses: 1390
% Input formulas: 107
% Initial clauses: 251
% 
% Duplicate literals: 76
% Trivial inequalities: 18
% Fw subsumption resolutions: 598
% Fw demodulations: 1358
% Bw demodulations: 90
% Condensations: 26
% Global subsumptions: 16
% 
% Simple tautologies: 192
% Equational tautologies: 8
% Forward subsumptions: 4771
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 1826
% Factoring: 63
% Forward superposition: 2554
% Backward superposition: 2290
% Self superposition: 28
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 18
% Split components: 37
% SAT solver clauses: 2397
% SAT solver unit clauses: 293
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 35
% SAT solver learnt literals: 14
% 
% Sat splits: 18
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 2330
% TWLsolver calls for satisfiability: 1664
% 
% Memory used [KB]: 11129
% Time elapsed: 0.309 s
% ------------------------------
% ------------------------------
% Success in time 0.885 s
