Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (84505)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 120
% Passive clauses: 1626
% Generated clauses: 19978
% Final active clauses: 116
% Final passive clauses: 1429
% Input formulas: 23
% Initial clauses: 24
% 
% Duplicate literals: 1412
% Trivial inequalities: 52
% Fw subsumption resolutions: 1897
% Fw demodulations: 3064
% Bw demodulations: 2
% 
% Simple tautologies: 196
% Equational tautologies: 18
% Forward subsumptions: 10777
% 
% Binary resolution: 11
% Forward superposition: 7838
% Backward superposition: 5358
% Self superposition: 167
% Equality factoring: 34
% Equality resolution: 5
% 
% Split clauses: 570
% Split components: 1145
% SAT solver clauses: 1010
% SAT solver unit clauses: 1
% SAT solver binary clauses: 20
% 
% Sat splits: 570
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1010
% TWLsolver calls for satisfiability: 83
% 
% Memory used [KB]: 14072
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f6931,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f6910,f71,f96,f6548,f6670,f80,f105,f2766,f2767,f76,f103,f89,f114,f90,f115])).
fof(f115,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))) ) | $spl34),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f90,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X18,X19,X20,X21] : s(X6,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X6),X18),s('type.quote.varmap'(X6),X19),s(X6,X20),s('type.list.list'('type.quote.index'),X21))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun(X6,fun('type.list.list'('type.quote.index'),X6)),happ(s(fun('type.quote.varmap'(X6),fun(X6,fun('type.list.list'('type.quote.index'),X6))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X6),X18))),s('type.quote.varmap'(X6),X19))),s(X6,X20))),s('type.list.list'('type.quote.index'),X21)))),
  file('thm.ringNorm.interp_m_def|split|0','ARITY_1_4_const.ringNorm.r_interp_m')).
fof(f114,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))) ) | $spl32),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f89,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X18,X19,X20,X21] : s(X6,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X6),X18),s('type.quote.varmap'(X6),X19),s(X6,X20),s('type.list.list'('type.quote.index'),X21))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun(X6,fun('type.list.list'('type.quote.index'),X6)),happ(s(fun('type.quote.varmap'(X6),fun(X6,fun('type.list.list'('type.quote.index'),X6))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X6),X18))),s('type.quote.varmap'(X6),X19))),s(X6,X20))),s('type.list.list'('type.quote.index'),X21)))),
  file('thm.ringNorm.interp_m_def|split|0','ARITY_1_4_const.canonical.interp_m')).
fof(f103,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) | $spl10),
  inference(cnf_transformation,[],[f103_D])).
fof(f103_D,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f76,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun(X6,fun('type.list.list'('type.quote.index'),X6))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun(X6,fun('type.list.list'('type.quote.index'),X6))),'const.canonical.interp_m_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_m_def|split|0','thm.ringNorm.r_interp_m_def')).
fof(f2767,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4)))) ) | $spl414),
  inference(cnf_transformation,[],[f2767_D])).
fof(f2767_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4)))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f2766,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4)))) ) | ($spl10 | $spl32 | $spl34)),
  inference(forward_demodulation,[],[f2634,f115])).
fof(f2634,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun(X0,fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun(X0,fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_interp_m_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s(X0,X3))),s('type.list.list'('type.quote.index'),X4)))) ) | ($spl10 | $spl32)),
  inference(superposition,[],[f114,f103])).
fof(f105,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) ) | $spl14),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f80,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X12,X8,X9] : s(X6,X9) = s(X6,'const.canonical.interp_m_4'(s('type.semi_ring.semi_ring'(X6),X12),s('type.quote.varmap'(X6),X8),s(X6,X9),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  file('thm.ringNorm.interp_m_def|split|0','thm.canonical.interp_m_def|split|0')).
fof(f6670,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X8) = s(X5,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X5),X6),s('type.quote.varmap'(X5),X7),s(X5,X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) ) | $spl984),
  inference(cnf_transformation,[],[f6670_D])).
fof(f6670_D,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X8) = s(X5,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X5),X6),s('type.quote.varmap'(X5),X7),s(X5,X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) ) <=> ~$spl984),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl984])])).
fof(f6548,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X8) = s(X5,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X5),X6),s('type.quote.varmap'(X5),X7),s(X5,X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) ) | ($spl14 | $spl414)),
  inference(superposition,[],[f2767,f105])).
fof(f96,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s(sK0,sK3),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))) | $spl1),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s(sK0,sK3),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f71,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s(sK0,sK3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s(sK0,sK3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2,X3] : s(X0,X3) != s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ~! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s(X0,X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,X9) = s(X6,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s(X6,X9),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,X9) = s(X6,'const.ringNorm.r_interp_m_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s(X6,X9),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  file('thm.ringNorm.interp_m_def|split|0','thm.ringNorm.interp_m_def|split|0')).
fof(f6910,plain,(
  $false | ($spl1 | $spl984)),
  inference(trivial_inequality_removal,[],[f6909])).
fof(f6909,plain,(
  s(sK0,sK3) != s(sK0,sK3) | ($spl1 | $spl984)),
  inference(superposition,[],[f96,f6670])).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 72
% Passive clauses: 504
% Generated clauses: 6782
% Final active clauses: 67
% Final passive clauses: 422
% Input formulas: 23
% Initial clauses: 24
% 
% Duplicate literals: 357
% Trivial inequalities: 17
% Fw subsumption resolutions: 218
% Fw demodulations: 506
% Bw demodulations: 11
% Condensations: 448
% 
% Simple tautologies: 78
% Equational tautologies: 16
% Forward subsumptions: 3884
% 
% Binary resolution: 8
% Forward superposition: 2699
% Backward superposition: 1838
% Self superposition: 112
% Equality factoring: 30
% Equality resolution: 2
% 
% Split clauses: 14
% Split components: 33
% SAT solver clauses: 741
% SAT solver unit clauses: 42
% SAT solver binary clauses: 14
% SAT solver learnt clauses: 1
% 
% Sat splits: 14
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 740
% TWLsolver calls for satisfiability: 94
% 
% Memory used [KB]: 8699
% Time elapsed: 0.281 s
% ------------------------------
% ------------------------------
% Success in time 0.897 s
