Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.option
% (51643)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 836
% Passive clauses: 6010
% Generated clauses: 24660
% Final active clauses: 786
% Final passive clauses: 4585
% Input formulas: 48
% Initial clauses: 149
% 
% Duplicate literals: 464
% Trivial inequalities: 231
% Fw subsumption resolutions: 2835
% Fw demodulations: 1898
% Bw demodulations: 47
% 
% Simple tautologies: 232
% Equational tautologies: 65
% Forward subsumptions: 12110
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 6978
% Factoring: 1614
% Forward superposition: 4757
% Backward superposition: 5294
% Self superposition: 46
% Equality factoring: 32
% Equality resolution: 40
% 
% Split clauses: 777
% Split components: 1561
% SAT solver clauses: 1216
% SAT solver unit clauses: 27
% SAT solver binary clauses: 276
% 
% Sat splits: 777
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1209
% TWLsolver calls for satisfiability: 429
% 
% Memory used [KB]: 10234
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.option
% (51778)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 441
% Passive clauses: 3009
% Generated clauses: 18705
% Final active clauses: 208
% Final passive clauses: 1377
% Input formulas: 48
% Initial clauses: 149
% 
% Duplicate literals: 269
% Trivial inequalities: 242
% Fw subsumption resolutions: 1632
% Fw demodulations: 792
% Bw demodulations: 1043
% Condensations: 553
% Global subsumptions: 18
% 
% Simple tautologies: 210
% Equational tautologies: 25
% Forward subsumptions: 8863
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 195
% 
% Binary resolution: 3274
% Factoring: 342
% Forward superposition: 2356
% Backward superposition: 1855
% Self superposition: 37
% Equality factoring: 28
% Equality resolution: 33
% 
% Split clauses: 298
% Split components: 611
% SAT solver clauses: 2950
% SAT solver unit clauses: 164
% SAT solver binary clauses: 78
% SAT solver learnt clauses: 38
% SAT solver learnt literals: 15
% 
% Sat splits: 298
% Sat splitting refutations: 53
% 
% TWLsolver clauses: 2876
% TWLsolver calls for satisfiability: 1122
% 
% Memory used [KB]: 10746
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.option
% (51815)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2294
% Passive clauses: 10259
% Generated clauses: 175792
% Final active clauses: 327
% Final passive clauses: 377
% Input formulas: 48
% Initial clauses: 149
% 
% Duplicate literals: 4404
% Trivial inequalities: 538
% Fw subsumption resolutions: 3807
% Fw demodulations: 18910
% Bw demodulations: 6599
% Condensations: 7716
% Global subsumptions: 2116
% 
% Simple tautologies: 251
% Equational tautologies: 458
% Forward subsumptions: 122320
% Fw demodulations to eq. taut.: 1132
% Bw demodulations to eq. taut.: 110
% 
% Unit resulting resolution: 13085
% Factoring: 1
% Forward superposition: 53608
% Backward superposition: 57065
% Self superposition: 163
% Equality factoring: 78
% Equality resolution: 54
% 
% Split clauses: 475
% Split components: 959
% SAT solver clauses: 15568
% SAT solver unit clauses: 644
% SAT solver binary clauses: 2548
% 
% Sat splits: 475
% Sat splitting refutations: 2240
% 
% Memory used [KB]: 143665
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.option
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.option
% SZS output start Proof for thm.option
fof(f22314,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f215,f367,f214,f369,f365,f213,f366,f368,f364,f370,f371,f372,f373,f374,f376,f377,f422,f460,f461,f462,f463,f464,f466,f467,f509,f537,f538,f539,f540,f541,f543,f544,f589,f647,f648,f649,f650,f651,f653,f654,f699,f804,f801,f846,f847,f848,f857,f858,f859,f860,f861,f863,f447,f864,f906,f532,f965,f966,f967,f968,f969,f971,f972,f1014,f1044,f1062,f1063,f1104,f1102,f1150,f1169,f1170,f1171,f1172,f1173,f1175,f1267,f1176,f1218,f1266,f1222,f1223,f1224,f1273,f1226,f1227,f1228,f1279,f1263,f1262,f1272,f1289,f1290,f1291,f1292,f1293,f1295,f1387,f1296,f1338,f1386,f1342,f1343,f1344,f1392,f1346,f1347,f1348,f1398,f1383,f1382,f1391,f1669,f1671,f1675,f1686,f1687,f1691,f1823,f1834,f1837,f1905,f1914,f1963,f1972,f2006,f2016,f2035,f2045,f2050,f2060,f2099,f2100,f2101,f2102,f2103,f2105,f2208,f2106,f2148,f2207,f2152,f2153,f2154,f2212,f2156,f2157,f2158,f2218,f2195,f2194,f2223,f2238,f2259,f2261,f2264,f2267,f2269,f2273,f2327,f2328,f2329,f2375,f2331,f2332,f2333,f2381,f2369,f2368,f2443,f2444,f2445,f2520,f2519,f2523,f2524,f2612,f2605,f2604,f2603,f2668,f2669,f2670,f2671,f2672,f2673,f2753,f2746,f2745,f2744,f2809,f2810,f2811,f2812,f2813,f2814,f2870,f2871,f2872,f2932,f2933,f2934,f2943,f2944,f2945,f2946,f3055,f2947,f3056,f2949,f2950,f2994,f2995,f3000,f3001,f3002,f3003,f3004,f3005,f3006,f3015,f3017,f3073,f3051,f3050,f3135,f3212,f3222,f3354,f3355,f3356,f3357,f3358,f3360,f3465,f3361,f3403,f3464,f3408,f3409,f3410,f3469,f3412,f3413,f3414,f3476,f3452,f3451,f3496,f3497,f3498,f3499,f3609,f3500,f3610,f3502,f3503,f3547,f3548,f3554,f3555,f3556,f3557,f3558,f3559,f3560,f3569,f3571,f3624,f3605,f3604,f3647,f3657,f3674,f3698,f3730,f3740,f3750,f3761,f3796,f3807,f4319,f4321,f4362,f4363,f4320,f4404,f4406,f4447,f4448,f4405,f4716,f4717,f4718,f4719,f4720,f4722,f4830,f4723,f4765,f4829,f4864,f4865,f4866,f4867,f4868,f4870,f4977,f4871,f4913,f4976,f5225,f5226,f5230,f5241,f5242,f5246,f1670,f5662,f5932,f6210,f6475,f5663,f6527,f6528,f6529,f6530,f6642,f6531,f6643,f6533,f6534,f6578,f6644,f6579,f6641,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6602,f6604,f6653,f6657,f6638,f6637,f6650,f6662,f6664,f6684,f6685,f6686,f6687,f6688,f6690,f6787,f6691,f6733,f6935,f6937,f6954,f6955,f6956,f6957,f7070,f6958,f7071,f6960,f6961,f7005,f7072,f7006,f7069,f7030,f7032,f7077,f7110,f7111,f7126,f7137,f7151,f7162,f7234,f7246,f7269,f7270,f7271,f7272,f730,f7273,f7275,f7276,f7321,f729,f7377,f7391,f7401,f7413,f7422,f7430,f7532,f7541,f7547,f7557,f7563,f7573,f7579,f7589,f7599,f7609,f7619,f7629,f7639,f7649,f7736,f7774,f7785,f7795,f7796,f7797,f7798,f7799,f7911,f7800,f7912,f7802,f7803,f7847,f7913,f7848,f7910,f7857,f7858,f7859,f7860,f7861,f7862,f7863,f7872,f7874,f7922,f7926,f7907,f7906,f7919,f7930,f7932,f7931,f8016,f8054,f8055,f8056,f8057,f8058,f8059,f8061,f8160,f8062,f8104,f8159,f8178,f8191,f8244,f8257,f8293,f8301,f8350,f8352,f8349,f8354,f8385,f8387,f8414,f8415,f8413,f8417,f8416,f8425,f8434,f8465,f8467,f8468,f8469,f8470,f8471,f8584,f8472,f8585,f8474,f8475,f8519,f8586,f8520,f8583,f8545,f8547,f8591,f8623,f8624,f8651,f8652,f8905,f8908,f8909,f8918,f8929,f8930,f2374,f9442,f9443,f9444,f9445,f9559,f9446,f9560,f9448,f9449,f9493,f9561,f9494,f9519,f9521,f9595,f9643,f9667,f9668,f9669,f9670,f9793,f9671,f9794,f9673,f9674,f9718,f9795,f9719,f9788,f9786,f9828,f9874,f9876,f9879,f9881,f9905,f9929,f9974,f9900,f9979,f9902,f9981,f9901,f9982,f9903,f9986,f10001,f10047,f10052,f10054,f10055,f10084,f10086,f10089,f10091,f10175,f10176,f10177,f10178,f10522,f10524,f10548,f10550,f10580,f10582,f10607,f10609,f10637,f10639,f10663,f10665,f10693,f10695,f10720,f10569,f10722,f10740,f10741,f10742,f10743,f10866,f10744,f10867,f10746,f10747,f10791,f10868,f10792,f10865,f10861,f10859,f10878,f10879,f10880,f10881,f10995,f10882,f10996,f10884,f10885,f10929,f10997,f448,f10930,f10994,f10955,f10957,f11078,f11085,f7238,f7236,f11260,f11265,f11267,f11268,f11297,f11299,f11302,f11304,f11389,f11435,f11436,f11437,f11490,f11491,f11492,f11493,f11524,f11525,f11496,f11529,f11530,f11531,f11532,f11563,f11564,f11535,f11569,f11570,f11571,f11572,f11603,f11604,f11575,f11608,f11609,f11610,f11611,f11642,f11643,f11614,f11648,f11649,f11650,f11651,f11681,f11682,f11654,f11691,f11522,f11692,f11693,f11523,f11694,f11722,f11723,f11698,f11526,f11725,f11561,f11726,f11727,f11562,f11728,f11756,f11757,f11732,f11565,f11761,f11601,f11762,f11763,f11602,f11764,f11792,f11793,f11768,f11605,f11796,f11640,f11797,f11798,f11641,f11799,f11827,f11828,f11803,f11644,f11832,f11679,f11833,f11834,f11680,f11835,f11862,f11863,f11678,f11839,f11683,f11871,f11872,f11873,f11874,f11905,f11906,f11877,f11909,f11910,f11911,f11912,f11943,f11944,f11915,f11949,f11950,f11951,f11952,f11983,f11984,f11955,f11988,f11989,f11990,f11991,f12022,f12023,f11994,f12028,f12029,f12030,f12031,f12061,f12062,f12034,f12071,f11903,f12072,f12073,f11904,f12074,f12102,f12103,f12078,f11907,f12106,f11941,f12107,f12108,f11942,f12109,f12137,f12138,f12113,f11945,f12142,f11981,f12143,f12144,f11982,f12145,f12173,f12174,f12149,f11985,f12177,f12020,f12178,f12179,f12021,f12180,f12208,f12209,f12184,f12024,f12213,f12059,f12214,f12215,f12060,f12216,f12243,f12244,f12058,f12220,f12063,f12401,f12403,f12488,f12490,f12509,f12536,f12538,f12582,f12584,f12602,f12660,f12662,f12673,f12680,f12757,f12759,f12778,f12770,f12779,f12815,f12817,f12828,f12836,f12914,f12916,f12934,f12927,f12936,f13029,f13031,f13073,f13075,f13092,f13119,f13121,f13164,f13166,f13182,f13263,f13265,f13307,f13309,f13326,f13353,f13355,f13398,f13400,f13416,f3611,f13557,f13556,f13564,f13573,f13624,f13626,f13628,f13630,f13631,f13676,f13740,f13686,f13687,f13688,f13689,f13690,f13691,f13692,f13749,f13750,f8353,f13754,f13737,f13736,f13746,f14018,f14017,f14023,f13563,f14943,f14944,f14945,f14946,f14957,f14958,f14959,f14960,f14961,f14962,f14963,f14964,f14965,f14966,f15012,f15018,f15019,f15020,f15082,f15021,f15083,f15022,f15084,f15023,f15085,f15024,f15086,f15025,f15087,f15026,f15088,f15031,f15089,f15032,f15090,f15033,f15091,f15034,f15092,f15035,f15093,f15043,f8433,f15095,f15115,f15118,f15121,f15124,f15126,f15394,f13703,f13701,f15395,f8155,f8156,f15405,f8119,f8118,f8117,f15408,f8115,f8114,f8113,f3618,f6782,f6783,f15410,f6746,f6745,f6744,f15413,f6742,f6741,f6740,f3064,f1270,f15446,f15448,f15450,f15452,f15453,f15498,f15527,f15554,f15966,f15967,f15969,f15971,f15973,f15974,f16019,f445,f16052,f16060,f16073,f1100,f1098,f16112,f16130,f16137,f16145,f16176,f16177,f16187,f16188,f16197,f16198,f6663,f16195,f16200,f16598,f16600,f16602,f16604,f16605,f16650,f16707,f16718,f16720,f16722,f16724,f16725,f16770,f16827,f17054,f17100,f17101,f17102,f17245,f17247,f17249,f17251,f17252,f17297,f3054,f17315,f17317,f17462,f17464,f17465,f17466,f17469,f17511,f6786,f17780,f17782,f17784,f17786,f17787,f17832,f17890,f17851,f17853,f17895,f18012,f18014,f18016,f18018,f18019,f18064,f18123,f18083,f18085,f18128,f18227,f18229,f18231,f18233,f18234,f18279,f3608,f18298,f18300,f18923,f18925,f18927,f18929,f18930,f18975,f19034,f18994,f18996,f3468,f19044,f19046,f19048,f19050,f19051,f19096,f19154,f19115,f19117,f2211,f19334,f19436,f19570,f19672,f19803,f19903,f20035,f20135,f5664,f20179,f20180,f6211,f20345,f8386,f20398,f8466,f20460,f6936,f20515,f8351,f20720,f20722,f20724,f20726,f20727,f20772,f9558,f20791,f20793,f9566,f20839,f20841,f20843,f446,f20845,f20846,f449,f20891,f9792,f20954,f20952,f3067,f21054,f21058,f21062,f21063,f21099,f11521,f21146,f21150,f21154,f21155,f21191,f11560,f21211,f21240,f21244,f21248,f21249,f21288,f11600,f21336,f21340,f21344,f21345,f10570,f21381,f11639,f21400,f21494,f21498,f21502,f21503,f21539,f11902,f21586,f21590,f21594,f21595,f21631,f11940,f21651,f12510,f21680,f21684,f21688,f21689,f21772,f11980,f21820,f21824,f21828,f21829,f10683,f21865,f12019,f3069,f21884,f3057,f12603,f22087,f22120,f22129,f22131,f22133,f22150,f21875,f21870,f22158,f22160,f21782,f21777,f22163,f22164,f21641,f21636,f22165,f22167,f21549,f21544,f22157,f22170,f22171,f21391,f21386,f22179,f22180,f21298,f21293,f22183,f22184,f22159,f21201,f21196,f22185,f22186,f21109,f21104,f22178,f451,f22189,f22190,f2239,f22166,f22275,f22146,f22277,f452,f22148,f22279,f1846,f1833,f22313])).
fof(f22313,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f7245,f4985])).
fof(f4985,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(backward_demodulation,[],[f1008,f1006])).
fof(f1006,plain,(
  s(bool,'F_0') != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | $spl2),
  inference(resolution,[],[f366,f320])).
fof(f320,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f1008,plain,(
  s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | $spl2),
  inference(resolution,[],[f366,f325])).
fof(f325,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f7245,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f7178,f1008])).
fof(f7178,plain,(
  s(bool,'F_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | $spl1),
  inference(superposition,[],[f216,f1105])).
fof(f1105,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl1),
  inference(subsumption_resolution,[],[f1070,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X14,X11] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,X14),s(X6,X11),s(X6,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.COND_ID')).
fof(f1070,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),'const.option.NONE_0'),s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl1),
  inference(superposition,[],[f365,f357])).
fof(f357,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK9(X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK9(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X18] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X18) | ? [X12] : s('type.option.option'(X6),X18) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.option.option_nchotomy')).
fof(f216,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X12] : s(bool,'F_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12)))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.option.IS_NONE_DEF|split|0')).
fof(f1833,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl52),
  inference(cnf_transformation,[],[f1833_D])).
fof(f1833_D,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1846,plain,(
  ( ! [X0] : (sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0)) ) | $spl56),
  inference(cnf_transformation,[],[f1846_D])).
fof(f1846_D,plain,(
  ( ! [X0] : (sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f22279,plain,(
  ( ! [X0] : (sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f996,f357])).
fof(f996,plain,(
  ( ! [X61] : (sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X61)) ) | $spl2),
  inference(resolution,[],[f366,f273])).
fof(f273,plain,(
  ( ! [X0,X1] : (sP5(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ((~sP5(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP5(X0,X1)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((~sP5(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP5(X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : (sP5(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f22148,plain,(
  sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | $spl430),
  inference(cnf_transformation,[],[f22148_D])).
fof(f22148_D,plain,(
  sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f452,plain,(
  ( ! [X18] : (p(s(bool,X18))) ) | $spl18),
  inference(cnf_transformation,[],[f452_D])).
fof(f452_D,plain,(
  ( ! [X18] : (p(s(bool,X18))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f22277,plain,(
  ( ! [X5] : (p(s(bool,X5)) | sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f996,f16829])).
fof(f16829,plain,(
  ( ! [X92,X93] : (~sP5(X92,X93) | p(s(bool,X93)) | sP4(sK8,X92)) ) | $spl5),
  inference(subsumption_resolution,[],[f16749,f309])).
fof(f309,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.AND_CLAUSES|split|3')).
fof(f16749,plain,(
  ( ! [X92,X93] : (sP4(sK8,X92) | p(s(bool,'F_0')) | p(s(bool,X93)) | ~sP5(X92,X93)) ) | $spl5),
  inference(resolution,[],[f15471,f271])).
fof(f271,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f15471,plain,(
  ( ! [X55] : (~p(s(bool,X55)) | sP4(sK8,X55)) ) | $spl5),
  inference(resolution,[],[f368,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (sP4(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : ((~sP4(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1)))),
  inference(rectify,[],[f166])).
fof(f166,plain,(
  ! [X1,X0] : ((~sP4(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X1,X0] : ((~sP4(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP4(X1,X0)))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X1,X0] : (sP4(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f22146,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),sK8) | $spl428),
  inference(cnf_transformation,[],[f22146_D])).
fof(f22146_D,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),sK8) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f22275,plain,(
  ( ! [X3] : (p(s(bool,X3)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),sK8)) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f996,f16709])).
fof(f16709,plain,(
  ( ! [X92,X93] : (~sP5(X92,X93) | p(s(bool,X93)) | sP4(X92,sK8)) ) | $spl5),
  inference(subsumption_resolution,[],[f16629,f309])).
fof(f16629,plain,(
  ( ! [X92,X93] : (sP4(X92,sK8) | p(s(bool,'F_0')) | p(s(bool,X93)) | ~sP5(X92,X93)) ) | $spl5),
  inference(resolution,[],[f15470,f271])).
fof(f15470,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | sP4(X54,sK8)) ) | $spl5),
  inference(resolution,[],[f368,f257])).
fof(f257,plain,(
  ( ! [X0,X1] : (sP4(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f167])).
fof(f22166,plain,(
  ( ! [X10,X11,X9] : (sP2(X9,X10,X11)) ) | $spl438),
  inference(cnf_transformation,[],[f22166_D])).
fof(f22166_D,plain,(
  ( ! [X10,X11,X9] : (sP2(X9,X10,X11)) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f2239,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | sP4(X2,'F_0')) ) | $spl76),
  inference(cnf_transformation,[],[f2239_D])).
fof(f2239_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | sP4(X2,'F_0')) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f22190,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP2(X5,X6,X7) | ~p(s(bool,X4))) ) | $spl288),
  inference(subsumption_resolution,[],[f21074,f235])).
fof(f235,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f155])).
fof(f155,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f21074,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP2(X5,X6,X7) | ~p(s(bool,X6)) | ~p(s(bool,X4))) ) | $spl288),
  inference(resolution,[],[f11523,f256])).
fof(f256,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP4(X0,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f22189,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP2(X9,X10,X11) | ~p(s(bool,X8))) ) | $spl288),
  inference(subsumption_resolution,[],[f21075,f235])).
fof(f21075,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP2(X9,X10,X11) | ~p(s(bool,X8)) | ~p(s(bool,X10))) ) | $spl288),
  inference(resolution,[],[f11523,f263])).
fof(f263,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP4(X1,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP4(X1,X0)) & (~sP4(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP4(X1,X0))),
  inference(definition_folding,[],[f60,f142])).
fof(f60,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X21,X22] : ((p(s(bool,X22)) <=> ~p(s(bool,X21))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.sat.sat_29_prove_ltb')).
fof(f451,plain,(
  ( ! [X20] : (~p(s(bool,X20))) ) | $spl16),
  inference(cnf_transformation,[],[f451_D])).
fof(f451_D,plain,(
  ( ! [X20] : (~p(s(bool,X20))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f22178,plain,(
  ( ! [X19,X20] : (sP4(X19,X20) | ~p(s(bool,X19)) | sP4(X20,'F_0')) ) | $spl440),
  inference(cnf_transformation,[],[f22178_D])).
fof(f22178_D,plain,(
  ( ! [X19,X20] : (sP4(X19,X20) | ~p(s(bool,X19)) | sP4(X20,'F_0')) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f21104,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,'F_0') | ~p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X3))) ) | $spl288),
  inference(resolution,[],[f11523,f233])).
fof(f233,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f156])).
fof(f21109,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4(X20,'F_0') | ~p(s(bool,X19)) | ~p(s(bool,X21))) ) | ($spl5 | $spl288)),
  inference(resolution,[],[f11523,f15457])).
fof(f15457,plain,(
  ( ! [X30,X31] : (~sP2(X30,X31,sK8) | ~p(s(bool,X31)) | ~p(s(bool,X30))) ) | $spl5),
  inference(resolution,[],[f368,f233])).
fof(f22186,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP2(X5,X6,X7) | ~p(s(bool,X4))) ) | $spl296),
  inference(subsumption_resolution,[],[f21166,f236])).
fof(f236,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f156])).
fof(f21166,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP2(X5,X6,X7) | ~p(s(bool,X5)) | ~p(s(bool,X4))) ) | $spl296),
  inference(resolution,[],[f11562,f256])).
fof(f22185,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP2(X9,X10,X11) | ~p(s(bool,X8))) ) | $spl296),
  inference(subsumption_resolution,[],[f21167,f236])).
fof(f21167,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP2(X9,X10,X11) | ~p(s(bool,X8)) | ~p(s(bool,X9))) ) | $spl296),
  inference(resolution,[],[f11562,f263])).
fof(f21196,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl296),
  inference(resolution,[],[f11562,f233])).
fof(f21201,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4(X20,'F_0') | ~p(s(bool,X21)) | ~p(s(bool,X19))) ) | ($spl5 | $spl296)),
  inference(resolution,[],[f11562,f15457])).
fof(f22159,plain,(
  ( ! [X10,X11,X9] : (sP1(X9,X10,X11)) ) | $spl436),
  inference(cnf_transformation,[],[f22159_D])).
fof(f22159_D,plain,(
  ( ! [X10,X11,X9] : (sP1(X9,X10,X11)) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f22184,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP1(X5,X6,X7) | ~p(s(bool,X4))) ) | $spl304),
  inference(subsumption_resolution,[],[f21263,f239])).
fof(f239,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f158])).
fof(f158,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f21263,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP1(X5,X6,X7) | ~p(s(bool,X6)) | ~p(s(bool,X4))) ) | $spl304),
  inference(resolution,[],[f11602,f256])).
fof(f22183,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP1(X9,X10,X11) | ~p(s(bool,X8))) ) | $spl304),
  inference(subsumption_resolution,[],[f21264,f239])).
fof(f21264,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP1(X9,X10,X11) | ~p(s(bool,X8)) | ~p(s(bool,X10))) ) | $spl304),
  inference(resolution,[],[f11602,f263])).
fof(f21293,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,'F_0') | ~p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X3))) ) | $spl304),
  inference(resolution,[],[f11602,f237])).
fof(f237,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f159])).
fof(f21298,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4(X20,'F_0') | ~p(s(bool,X19)) | ~p(s(bool,X21))) ) | ($spl5 | $spl304)),
  inference(resolution,[],[f11602,f15460])).
fof(f15460,plain,(
  ( ! [X37,X36] : (~sP1(X36,X37,sK8) | ~p(s(bool,X37)) | ~p(s(bool,X36))) ) | $spl5),
  inference(resolution,[],[f368,f237])).
fof(f22180,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP1(X5,X6,X7) | ~p(s(bool,X4))) ) | $spl312),
  inference(subsumption_resolution,[],[f21356,f240])).
fof(f240,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f159])).
fof(f21356,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP1(X5,X6,X7) | ~p(s(bool,X5)) | ~p(s(bool,X4))) ) | $spl312),
  inference(resolution,[],[f11641,f256])).
fof(f22179,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP1(X9,X10,X11) | ~p(s(bool,X8))) ) | $spl312),
  inference(subsumption_resolution,[],[f21357,f240])).
fof(f21357,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP1(X9,X10,X11) | ~p(s(bool,X8)) | ~p(s(bool,X9))) ) | $spl312),
  inference(resolution,[],[f11641,f263])).
fof(f21386,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl312),
  inference(resolution,[],[f11641,f237])).
fof(f21391,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4(X20,'F_0') | ~p(s(bool,X21)) | ~p(s(bool,X19))) ) | ($spl5 | $spl312)),
  inference(resolution,[],[f11641,f15460])).
fof(f22171,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP2(X5,X6,X7) | ~p(s(bool,X4))) ) | $spl328),
  inference(subsumption_resolution,[],[f21514,f235])).
fof(f21514,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP2(X5,X6,X7) | ~p(s(bool,X4)) | ~p(s(bool,X6))) ) | $spl328),
  inference(resolution,[],[f11904,f256])).
fof(f22170,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP2(X9,X10,X11) | ~p(s(bool,X8))) ) | $spl328),
  inference(subsumption_resolution,[],[f21515,f235])).
fof(f21515,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP2(X9,X10,X11) | ~p(s(bool,X10)) | ~p(s(bool,X8))) ) | $spl328),
  inference(resolution,[],[f11904,f263])).
fof(f22157,plain,(
  ( ! [X19,X20] : (sP4(X19,X20) | ~p(s(bool,X20)) | sP4(X19,'F_0')) ) | $spl434),
  inference(cnf_transformation,[],[f22157_D])).
fof(f22157_D,plain,(
  ( ! [X19,X20] : (sP4(X19,X20) | ~p(s(bool,X20)) | sP4(X19,'F_0')) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f21544,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,'F_0') | ~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X3))) ) | $spl328),
  inference(resolution,[],[f11904,f233])).
fof(f21549,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4(X19,'F_0') | ~p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl5 | $spl328)),
  inference(resolution,[],[f11904,f15457])).
fof(f22167,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP2(X5,X6,X7) | ~p(s(bool,X4))) ) | $spl336),
  inference(subsumption_resolution,[],[f21606,f236])).
fof(f21606,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP2(X5,X6,X7) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl336),
  inference(resolution,[],[f11942,f256])).
fof(f22165,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP2(X9,X10,X11) | ~p(s(bool,X8))) ) | $spl336),
  inference(subsumption_resolution,[],[f21607,f236])).
fof(f21607,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP2(X9,X10,X11) | ~p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl336),
  inference(resolution,[],[f11942,f263])).
fof(f21636,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,'F_0') | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl336),
  inference(resolution,[],[f11942,f233])).
fof(f21641,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4(X19,'F_0') | ~p(s(bool,X21)) | ~p(s(bool,X20))) ) | ($spl5 | $spl336)),
  inference(resolution,[],[f11942,f15457])).
fof(f22164,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP1(X5,X6,X7) | ~p(s(bool,X4))) ) | $spl344),
  inference(subsumption_resolution,[],[f21747,f239])).
fof(f21747,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP1(X5,X6,X7) | ~p(s(bool,X4)) | ~p(s(bool,X6))) ) | $spl344),
  inference(resolution,[],[f11982,f256])).
fof(f22163,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP1(X9,X10,X11) | ~p(s(bool,X8))) ) | $spl344),
  inference(subsumption_resolution,[],[f21748,f239])).
fof(f21748,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP1(X9,X10,X11) | ~p(s(bool,X10)) | ~p(s(bool,X8))) ) | $spl344),
  inference(resolution,[],[f11982,f263])).
fof(f21777,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,'F_0') | ~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X3))) ) | $spl344),
  inference(resolution,[],[f11982,f237])).
fof(f21782,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4(X19,'F_0') | ~p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl5 | $spl344)),
  inference(resolution,[],[f11982,f15460])).
fof(f22160,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP1(X5,X6,X7) | ~p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f21840,f240])).
fof(f21840,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | sP1(X5,X6,X7) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl352),
  inference(resolution,[],[f12021,f256])).
fof(f22158,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP1(X9,X10,X11) | ~p(s(bool,X8))) ) | $spl352),
  inference(subsumption_resolution,[],[f21841,f240])).
fof(f21841,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP1(X9,X10,X11) | ~p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl352),
  inference(resolution,[],[f12021,f263])).
fof(f21870,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,'F_0') | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl352),
  inference(resolution,[],[f12021,f237])).
fof(f21875,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4(X19,'F_0') | ~p(s(bool,X21)) | ~p(s(bool,X20))) ) | ($spl5 | $spl352)),
  inference(resolution,[],[f12021,f15460])).
fof(f22150,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'F_0') | $spl432),
  inference(cnf_transformation,[],[f22150_D])).
fof(f22150_D,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'F_0') <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f22133,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'F_0')) ) | ($spl2 | $spl76)),
  inference(resolution,[],[f995,f17354])).
fof(f17354,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4(X90,'F_0')) ) | $spl76),
  inference(subsumption_resolution,[],[f17275,f309])).
fof(f17275,plain,(
  ( ! [X90,X91] : (sP4(X90,'F_0') | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl76),
  inference(resolution,[],[f2239,f271])).
fof(f995,plain,(
  ( ! [X60] : (sP5(X60,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f272])).
fof(f272,plain,(
  ( ! [X0,X1] : (sP5(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f172])).
fof(f22131,plain,(
  ( ! [X4] : (p(s(bool,X4)) | sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f995,f16828])).
fof(f16828,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4(sK8,X90)) ) | $spl5),
  inference(subsumption_resolution,[],[f16748,f309])).
fof(f16748,plain,(
  ( ! [X90,X91] : (sP4(sK8,X90) | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl5),
  inference(resolution,[],[f15471,f271])).
fof(f22129,plain,(
  ( ! [X2] : (p(s(bool,X2)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),sK8)) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f995,f16708])).
fof(f16708,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4(X90,sK8)) ) | $spl5),
  inference(subsumption_resolution,[],[f16628,f309])).
fof(f16628,plain,(
  ( ! [X90,X91] : (sP4(X90,sK8) | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl5),
  inference(resolution,[],[f15470,f271])).
fof(f22120,plain,(
  $false | ($spl5 | $spl32)),
  inference(subsumption_resolution,[],[f1102,f15490])).
fof(f15490,plain,(
  s(bool,'T_0') != s(bool,sK8) | $spl5),
  inference(resolution,[],[f368,f324])).
fof(f324,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f22087,plain,(
  $false | ($spl3 | $spl5 | $spl28 | $spl35)),
  inference(subsumption_resolution,[],[f22025,f1104])).
fof(f22025,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | ($spl3 | $spl5 | $spl28)),
  inference(superposition,[],[f21947,f7364])).
fof(f7364,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl3),
  inference(subsumption_resolution,[],[f7363,f361])).
fof(f361,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f284])).
fof(f284,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X12] : (s(X6,X12) = s(X6,X12) <=> p(s(bool,'T_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.REFL_CLAUSE')).
fof(f7363,plain,(
  ~p(s(bool,'T_0')) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl3),
  inference(forward_demodulation,[],[f7329,f217])).
fof(f217,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X6),'const.option.NONE_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.option.IS_NONE_DEF|split|1')).
fof(f7329,plain,(
  ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | $spl3),
  inference(superposition,[],[f367,f357])).
fof(f21947,plain,(
  ( ! [X72,X73] : (s('type.option.option'(X72),'const.option.NONE_0') = s('type.option.option'(X72),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(X72),'const.option.SOME_1'(s(X72,X73))),s('type.option.option'(X72),'const.option.NONE_0')))) ) | ($spl5 | $spl28)),
  inference(forward_demodulation,[],[f15495,f1098])).
fof(f15495,plain,(
  ( ! [X72,X73] : (s('type.option.option'(X72),'const.option.NONE_0') = s('type.option.option'(X72),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(X72),'const.option.SOME_1'(s(X72,X73))),s('type.option.option'(X72),'const.option.NONE_0')))) ) | $spl5),
  inference(resolution,[],[f368,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) | ~p(s(bool,X2))) & (p(s(bool,X2)) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0')))))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) <=> ~p(s(bool,X2)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0'))) <=> ~p(s(bool,X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X12,X8] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))),s('type.option.option'(X6),'const.option.NONE_0'))) <=> ~p(s(bool,X8)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.option.IF_EQUALS_OPTION|split|0')).
fof(f12603,plain,(
  ( ! [X10,X8,X11] : (sP3(X10,X8,X11) | ~sP0(X11,X8,X10) | ~sP2(X10,X11,X8) | p(s(bool,X8))) ) | $spl366),
  inference(cnf_transformation,[],[f12603_D])).
fof(f12603_D,plain,(
  ( ! [X10,X8,X11] : (sP3(X10,X8,X11) | ~sP0(X11,X8,X10) | ~sP2(X10,X11,X8) | p(s(bool,X8))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f3057,plain,(
  ( ! [X20] : (sP4(X20,'F_0')) ) | $spl84),
  inference(cnf_transformation,[],[f3057_D])).
fof(f3057_D,plain,(
  ( ! [X20] : (sP4(X20,'F_0')) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f21884,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | ($spl76 | $spl352)),
  inference(subsumption_resolution,[],[f21872,f17268])).
fof(f17268,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4(X76,'F_0')) ) | $spl76),
  inference(resolution,[],[f2239,f255])).
fof(f255,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP4(X0,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f21872,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(X8,'F_0') | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl352),
  inference(resolution,[],[f12021,f231])).
fof(f231,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f152])).
fof(f152,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f3069,plain,(
  sP4('F_0','F_0') | $spl90),
  inference(cnf_transformation,[],[f3069_D])).
fof(f3069_D,plain,(
  sP4('F_0','F_0') <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f12019,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP1(X1,X2,X3)) ) | $spl348),
  inference(cnf_transformation,[],[f12019_D])).
fof(f12019_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP1(X1,X2,X3)) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f21865,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP1(X33,X34,X35) | p(s(bool,X32)) | sP4('F_0','F_0')) ) | ($spl76 | $spl352)),
  inference(resolution,[],[f12021,f17267])).
fof(f17267,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4(X74,'F_0')) ) | $spl76),
  inference(resolution,[],[f2239,f255])).
fof(f10683,plain,(
  ( ! [X8,X9] : (sP4('F_0',X8) | ~p(s(bool,X9)) | sP4(X8,X9)) ) | $spl270),
  inference(cnf_transformation,[],[f10683_D])).
fof(f10683_D,plain,(
  ( ! [X8,X9] : (sP4('F_0',X8) | ~p(s(bool,X9)) | sP4(X8,X9)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f21829,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4('F_0',X19) | ~p(s(bool,X21)) | ~p(s(bool,X20))) ) | ($spl5 | $spl350)),
  inference(resolution,[],[f12020,f15460])).
fof(f21828,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17) | sP4('F_0',X16) | ~p(s(bool,X18)) | ~p(s(bool,X17))) ) | ($spl3 | $spl350)),
  inference(resolution,[],[f12020,f7283])).
fof(f7283,plain,(
  ( ! [X37,X36] : (~sP1(X36,X37,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ~p(s(bool,X37)) | ~p(s(bool,X36))) ) | $spl3),
  inference(resolution,[],[f367,f237])).
fof(f21824,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl350),
  inference(resolution,[],[f12020,f237])).
fof(f21820,plain,(
  ( ! [X39,X41,X42,X40] : (sP4(X39,X40) | sP1(X40,X41,X42) | p(s(bool,X39)) | sP4('F_0','F_0')) ) | ($spl76 | $spl350)),
  inference(resolution,[],[f12020,f17268])).
fof(f11980,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP1(X2,X1,X3)) ) | $spl340),
  inference(cnf_transformation,[],[f11980_D])).
fof(f11980_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP1(X2,X1,X3)) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f21772,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP1(X34,X33,X35) | p(s(bool,X32)) | sP4('F_0','F_0')) ) | ($spl76 | $spl344)),
  inference(resolution,[],[f11982,f17267])).
fof(f21689,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4('F_0',X19) | ~p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl5 | $spl342)),
  inference(resolution,[],[f11981,f15460])).
fof(f21688,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17) | sP4('F_0',X16) | ~p(s(bool,X17)) | ~p(s(bool,X18))) ) | ($spl3 | $spl342)),
  inference(resolution,[],[f11981,f7283])).
fof(f21684,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X0) | ~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X3))) ) | $spl342),
  inference(resolution,[],[f11981,f237])).
fof(f21680,plain,(
  ( ! [X39,X41,X42,X40] : (sP4(X39,X40) | sP1(X41,X40,X42) | p(s(bool,X39)) | sP4('F_0','F_0')) ) | ($spl76 | $spl342)),
  inference(resolution,[],[f11981,f17268])).
fof(f12510,plain,(
  ( ! [X14,X12,X15] : (sP3(X12,X14,X15) | ~sP0(X15,X14,X12) | ~sP1(X14,X15,X12) | p(s(bool,X12))) ) | $spl364),
  inference(cnf_transformation,[],[f12510_D])).
fof(f12510_D,plain,(
  ( ! [X14,X12,X15] : (sP3(X12,X14,X15) | ~sP0(X15,X14,X12) | ~sP1(X14,X15,X12) | p(s(bool,X12))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f21651,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'F_0') | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | ($spl76 | $spl336)),
  inference(subsumption_resolution,[],[f21639,f17268])).
fof(f21639,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(X12,'F_0') | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl336),
  inference(resolution,[],[f11942,f232])).
fof(f232,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f11940,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP2(X1,X2,X3)) ) | $spl332),
  inference(cnf_transformation,[],[f11940_D])).
fof(f11940_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP2(X1,X2,X3)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f21631,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP2(X33,X34,X35) | p(s(bool,X32)) | sP4('F_0','F_0')) ) | ($spl76 | $spl336)),
  inference(resolution,[],[f11942,f17267])).
fof(f21595,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4('F_0',X19) | ~p(s(bool,X21)) | ~p(s(bool,X20))) ) | ($spl5 | $spl334)),
  inference(resolution,[],[f11941,f15457])).
fof(f21594,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17) | sP4('F_0',X16) | ~p(s(bool,X18)) | ~p(s(bool,X17))) ) | ($spl3 | $spl334)),
  inference(resolution,[],[f11941,f7280])).
fof(f7280,plain,(
  ( ! [X30,X31] : (~sP2(X30,X31,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ~p(s(bool,X31)) | ~p(s(bool,X30))) ) | $spl3),
  inference(resolution,[],[f367,f233])).
fof(f21590,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl334),
  inference(resolution,[],[f11941,f233])).
fof(f21586,plain,(
  ( ! [X39,X41,X42,X40] : (sP4(X39,X40) | sP2(X40,X41,X42) | p(s(bool,X39)) | sP4('F_0','F_0')) ) | ($spl76 | $spl334)),
  inference(resolution,[],[f11941,f17268])).
fof(f11902,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP2(X2,X1,X3)) ) | $spl324),
  inference(cnf_transformation,[],[f11902_D])).
fof(f11902_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP2(X2,X1,X3)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f21539,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP2(X34,X33,X35) | p(s(bool,X32)) | sP4('F_0','F_0')) ) | ($spl76 | $spl328)),
  inference(resolution,[],[f11904,f17267])).
fof(f21503,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4('F_0',X19) | ~p(s(bool,X20)) | ~p(s(bool,X21))) ) | ($spl5 | $spl326)),
  inference(resolution,[],[f11903,f15457])).
fof(f21502,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17) | sP4('F_0',X16) | ~p(s(bool,X17)) | ~p(s(bool,X18))) ) | ($spl3 | $spl326)),
  inference(resolution,[],[f11903,f7280])).
fof(f21498,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X0) | ~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X3))) ) | $spl326),
  inference(resolution,[],[f11903,f233])).
fof(f21494,plain,(
  ( ! [X39,X41,X42,X40] : (sP4(X39,X40) | sP2(X41,X40,X42) | p(s(bool,X39)) | sP4('F_0','F_0')) ) | ($spl76 | $spl326)),
  inference(resolution,[],[f11903,f17268])).
fof(f21400,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X9,'F_0') | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | ($spl76 | $spl312)),
  inference(subsumption_resolution,[],[f21388,f17267])).
fof(f21388,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(X9,'F_0') | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl312),
  inference(resolution,[],[f11641,f231])).
fof(f11639,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X1)) | sP1(X0,X2,X3)) ) | $spl308),
  inference(cnf_transformation,[],[f11639_D])).
fof(f11639_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X1)) | sP1(X0,X2,X3)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f21381,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP1(X32,X34,X35) | p(s(bool,X33)) | sP4('F_0','F_0')) ) | ($spl76 | $spl312)),
  inference(resolution,[],[f11641,f17267])).
fof(f10570,plain,(
  ( ! [X8,X9] : (sP4('F_0',X8) | ~p(s(bool,X9)) | sP4(X9,X8)) ) | $spl268),
  inference(cnf_transformation,[],[f10570_D])).
fof(f10570_D,plain,(
  ( ! [X8,X9] : (sP4('F_0',X8) | ~p(s(bool,X9)) | sP4(X9,X8)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f21345,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4('F_0',X20) | ~p(s(bool,X21)) | ~p(s(bool,X19))) ) | ($spl5 | $spl310)),
  inference(resolution,[],[f11640,f15460])).
fof(f21344,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17) | sP4('F_0',X17) | ~p(s(bool,X18)) | ~p(s(bool,X16))) ) | ($spl3 | $spl310)),
  inference(resolution,[],[f11640,f7283])).
fof(f21340,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X1) | ~p(s(bool,X0)) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl310),
  inference(resolution,[],[f11640,f237])).
fof(f21336,plain,(
  ( ! [X39,X41,X42,X40] : (sP4(X39,X40) | sP1(X39,X41,X42) | p(s(bool,X40)) | sP4('F_0','F_0')) ) | ($spl76 | $spl310)),
  inference(resolution,[],[f11640,f17268])).
fof(f11600,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X1)) | sP1(X2,X0,X3)) ) | $spl300),
  inference(cnf_transformation,[],[f11600_D])).
fof(f11600_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X1)) | sP1(X2,X0,X3)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f21288,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP1(X34,X32,X35) | p(s(bool,X33)) | sP4('F_0','F_0')) ) | ($spl76 | $spl304)),
  inference(resolution,[],[f11602,f17267])).
fof(f21249,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4('F_0',X20) | ~p(s(bool,X19)) | ~p(s(bool,X21))) ) | ($spl5 | $spl302)),
  inference(resolution,[],[f11601,f15460])).
fof(f21248,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17) | sP4('F_0',X17) | ~p(s(bool,X16)) | ~p(s(bool,X18))) ) | ($spl3 | $spl302)),
  inference(resolution,[],[f11601,f7283])).
fof(f21244,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X1) | ~p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X3))) ) | $spl302),
  inference(resolution,[],[f11601,f237])).
fof(f21240,plain,(
  ( ! [X39,X41,X42,X40] : (sP4(X39,X40) | sP1(X41,X39,X42) | p(s(bool,X40)) | sP4('F_0','F_0')) ) | ($spl76 | $spl302)),
  inference(resolution,[],[f11601,f17268])).
fof(f21211,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X13,'F_0') | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | ($spl76 | $spl296)),
  inference(subsumption_resolution,[],[f21199,f17267])).
fof(f21199,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(X13,'F_0') | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl296),
  inference(resolution,[],[f11562,f232])).
fof(f11560,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X1)) | sP2(X0,X2,X3)) ) | $spl292),
  inference(cnf_transformation,[],[f11560_D])).
fof(f11560_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X1)) | sP2(X0,X2,X3)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f21191,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP2(X32,X34,X35) | p(s(bool,X33)) | sP4('F_0','F_0')) ) | ($spl76 | $spl296)),
  inference(resolution,[],[f11562,f17267])).
fof(f21155,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4('F_0',X20) | ~p(s(bool,X21)) | ~p(s(bool,X19))) ) | ($spl5 | $spl294)),
  inference(resolution,[],[f11561,f15457])).
fof(f21154,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17) | sP4('F_0',X17) | ~p(s(bool,X18)) | ~p(s(bool,X16))) ) | ($spl3 | $spl294)),
  inference(resolution,[],[f11561,f7280])).
fof(f21150,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X1) | ~p(s(bool,X0)) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl294),
  inference(resolution,[],[f11561,f233])).
fof(f21146,plain,(
  ( ! [X39,X41,X42,X40] : (sP4(X39,X40) | sP2(X39,X41,X42) | p(s(bool,X40)) | sP4('F_0','F_0')) ) | ($spl76 | $spl294)),
  inference(resolution,[],[f11561,f17268])).
fof(f11521,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X1)) | sP2(X2,X0,X3)) ) | $spl284),
  inference(cnf_transformation,[],[f11521_D])).
fof(f11521_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | p(s(bool,X1)) | sP2(X2,X0,X3)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f21099,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP2(X34,X32,X35) | p(s(bool,X33)) | sP4('F_0','F_0')) ) | ($spl76 | $spl288)),
  inference(resolution,[],[f11523,f17267])).
fof(f21063,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20) | sP4('F_0',X20) | ~p(s(bool,X19)) | ~p(s(bool,X21))) ) | ($spl5 | $spl286)),
  inference(resolution,[],[f11522,f15457])).
fof(f21062,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17) | sP4('F_0',X17) | ~p(s(bool,X16)) | ~p(s(bool,X18))) ) | ($spl3 | $spl286)),
  inference(resolution,[],[f11522,f7280])).
fof(f21058,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X1) | ~p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X3))) ) | $spl286),
  inference(resolution,[],[f11522,f233])).
fof(f21054,plain,(
  ( ! [X39,X41,X42,X40] : (sP4(X39,X40) | sP2(X41,X39,X42) | p(s(bool,X40)) | sP4('F_0','F_0')) ) | ($spl76 | $spl286)),
  inference(resolution,[],[f11522,f17268])).
fof(f3067,plain,(
  ( ! [X144,X145] : (p(s(bool,X144)) | ~sP5(X144,X145) | p(s(bool,X145))) ) | $spl88),
  inference(cnf_transformation,[],[f3067_D])).
fof(f3067_D,plain,(
  ( ! [X144,X145] : (p(s(bool,X144)) | ~sP5(X144,X145) | p(s(bool,X145))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f20952,plain,(
  ( ! [X220,X219,X221] : (sP4(X219,'F_0') | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f20912])).
fof(f20912,plain,(
  ( ! [X220,X219,X221] : (sP4(X219,'F_0') | sP4(X219,'F_0') | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl76),
  inference(resolution,[],[f17270,f271])).
fof(f17270,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4(X80,'F_0')) ) | $spl76),
  inference(resolution,[],[f2239,f260])).
fof(f20954,plain,(
  ( ! [X215,X213,X214] : (sP4(X213,'F_0') | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f20910])).
fof(f20910,plain,(
  ( ! [X215,X213,X214] : (sP4(X213,'F_0') | sP4(X213,'F_0') | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl76),
  inference(resolution,[],[f17270,f275])).
fof(f275,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : ((~sP5(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP5(X0,X1)))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : ((~sP5(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP5(X0,X1)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : (sP5(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f124,f144])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X19,X6] : ((~(p(s(bool,X6)) | p(s(bool,X19))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X19)) => p(s(bool,'F_0')))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.sat.sat_9_prove_ltb')).
fof(f9792,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X0,'F_0')) ) | $spl250),
  inference(cnf_transformation,[],[f9792_D])).
fof(f9792_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X0,'F_0')) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f20891,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X182,'F_0') | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl76),
  inference(resolution,[],[f17270,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f222])).
fof(f222,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X12,X16,X13,X17] : ((s(bool,X12) = s(bool,X16) & (p(s(bool,X16)) => s(bool,X13) = s(bool,X17))) => ((p(s(bool,X12)) => p(s(bool,X13))) <=> (p(s(bool,X16)) => p(s(bool,X17)))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.IMP_CONG')).
fof(f449,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) | $spl14),
  inference(cnf_transformation,[],[f449_D])).
fof(f449_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f20846,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X35,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl76),
  inference(resolution,[],[f17270,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f20845,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X30,'F_0') | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl76),
  inference(resolution,[],[f17270,f224])).
fof(f224,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f446,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) | $spl8),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f20843,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X20,'F_0') | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl76),
  inference(resolution,[],[f17270,f223])).
fof(f223,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f20841,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X10,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl76),
  inference(resolution,[],[f17270,f221])).
fof(f221,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f20839,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl76),
  inference(resolution,[],[f17270,f219])).
fof(f219,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f150])).
fof(f9566,plain,(
  ( ! [X229] : (sP4('F_0',X229) | sP4(X229,'F_0')) ) | $spl248),
  inference(cnf_transformation,[],[f9566_D])).
fof(f9566_D,plain,(
  ( ! [X229] : (sP4('F_0',X229) | sP4(X229,'F_0')) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f20793,plain,(
  ( ! [X220,X219,X221] : (sP4('F_0',X219) | sP4(X219,'F_0') | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl76),
  inference(resolution,[],[f17269,f271])).
fof(f17269,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4(X78,'F_0')) ) | $spl76),
  inference(resolution,[],[f2239,f257])).
fof(f20791,plain,(
  ( ! [X215,X213,X214] : (sP4('F_0',X213) | sP4(X213,'F_0') | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl76),
  inference(resolution,[],[f17269,f275])).
fof(f9558,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X1,'F_0')) ) | $spl242),
  inference(cnf_transformation,[],[f9558_D])).
fof(f9558_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X1,'F_0')) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f20772,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X183,'F_0') | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl76),
  inference(resolution,[],[f17269,f360])).
fof(f20727,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X36,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl76),
  inference(resolution,[],[f17269,f225])).
fof(f20726,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X31,'F_0') | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl76),
  inference(resolution,[],[f17269,f224])).
fof(f20724,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X21,'F_0') | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl76),
  inference(resolution,[],[f17269,f223])).
fof(f20722,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X11,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl76),
  inference(resolution,[],[f17269,f221])).
fof(f20720,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl76),
  inference(resolution,[],[f17269,f219])).
fof(f8351,plain,(
  ( ! [X0,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP4('T_0',X1)) ) | $spl218),
  inference(cnf_transformation,[],[f8351_D])).
fof(f8351_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1) | p(s(bool,X0)) | sP4('T_0',X1)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f20515,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4('T_0',X17) | p(s(bool,X16)) | sP4('F_0','F_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204 | $spl380)),
  inference(resolution,[],[f20320,f17267])).
fof(f20320,plain,(
  ( ! [X182,X181] : (sP4(X181,'F_0') | sP4(X181,X182) | sP4('T_0',X182)) ) | ($spl5 | $spl28 | $spl76 | $spl204 | $spl380)),
  inference(subsumption_resolution,[],[f20319,f334])).
fof(f334,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0',aHL_FALSITY)).
fof(f20319,plain,(
  ( ! [X182,X181] : (p(s(bool,'F_0')) | sP4(X181,'F_0') | sP4(X181,X182) | sP4('T_0',X182)) ) | ($spl5 | $spl28 | $spl76 | $spl204 | $spl380)),
  inference(forward_demodulation,[],[f20265,f1098])).
fof(f20265,plain,(
  ( ! [X182,X181] : (p(s(bool,sK8)) | sP4(X181,'F_0') | sP4(X181,X182) | sP4('T_0',X182)) ) | ($spl5 | $spl28 | $spl76 | $spl204 | $spl380)),
  inference(resolution,[],[f17267,f18775])).
fof(f18775,plain,(
  ( ! [X0,X1] : (sP4(sK8,X0) | sP4(X0,X1) | sP4('T_0',X1)) ) | ($spl5 | $spl28 | $spl204 | $spl380)),
  inference(resolution,[],[f18676,f13746])).
fof(f18676,plain,(
  ( ! [X33,X34] : (sP4(sK8,X33) | sP4(X33,X34) | sP4(X34,'T_0')) ) | ($spl5 | $spl28 | $spl204)),
  inference(subsumption_resolution,[],[f18675,f334])).
fof(f18675,plain,(
  ( ! [X33,X34] : (p(s(bool,'F_0')) | sP4(sK8,X33) | sP4(X33,X34) | sP4(X34,'T_0')) ) | ($spl5 | $spl28 | $spl204)),
  inference(forward_demodulation,[],[f18639,f1098])).
fof(f18639,plain,(
  ( ! [X33,X34] : (p(s(bool,sK8)) | sP4(sK8,X33) | sP4(X33,X34) | sP4(X34,'T_0')) ) | ($spl5 | $spl204)),
  inference(resolution,[],[f16741,f18065])).
fof(f18065,plain,(
  ( ! [X187,X186] : (sP4(X186,X187) | sP4(X186,sK8) | sP4(X187,'T_0')) ) | ($spl5 | $spl204)),
  inference(resolution,[],[f16623,f7931])).
fof(f16623,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4(X80,sK8)) ) | $spl5),
  inference(resolution,[],[f15470,f260])).
fof(f16741,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4(sK8,X76)) ) | $spl5),
  inference(resolution,[],[f15471,f255])).
fof(f6936,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0) | p(s(bool,X1)) | sP4(X0,X1)) ) | $spl152),
  inference(cnf_transformation,[],[f6936_D])).
fof(f6936_D,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0) | p(s(bool,X1)) | sP4(X0,X1)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f20460,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4('T_0',X16) | p(s(bool,X17)) | sP4('F_0','F_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204 | $spl380)),
  inference(resolution,[],[f20318,f17267])).
fof(f20318,plain,(
  ( ! [X180,X179] : (sP4(X179,'F_0') | sP4(X180,X179) | sP4('T_0',X180)) ) | ($spl5 | $spl28 | $spl76 | $spl204 | $spl380)),
  inference(subsumption_resolution,[],[f20317,f334])).
fof(f20317,plain,(
  ( ! [X180,X179] : (p(s(bool,'F_0')) | sP4(X179,'F_0') | sP4(X180,X179) | sP4('T_0',X180)) ) | ($spl5 | $spl28 | $spl76 | $spl204 | $spl380)),
  inference(forward_demodulation,[],[f20264,f1098])).
fof(f20264,plain,(
  ( ! [X180,X179] : (p(s(bool,sK8)) | sP4(X179,'F_0') | sP4(X180,X179) | sP4('T_0',X180)) ) | ($spl5 | $spl28 | $spl76 | $spl204 | $spl380)),
  inference(resolution,[],[f17267,f18713])).
fof(f18713,plain,(
  ( ! [X0,X1] : (sP4(sK8,X0) | sP4(X1,X0) | sP4('T_0',X1)) ) | ($spl5 | $spl28 | $spl204 | $spl380)),
  inference(resolution,[],[f18672,f13746])).
fof(f18672,plain,(
  ( ! [X30,X29] : (sP4(sK8,X29) | sP4(X30,X29) | sP4(X30,'T_0')) ) | ($spl5 | $spl28 | $spl204)),
  inference(subsumption_resolution,[],[f18671,f334])).
fof(f18671,plain,(
  ( ! [X30,X29] : (p(s(bool,'F_0')) | sP4(sK8,X29) | sP4(X30,X29) | sP4(X30,'T_0')) ) | ($spl5 | $spl28 | $spl204)),
  inference(forward_demodulation,[],[f18637,f1098])).
fof(f18637,plain,(
  ( ! [X30,X29] : (p(s(bool,sK8)) | sP4(sK8,X29) | sP4(X30,X29) | sP4(X30,'T_0')) ) | ($spl5 | $spl204)),
  inference(resolution,[],[f16741,f17833])).
fof(f17833,plain,(
  ( ! [X187,X186] : (sP4(X186,X187) | sP4(X187,sK8) | sP4(X186,'T_0')) ) | ($spl5 | $spl204)),
  inference(resolution,[],[f16622,f7931])).
fof(f16622,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4(X78,sK8)) ) | $spl5),
  inference(resolution,[],[f15470,f257])).
fof(f8466,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0') | p(s(bool,X1)) | sP4(X1,X0)) ) | $spl232),
  inference(cnf_transformation,[],[f8466_D])).
fof(f8466_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0') | p(s(bool,X1)) | sP4(X1,X0)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f20398,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4(X17,'T_0') | p(s(bool,X16)) | sP4('F_0','F_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204)),
  inference(resolution,[],[f20316,f17267])).
fof(f20316,plain,(
  ( ! [X177,X178] : (sP4(X177,'F_0') | sP4(X177,X178) | sP4(X178,'T_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204)),
  inference(subsumption_resolution,[],[f20315,f334])).
fof(f20315,plain,(
  ( ! [X177,X178] : (p(s(bool,'F_0')) | sP4(X177,'F_0') | sP4(X177,X178) | sP4(X178,'T_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204)),
  inference(forward_demodulation,[],[f20263,f1098])).
fof(f20263,plain,(
  ( ! [X177,X178] : (p(s(bool,sK8)) | sP4(X177,'F_0') | sP4(X177,X178) | sP4(X178,'T_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204)),
  inference(resolution,[],[f17267,f18676])).
fof(f8386,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0') | p(s(bool,X1)) | sP4(X0,X1)) ) | $spl224),
  inference(cnf_transformation,[],[f8386_D])).
fof(f8386_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0') | p(s(bool,X1)) | sP4(X0,X1)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f20345,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4(X16,'T_0') | p(s(bool,X17)) | sP4('F_0','F_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204)),
  inference(resolution,[],[f20314,f17267])).
fof(f20314,plain,(
  ( ! [X175,X176] : (sP4(X175,'F_0') | sP4(X176,X175) | sP4(X176,'T_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204)),
  inference(subsumption_resolution,[],[f20313,f334])).
fof(f20313,plain,(
  ( ! [X175,X176] : (p(s(bool,'F_0')) | sP4(X175,'F_0') | sP4(X176,X175) | sP4(X176,'T_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204)),
  inference(forward_demodulation,[],[f20262,f1098])).
fof(f20262,plain,(
  ( ! [X175,X176] : (p(s(bool,sK8)) | sP4(X175,'F_0') | sP4(X176,X175) | sP4(X176,'T_0')) ) | ($spl5 | $spl28 | $spl76 | $spl204)),
  inference(resolution,[],[f17267,f18672])).
fof(f6211,plain,(
  ( ! [X4] : (sP4(X4,sK8)) ) | $spl134),
  inference(cnf_transformation,[],[f6211_D])).
fof(f6211_D,plain,(
  ( ! [X4] : (sP4(X4,sK8)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f20180,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1))) ) | ($spl5 | $spl120)),
  inference(subsumption_resolution,[],[f4405,f16620])).
fof(f16620,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4(X74,sK8)) ) | $spl5),
  inference(resolution,[],[f15470,f255])).
fof(f20179,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1))) ) | ($spl5 | $spl118)),
  inference(subsumption_resolution,[],[f4320,f16621])).
fof(f16621,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4(X76,sK8)) ) | $spl5),
  inference(resolution,[],[f15470,f255])).
fof(f5664,plain,(
  ( ! [X4] : (sP4(sK8,X4)) ) | $spl132),
  inference(cnf_transformation,[],[f5664_D])).
fof(f5664_D,plain,(
  ( ! [X4] : (sP4(sK8,X4)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f20135,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl5),
  inference(subsumption_resolution,[],[f20118,f16741])).
fof(f20118,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(sK8,X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl5),
  inference(resolution,[],[f19060,f231])).
fof(f19060,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(sK8,X72) | sP1(X73,X74,X75)) ) | $spl5),
  inference(resolution,[],[f16743,f240])).
fof(f16743,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4(sK8,X80)) ) | $spl5),
  inference(resolution,[],[f15471,f260])).
fof(f20035,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(sK8,X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl5),
  inference(subsumption_resolution,[],[f20018,f16741])).
fof(f20018,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(sK8,X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl5),
  inference(resolution,[],[f19057,f232])).
fof(f19057,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(sK8,X60) | sP2(X61,X62,X63)) ) | $spl5),
  inference(resolution,[],[f16743,f236])).
fof(f19903,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X9) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl5),
  inference(subsumption_resolution,[],[f19886,f16740])).
fof(f16740,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4(sK8,X74)) ) | $spl5),
  inference(resolution,[],[f15471,f255])).
fof(f19886,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(sK8,X9) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl5),
  inference(resolution,[],[f18939,f231])).
fof(f18939,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(sK8,X73) | sP1(X72,X74,X75)) ) | $spl5),
  inference(resolution,[],[f16742,f240])).
fof(f16742,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4(sK8,X78)) ) | $spl5),
  inference(resolution,[],[f15471,f257])).
fof(f19803,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(sK8,X13) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl5),
  inference(subsumption_resolution,[],[f19786,f16740])).
fof(f19786,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(sK8,X13) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl5),
  inference(resolution,[],[f18936,f232])).
fof(f18936,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(sK8,X61) | sP2(X60,X62,X63)) ) | $spl5),
  inference(resolution,[],[f16742,f236])).
fof(f19672,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,sK8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl5),
  inference(subsumption_resolution,[],[f19653,f16621])).
fof(f19653,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(X8,sK8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl5),
  inference(resolution,[],[f18028,f231])).
fof(f18028,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(X72,sK8) | sP1(X73,X74,X75)) ) | $spl5),
  inference(resolution,[],[f16623,f240])).
fof(f19570,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,sK8) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl5),
  inference(subsumption_resolution,[],[f19551,f16621])).
fof(f19551,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(X12,sK8) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl5),
  inference(resolution,[],[f18025,f232])).
fof(f18025,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(X60,sK8) | sP2(X61,X62,X63)) ) | $spl5),
  inference(resolution,[],[f16623,f236])).
fof(f19436,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X9,sK8) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl5),
  inference(subsumption_resolution,[],[f19417,f16620])).
fof(f19417,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(X9,sK8) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl5),
  inference(resolution,[],[f17796,f231])).
fof(f17796,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(X73,sK8) | sP1(X72,X74,X75)) ) | $spl5),
  inference(resolution,[],[f16622,f240])).
fof(f19334,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X13,sK8) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl5),
  inference(subsumption_resolution,[],[f19315,f16620])).
fof(f19315,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(X13,sK8) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl5),
  inference(resolution,[],[f17793,f232])).
fof(f17793,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(X61,sK8) | sP2(X60,X62,X63)) ) | $spl5),
  inference(resolution,[],[f16622,f236])).
fof(f2211,plain,(
  ( ! [X182] : (sP4(sK8,X182) | sP4(X182,'F_0')) ) | $spl72),
  inference(cnf_transformation,[],[f2211_D])).
fof(f2211_D,plain,(
  ( ! [X182] : (sP4(sK8,X182) | sP4(X182,'F_0')) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f19117,plain,(
  ( ! [X220,X219,X221] : (sP4(X219,'F_0') | sP4(sK8,X219) | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl5),
  inference(resolution,[],[f16743,f271])).
fof(f19115,plain,(
  ( ! [X215,X213,X214] : (sP4(X213,'F_0') | sP4(sK8,X213) | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl5),
  inference(resolution,[],[f16743,f275])).
fof(f19154,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(sK8,X0)) ) | $spl426),
  inference(cnf_transformation,[],[f19154_D])).
fof(f19154_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(sK8,X0)) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f19096,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(sK8,X182) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl5),
  inference(resolution,[],[f16743,f360])).
fof(f19051,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(sK8,X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl5),
  inference(resolution,[],[f16743,f225])).
fof(f19050,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(sK8,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl5),
  inference(resolution,[],[f16743,f224])).
fof(f19048,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(sK8,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f16743,f223])).
fof(f19046,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(sK8,X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl5),
  inference(resolution,[],[f16743,f221])).
fof(f19044,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(sK8,X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl5),
  inference(resolution,[],[f16743,f219])).
fof(f3468,plain,(
  ( ! [X184] : (sP4(sK8,X184) | sP4('F_0',X184)) ) | $spl98),
  inference(cnf_transformation,[],[f3468_D])).
fof(f3468_D,plain,(
  ( ! [X184] : (sP4(sK8,X184) | sP4('F_0',X184)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f18996,plain,(
  ( ! [X220,X219,X221] : (sP4('F_0',X219) | sP4(sK8,X219) | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl5),
  inference(resolution,[],[f16742,f271])).
fof(f18994,plain,(
  ( ! [X215,X213,X214] : (sP4('F_0',X213) | sP4(sK8,X213) | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl5),
  inference(resolution,[],[f16742,f275])).
fof(f19034,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(sK8,X1)) ) | $spl424),
  inference(cnf_transformation,[],[f19034_D])).
fof(f19034_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(sK8,X1)) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f18975,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(sK8,X183) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl5),
  inference(resolution,[],[f16742,f360])).
fof(f18930,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(sK8,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl5),
  inference(resolution,[],[f16742,f225])).
fof(f18929,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(sK8,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl5),
  inference(resolution,[],[f16742,f224])).
fof(f18927,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(sK8,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f16742,f223])).
fof(f18925,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(sK8,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl5),
  inference(resolution,[],[f16742,f221])).
fof(f18923,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(sK8,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl5),
  inference(resolution,[],[f16742,f219])).
fof(f18300,plain,(
  ( ! [X146,X147] : (sP4('F_0','F_0') | p(s(bool,X146)) | p(s(bool,X147)) | ~sP5(X146,X147)) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18202,f271])).
fof(f18202,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | sP4('F_0',X2)) ) | ($spl5 | $spl49 | $spl204)),
  inference(subsumption_resolution,[],[f18195,f361])).
fof(f18195,plain,(
  ( ! [X2] : (sP4('F_0',X2) | ~p(s(bool,X2)) | ~p(s(bool,'T_0'))) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18155,f256])).
fof(f18155,plain,(
  ( ! [X14] : (sP4('F_0',X14) | sP4(X14,'T_0')) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18065,f1390])).
fof(f1390,plain,(
  ~sP4('F_0',sK8) | $spl49),
  inference(cnf_transformation,[],[f1390_D])).
fof(f1390_D,plain,(
  ~sP4('F_0',sK8) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f18298,plain,(
  ( ! [X142,X143] : (sP4('F_0','F_0') | p(s(bool,X142)) | p(s(bool,X143)) | ~sP5(X142,X143)) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18202,f275])).
fof(f3608,plain,(
  ( ! [X2,X0] : (sP4('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl100),
  inference(cnf_transformation,[],[f3608_D])).
fof(f3608_D,plain,(
  ( ! [X2,X0] : (sP4('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f18279,plain,(
  ( ! [X132,X130,X131] : (sP4('F_0',X130) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18202,f360])).
fof(f18234,plain,(
  ( ! [X30,X28,X31,X29] : (sP4('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18202,f225])).
fof(f18233,plain,(
  ( ! [X26,X24,X27,X25] : (sP4('F_0',X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18202,f224])).
fof(f18231,plain,(
  ( ! [X19,X17,X18,X16] : (sP4('F_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18202,f223])).
fof(f18229,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18202,f221])).
fof(f18227,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | ($spl5 | $spl49 | $spl204)),
  inference(resolution,[],[f18202,f219])).
fof(f18128,plain,(
  ( ! [X213] : (sP4(X213,'F_0') | sP4(X213,sK8)) ) | $spl422),
  inference(cnf_transformation,[],[f18128_D])).
fof(f18128_D,plain,(
  ( ! [X213] : (sP4(X213,'F_0') | sP4(X213,sK8)) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f18085,plain,(
  ( ! [X220,X219,X221] : (sP4(X219,'F_0') | sP4(X219,sK8) | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl5),
  inference(resolution,[],[f16623,f271])).
fof(f18083,plain,(
  ( ! [X215,X213,X214] : (sP4(X213,'F_0') | sP4(X213,sK8) | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl5),
  inference(resolution,[],[f16623,f275])).
fof(f18123,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X0,sK8)) ) | $spl420),
  inference(cnf_transformation,[],[f18123_D])).
fof(f18123_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X0,sK8)) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f18064,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X182,sK8) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl5),
  inference(resolution,[],[f16623,f360])).
fof(f18019,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X35,sK8) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl5),
  inference(resolution,[],[f16623,f225])).
fof(f18018,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X30,sK8) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl5),
  inference(resolution,[],[f16623,f224])).
fof(f18016,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X20,sK8) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f16623,f223])).
fof(f18014,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X10,sK8) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl5),
  inference(resolution,[],[f16623,f221])).
fof(f18012,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,sK8) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl5),
  inference(resolution,[],[f16623,f219])).
fof(f17895,plain,(
  ( ! [X213] : (sP4('F_0',X213) | sP4(X213,sK8)) ) | $spl418),
  inference(cnf_transformation,[],[f17895_D])).
fof(f17895_D,plain,(
  ( ! [X213] : (sP4('F_0',X213) | sP4(X213,sK8)) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f17853,plain,(
  ( ! [X220,X219,X221] : (sP4('F_0',X219) | sP4(X219,sK8) | p(s(bool,X220)) | p(s(bool,X221)) | ~sP5(X220,X221)) ) | $spl5),
  inference(resolution,[],[f16622,f271])).
fof(f17851,plain,(
  ( ! [X215,X213,X214] : (sP4('F_0',X213) | sP4(X213,sK8) | p(s(bool,X214)) | p(s(bool,X215)) | ~sP5(X214,X215)) ) | $spl5),
  inference(resolution,[],[f16622,f275])).
fof(f17890,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X1,sK8)) ) | $spl416),
  inference(cnf_transformation,[],[f17890_D])).
fof(f17890_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4(X1,sK8)) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f17832,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X183,sK8) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl5),
  inference(resolution,[],[f16622,f360])).
fof(f17787,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X36,sK8) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl5),
  inference(resolution,[],[f16622,f225])).
fof(f17786,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X31,sK8) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl5),
  inference(resolution,[],[f16622,f224])).
fof(f17784,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X21,sK8) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f16622,f223])).
fof(f17782,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X11,sK8) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl5),
  inference(resolution,[],[f16622,f221])).
fof(f17780,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,sK8) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl5),
  inference(resolution,[],[f16622,f219])).
fof(f6786,plain,(
  ( ! [X2,X0] : (sP4('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl148),
  inference(cnf_transformation,[],[f6786_D])).
fof(f6786_D,plain,(
  ( ! [X2,X0] : (sP4('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f17511,plain,(
  ( ! [X127,X125,X126] : (sP4('T_0',X125) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | ($spl5 | $spl204 | $spl380)),
  inference(resolution,[],[f16876,f360])).
fof(f16876,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP4('T_0',X0)) ) | ($spl5 | $spl204 | $spl380)),
  inference(resolution,[],[f16861,f15469])).
fof(f15469,plain,(
  ( ! [X53] : (~sP4(sK8,X53) | p(s(bool,X53))) ) | $spl5),
  inference(resolution,[],[f368,f255])).
fof(f16861,plain,(
  ( ! [X0] : (sP4(sK8,X0) | sP4('T_0',X0)) ) | ($spl5 | $spl204 | $spl380)),
  inference(resolution,[],[f16771,f13746])).
fof(f16771,plain,(
  ( ! [X133] : (sP4(sK8,X133) | sP4(X133,'T_0')) ) | ($spl5 | $spl204)),
  inference(resolution,[],[f15471,f7931])).
fof(f17469,plain,(
  ( ! [X30,X28,X31,X29] : (sP4('T_0',X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | ($spl5 | $spl204 | $spl380)),
  inference(resolution,[],[f16876,f225])).
fof(f17466,plain,(
  ( ! [X19,X17,X18,X16] : (sP4('T_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | ($spl5 | $spl204 | $spl380)),
  inference(resolution,[],[f16876,f224])).
fof(f17465,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('T_0',X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | ($spl5 | $spl204 | $spl380)),
  inference(resolution,[],[f16876,f223])).
fof(f17464,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('T_0',X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | ($spl5 | $spl204 | $spl380)),
  inference(resolution,[],[f16876,f221])).
fof(f17462,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | ($spl5 | $spl204 | $spl380)),
  inference(resolution,[],[f16876,f219])).
fof(f17317,plain,(
  ( ! [X146,X145] : (sP4('F_0','F_0') | p(s(bool,X145)) | p(s(bool,X146)) | ~sP5(X145,X146)) ) | $spl76),
  inference(resolution,[],[f2239,f271])).
fof(f17315,plain,(
  ( ! [X142,X141] : (sP4('F_0','F_0') | p(s(bool,X141)) | p(s(bool,X142)) | ~sP5(X141,X142)) ) | $spl76),
  inference(resolution,[],[f2239,f275])).
fof(f3054,plain,(
  ( ! [X2,X0] : (sP4(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl80),
  inference(cnf_transformation,[],[f3054_D])).
fof(f3054_D,plain,(
  ( ! [X2,X0] : (sP4(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f17297,plain,(
  ( ! [X132,X130,X131] : (sP4(X130,'F_0') | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl76),
  inference(resolution,[],[f2239,f360])).
fof(f17252,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl76),
  inference(resolution,[],[f2239,f225])).
fof(f17251,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,'F_0') | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl76),
  inference(resolution,[],[f2239,f224])).
fof(f17249,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl76),
  inference(resolution,[],[f2239,f223])).
fof(f17247,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl76),
  inference(resolution,[],[f2239,f221])).
fof(f17245,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl76),
  inference(resolution,[],[f2239,f219])).
fof(f17102,plain,(
  $false | ($spl20 | $spl28)),
  inference(subsumption_resolution,[],[f17056,f16991])).
fof(f16991,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl20 | $spl28)),
  inference(subsumption_resolution,[],[f16990,f323])).
fof(f323,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f16990,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1))) ) | ($spl20 | $spl28)),
  inference(forward_demodulation,[],[f532,f1098])).
fof(f17056,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl28)),
  inference(resolution,[],[f16991,f338])).
fof(f338,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f17101,plain,(
  $false | ($spl20 | $spl28)),
  inference(subsumption_resolution,[],[f17055,f16991])).
fof(f17055,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl28)),
  inference(resolution,[],[f16991,f336])).
fof(f336,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f22])).
fof(f22,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.NOT_CLAUSES|split|2')).
fof(f17100,plain,(
  $false | ($spl20 | $spl28)),
  inference(subsumption_resolution,[],[f17049,f16991])).
fof(f17049,plain,(
  ( ! [X129] : (~p(s(bool,X129))) ) | ($spl20 | $spl28)),
  inference(resolution,[],[f16991,f291])).
fof(f291,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.IMP_CLAUSES|split|4')).
fof(f17054,plain,(
  $false | ($spl20 | $spl28)),
  inference(resolution,[],[f16991,f334])).
fof(f16827,plain,(
  ( ! [X2,X0] : (sP4(sK8,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl414),
  inference(cnf_transformation,[],[f16827_D])).
fof(f16827_D,plain,(
  ( ! [X2,X0] : (sP4(sK8,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f16770,plain,(
  ( ! [X132,X130,X131] : (sP4(sK8,X130) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl5),
  inference(resolution,[],[f15471,f360])).
fof(f16725,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(sK8,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl5),
  inference(resolution,[],[f15471,f225])).
fof(f16724,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(sK8,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl5),
  inference(resolution,[],[f15471,f224])).
fof(f16722,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(sK8,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl5),
  inference(resolution,[],[f15471,f223])).
fof(f16720,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl5),
  inference(resolution,[],[f15471,f221])).
fof(f16718,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK8,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl5),
  inference(resolution,[],[f15471,f219])).
fof(f16707,plain,(
  ( ! [X2,X0] : (sP4(X0,sK8) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl412),
  inference(cnf_transformation,[],[f16707_D])).
fof(f16707_D,plain,(
  ( ! [X2,X0] : (sP4(X0,sK8) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f16650,plain,(
  ( ! [X132,X130,X131] : (sP4(X130,sK8) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl5),
  inference(resolution,[],[f15470,f360])).
fof(f16605,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,sK8) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl5),
  inference(resolution,[],[f15470,f225])).
fof(f16604,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,sK8) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl5),
  inference(resolution,[],[f15470,f224])).
fof(f16602,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,sK8) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl5),
  inference(resolution,[],[f15470,f223])).
fof(f16600,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,sK8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl5),
  inference(resolution,[],[f15470,f221])).
fof(f16598,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl5),
  inference(resolution,[],[f15470,f219])).
fof(f16200,plain,(
  sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | $spl410),
  inference(cnf_transformation,[],[f16200_D])).
fof(f16200_D,plain,(
  sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f16195,plain,(
  ( ! [X4] : (sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X4))) ) | ($spl3 | $spl142)),
  inference(resolution,[],[f6650,f7359])).
fof(f7359,plain,(
  ( ! [X60] : (~sP5(X60,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X60))) ) | $spl3),
  inference(subsumption_resolution,[],[f7299,f309])).
fof(f7299,plain,(
  ( ! [X60] : (p(s(bool,'F_0')) | p(s(bool,X60)) | ~sP5(X60,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f271])).
fof(f6663,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP4('T_0',X0)) ) | $spl146),
  inference(cnf_transformation,[],[f6663_D])).
fof(f6663_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP4('T_0',X0)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f16198,plain,(
  ( ! [X2,X3] : (sP4('T_0',X2) | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl142),
  inference(subsumption_resolution,[],[f16194,f309])).
fof(f16194,plain,(
  ( ! [X2,X3] : (sP4('T_0',X2) | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl142),
  inference(resolution,[],[f6650,f271])).
fof(f16197,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl142),
  inference(subsumption_resolution,[],[f16193,f309])).
fof(f16193,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl142),
  inference(resolution,[],[f6650,f275])).
fof(f16188,plain,(
  $false | ($spl5 | $spl74)),
  inference(subsumption_resolution,[],[f16183,f368])).
fof(f16183,plain,(
  p(s(bool,sK8)) | $spl74),
  inference(duplicate_literal_removal,[],[f16182])).
fof(f16182,plain,(
  p(s(bool,sK8)) | p(s(bool,sK8)) | $spl74),
  inference(resolution,[],[f2238,f264])).
fof(f264,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP4(X1,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f16187,plain,(
  $false | ($spl5 | $spl74)),
  inference(subsumption_resolution,[],[f16186,f368])).
fof(f16186,plain,(
  p(s(bool,sK8)) | $spl74),
  inference(duplicate_literal_removal,[],[f16179])).
fof(f16179,plain,(
  p(s(bool,sK8)) | p(s(bool,sK8)) | $spl74),
  inference(resolution,[],[f2238,f255])).
fof(f16177,plain,(
  $false | ($spl5 | $spl29)),
  inference(subsumption_resolution,[],[f16170,f368])).
fof(f16170,plain,(
  p(s(bool,sK8)) | $spl29),
  inference(trivial_inequality_removal,[],[f16158])).
fof(f16158,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,sK8)) | $spl29),
  inference(superposition,[],[f1097,f323])).
fof(f1097,plain,(
  s(bool,'F_0') != s(bool,sK8) | $spl29),
  inference(cnf_transformation,[],[f1097_D])).
fof(f1097_D,plain,(
  s(bool,'F_0') != s(bool,sK8) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f16176,plain,(
  $false | ($spl5 | $spl29)),
  inference(subsumption_resolution,[],[f16171,f368])).
fof(f16171,plain,(
  p(s(bool,sK8)) | $spl29),
  inference(trivial_inequality_removal,[],[f16157])).
fof(f16157,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,sK8)) | $spl29),
  inference(superposition,[],[f1097,f321])).
fof(f321,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f16145,plain,(
  $false | ($spl49 | $spl75 | $spl98)),
  inference(subsumption_resolution,[],[f16138,f1390])).
fof(f16138,plain,(
  sP4('F_0',sK8) | ($spl75 | $spl98)),
  inference(resolution,[],[f3468,f2237])).
fof(f2237,plain,(
  ~sP4(sK8,sK8) | $spl75),
  inference(cnf_transformation,[],[f2237_D])).
fof(f2237_D,plain,(
  ~sP4(sK8,sK8) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f16137,plain,(
  $false | ($spl43 | $spl72 | $spl75)),
  inference(subsumption_resolution,[],[f16129,f2237])).
fof(f16129,plain,(
  sP4(sK8,sK8) | ($spl43 | $spl72)),
  inference(resolution,[],[f2211,f1271])).
fof(f1271,plain,(
  ~sP4(sK8,'F_0') | $spl43),
  inference(cnf_transformation,[],[f1271_D])).
fof(f1271_D,plain,(
  ~sP4(sK8,'F_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f16130,plain,(
  $false | ($spl43 | $spl72 | $spl75)),
  inference(subsumption_resolution,[],[f16117,f1271])).
fof(f16117,plain,(
  sP4(sK8,'F_0') | ($spl72 | $spl75)),
  inference(resolution,[],[f2211,f2237])).
fof(f16112,plain,(
  $false | ($spl5 | $spl231)),
  inference(subsumption_resolution,[],[f15255,f368])).
fof(f15255,plain,(
  p(s(bool,sK8)) | $spl231),
  inference(subsumption_resolution,[],[f15163,f361])).
fof(f15163,plain,(
  p(s(bool,sK8)) | ~p(s(bool,'T_0')) | $spl231),
  inference(resolution,[],[f8432,f257])).
fof(f8432,plain,(
  ~sP4('T_0',sK8) | $spl231),
  inference(cnf_transformation,[],[f8432_D])).
fof(f8432_D,plain,(
  ~sP4('T_0',sK8) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f1098,plain,(
  s(bool,'F_0') = s(bool,sK8) | $spl28),
  inference(cnf_transformation,[],[f1098_D])).
fof(f1098_D,plain,(
  s(bool,'F_0') = s(bool,sK8) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1100,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl31),
  inference(cnf_transformation,[],[f1100_D])).
fof(f1100_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f16073,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f365,f358])).
fof(f358,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0',aHL_BOOL_CASES)).
fof(f16060,plain,(
  $false | ($spl5 | $spl48)),
  inference(subsumption_resolution,[],[f15428,f368])).
fof(f15428,plain,(
  p(s(bool,sK8)) | $spl48),
  inference(subsumption_resolution,[],[f15423,f333])).
fof(f333,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11] : (p(s(bool,'F_0')) => p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.FALSITY')).
fof(f15423,plain,(
  p(s(bool,sK8)) | p(s(bool,'F_0')) | $spl48),
  inference(resolution,[],[f1391,f255])).
fof(f16052,plain,(
  $false | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f15420,f368])).
fof(f15420,plain,(
  p(s(bool,sK8)) | $spl42),
  inference(subsumption_resolution,[],[f15415,f333])).
fof(f15415,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK8)) | $spl42),
  inference(resolution,[],[f1272,f255])).
fof(f445,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK8) != s(bool,X1)) ) | $spl6),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK8) != s(bool,X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f16019,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,sK8) != s(bool,X79)) ) | $spl5),
  inference(resolution,[],[f368,f360])).
fof(f15974,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f368,f225])).
fof(f15973,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,sK8) != s(bool,X19)) ) | $spl5),
  inference(resolution,[],[f368,f224])).
fof(f15971,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,sK8) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f15969,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f15967,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f15966,plain,(
  $false | ($spl28 | $spl42)),
  inference(subsumption_resolution,[],[f15965,f334])).
fof(f15965,plain,(
  p(s(bool,'F_0')) | ($spl28 | $spl42)),
  inference(forward_demodulation,[],[f15420,f1098])).
fof(f15554,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f365,f358])).
fof(f15527,plain,(
  $false | ($spl5 | $spl146 | $spl231)),
  inference(subsumption_resolution,[],[f15499,f8432])).
fof(f15499,plain,(
  sP4('T_0',sK8) | ($spl5 | $spl146)),
  inference(resolution,[],[f368,f6663])).
fof(f15498,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,sK8) != s(bool,X79)) ) | $spl5),
  inference(resolution,[],[f368,f360])).
fof(f15453,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f368,f225])).
fof(f15452,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,sK8) != s(bool,X19)) ) | $spl5),
  inference(resolution,[],[f368,f224])).
fof(f15450,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,sK8) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f15448,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f15446,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f1270,plain,(
  ( ! [X130,X129] : (sP5(X129,X130)) ) | $spl40),
  inference(cnf_transformation,[],[f1270_D])).
fof(f1270_D,plain,(
  ( ! [X130,X129] : (sP5(X129,X130)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3064,plain,(
  sP4('T_0','F_0') | $spl86),
  inference(cnf_transformation,[],[f3064_D])).
fof(f3064_D,plain,(
  sP4('T_0','F_0') <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f6740,plain,(
  ( ! [X134,X135] : (sP4('T_0','F_0') | sP5(X134,X135)) ) | $spl146),
  inference(resolution,[],[f6663,f274])).
fof(f274,plain,(
  ( ! [X0,X1] : (sP5(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f172])).
fof(f6741,plain,(
  ( ! [X136,X137] : (sP4('T_0','F_0') | sP5(X136,X137)) ) | $spl146),
  inference(resolution,[],[f6663,f279])).
fof(f279,plain,(
  ( ! [X0,X1] : (sP5(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f174])).
fof(f6742,plain,(
  ( ! [X138] : (sP4('T_0','F_0') | ~p(s(bool,X138))) ) | $spl146),
  inference(resolution,[],[f6663,f291])).
fof(f15413,plain,(
  ( ! [X139] : (sP4('T_0','F_0') | p(s(bool,X139))) ) | $spl146),
  inference(subsumption_resolution,[],[f6743,f361])).
fof(f6743,plain,(
  ( ! [X139] : (sP4('T_0','F_0') | p(s(bool,X139)) | ~p(s(bool,'T_0'))) ) | $spl146),
  inference(resolution,[],[f6663,f295])).
fof(f295,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.IMP_CLAUSES|split|2')).
fof(f6744,plain,(
  ( ! [X140] : (sP4('T_0','F_0') | p(s(bool,X140))) ) | $spl146),
  inference(resolution,[],[f6663,f309])).
fof(f6745,plain,(
  ( ! [X141] : (sP4('T_0','F_0') | p(s(bool,X141))) ) | $spl146),
  inference(resolution,[],[f6663,f316])).
fof(f316,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) & p(s(bool,X11))) <=> p(s(bool,'F_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.AND_CLAUSES|split|2')).
fof(f6746,plain,(
  ( ! [X142] : (sP4('T_0','F_0') | p(s(bool,X142))) ) | $spl146),
  inference(resolution,[],[f6663,f333])).
fof(f15410,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('T_0','F_0')) ) | $spl146),
  inference(subsumption_resolution,[],[f6767,f324])).
fof(f6767,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('T_0','F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl146),
  inference(superposition,[],[f6663,f358])).
fof(f6783,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('T_0','F_0')) ) | $spl146),
  inference(duplicate_literal_removal,[],[f6768])).
fof(f6768,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('T_0','F_0') | p(s(bool,X30))) ) | $spl146),
  inference(superposition,[],[f6663,f323])).
fof(f6782,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('T_0','F_0')) ) | $spl146),
  inference(duplicate_literal_removal,[],[f6769])).
fof(f6769,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('T_0','F_0') | p(s(bool,X31))) ) | $spl146),
  inference(superposition,[],[f6663,f321])).
fof(f3618,plain,(
  sP4('F_0','T_0') | $spl106),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  sP4('F_0','T_0') <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f8113,plain,(
  ( ! [X140,X141] : (sP4('F_0','T_0') | sP5(X140,X141)) ) | $spl204),
  inference(resolution,[],[f7931,f274])).
fof(f8114,plain,(
  ( ! [X142,X143] : (sP4('F_0','T_0') | sP5(X142,X143)) ) | $spl204),
  inference(resolution,[],[f7931,f279])).
fof(f8115,plain,(
  ( ! [X144] : (sP4('F_0','T_0') | ~p(s(bool,X144))) ) | $spl204),
  inference(resolution,[],[f7931,f291])).
fof(f15408,plain,(
  ( ! [X145] : (sP4('F_0','T_0') | p(s(bool,X145))) ) | $spl204),
  inference(subsumption_resolution,[],[f8116,f361])).
fof(f8116,plain,(
  ( ! [X145] : (sP4('F_0','T_0') | p(s(bool,X145)) | ~p(s(bool,'T_0'))) ) | $spl204),
  inference(resolution,[],[f7931,f295])).
fof(f8117,plain,(
  ( ! [X146] : (sP4('F_0','T_0') | p(s(bool,X146))) ) | $spl204),
  inference(resolution,[],[f7931,f309])).
fof(f8118,plain,(
  ( ! [X147] : (sP4('F_0','T_0') | p(s(bool,X147))) ) | $spl204),
  inference(resolution,[],[f7931,f316])).
fof(f8119,plain,(
  ( ! [X148] : (sP4('F_0','T_0') | p(s(bool,X148))) ) | $spl204),
  inference(resolution,[],[f7931,f333])).
fof(f15405,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('F_0','T_0')) ) | $spl204),
  inference(subsumption_resolution,[],[f8141,f324])).
fof(f8141,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('F_0','T_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl204),
  inference(superposition,[],[f7931,f358])).
fof(f8156,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('F_0','T_0')) ) | $spl204),
  inference(duplicate_literal_removal,[],[f8142])).
fof(f8142,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('F_0','T_0') | p(s(bool,X30))) ) | $spl204),
  inference(superposition,[],[f7931,f323])).
fof(f8155,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('F_0','T_0')) ) | $spl204),
  inference(duplicate_literal_removal,[],[f8143])).
fof(f8143,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('F_0','T_0') | p(s(bool,X31))) ) | $spl204),
  inference(superposition,[],[f7931,f321])).
fof(f15395,plain,(
  ( ! [X235] : (sP4('T_0',X235) | ~sP4(X235,'F_0')) ) | $spl408),
  inference(cnf_transformation,[],[f15395_D])).
fof(f15395_D,plain,(
  ( ! [X235] : (sP4('T_0',X235) | ~sP4(X235,'F_0')) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f13701,plain,(
  ( ! [X231,X229,X230] : (sP4('T_0',X229) | ~sP4(X229,'F_0') | p(s(bool,X230)) | p(s(bool,X231)) | ~sP5(X230,X231)) ) | $spl146),
  inference(resolution,[],[f6706,f275])).
fof(f6706,plain,(
  ( ! [X74,X75] : (~p(s(bool,X75)) | sP4('T_0',X74) | ~sP4(X74,X75)) ) | $spl146),
  inference(resolution,[],[f6663,f256])).
fof(f13703,plain,(
  ( ! [X237,X235,X236] : (sP4('T_0',X235) | ~sP4(X235,'F_0') | p(s(bool,X236)) | p(s(bool,X237)) | ~sP5(X236,X237)) ) | $spl146),
  inference(resolution,[],[f6706,f271])).
fof(f15394,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f15365,f7268])).
fof(f7268,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f364,f900])).
fof(f900,plain,(
  s(bool,'T_0') = s(bool,sK8) | $spl4),
  inference(resolution,[],[f369,f325])).
fof(f15365,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | ($spl3 | $spl4)),
  inference(superposition,[],[f13622,f7364])).
fof(f13622,plain,(
  ( ! [X72,X73] : (s('type.option.option'(X72),'const.option.NONE_0') != s('type.option.option'(X72),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(X72),'const.option.SOME_1'(s(X72,X73))),s('type.option.option'(X72),'const.option.NONE_0')))) ) | $spl4),
  inference(forward_demodulation,[],[f904,f900])).
fof(f904,plain,(
  ( ! [X72,X73] : (s('type.option.option'(X72),'const.option.NONE_0') != s('type.option.option'(X72),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(X72),'const.option.SOME_1'(s(X72,X73))),s('type.option.option'(X72),'const.option.NONE_0')))) ) | $spl4),
  inference(resolution,[],[f369,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s('type.option.option'(X0),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f208])).
fof(f15126,plain,(
  $false | ($spl4 | $spl75 | $spl230)),
  inference(subsumption_resolution,[],[f15125,f2237])).
fof(f15125,plain,(
  sP4(sK8,sK8) | ($spl4 | $spl230)),
  inference(subsumption_resolution,[],[f15107,f361])).
fof(f15107,plain,(
  ~p(s(bool,'T_0')) | sP4(sK8,sK8) | ($spl4 | $spl230)),
  inference(resolution,[],[f8433,f1312])).
fof(f1312,plain,(
  ( ! [X76,X77] : (~sP4(X77,X76) | ~p(s(bool,X77)) | sP4(X76,sK8)) ) | $spl4),
  inference(resolution,[],[f882,f256])).
fof(f882,plain,(
  ( ! [X55] : (p(s(bool,X55)) | sP4(X55,sK8)) ) | $spl4),
  inference(resolution,[],[f369,f260])).
fof(f15124,plain,(
  $false | ($spl4 | $spl75 | $spl230)),
  inference(subsumption_resolution,[],[f15123,f2237])).
fof(f15123,plain,(
  sP4(sK8,sK8) | ($spl4 | $spl230)),
  inference(subsumption_resolution,[],[f15105,f361])).
fof(f15105,plain,(
  ~p(s(bool,'T_0')) | sP4(sK8,sK8) | ($spl4 | $spl230)),
  inference(resolution,[],[f8433,f1192])).
fof(f1192,plain,(
  ( ! [X76,X77] : (~sP4(X77,X76) | ~p(s(bool,X77)) | sP4(sK8,X76)) ) | $spl4),
  inference(resolution,[],[f881,f256])).
fof(f881,plain,(
  ( ! [X54] : (p(s(bool,X54)) | sP4(sK8,X54)) ) | $spl4),
  inference(resolution,[],[f369,f257])).
fof(f15121,plain,(
  $false | ($spl4 | $spl230)),
  inference(subsumption_resolution,[],[f15120,f340])).
fof(f340,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0',aHL_TRUTH)).
fof(f15120,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl230)),
  inference(forward_demodulation,[],[f15119,f900])).
fof(f15119,plain,(
  ~p(s(bool,sK8)) | $spl230),
  inference(subsumption_resolution,[],[f15102,f361])).
fof(f15102,plain,(
  ~p(s(bool,sK8)) | ~p(s(bool,'T_0')) | $spl230),
  inference(resolution,[],[f8433,f263])).
fof(f15118,plain,(
  $false | ($spl4 | $spl230)),
  inference(subsumption_resolution,[],[f15117,f340])).
fof(f15117,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl230)),
  inference(forward_demodulation,[],[f15116,f900])).
fof(f15116,plain,(
  ~p(s(bool,sK8)) | $spl230),
  inference(subsumption_resolution,[],[f15101,f361])).
fof(f15101,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK8)) | $spl230),
  inference(resolution,[],[f8433,f256])).
fof(f15115,plain,(
  $false | ($spl4 | $spl230)),
  inference(subsumption_resolution,[],[f15099,f340])).
fof(f15099,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl230)),
  inference(resolution,[],[f8433,f880])).
fof(f880,plain,(
  ( ! [X53] : (~sP4(X53,sK8) | ~p(s(bool,X53))) ) | $spl4),
  inference(resolution,[],[f369,f256])).
fof(f15095,plain,(
  ( ! [X569,X570] : (sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(X569),'const.option.SOME_1'(s(X569,X570)))))) ) | $spl406),
  inference(cnf_transformation,[],[f15095_D])).
fof(f15095_D,plain,(
  ( ! [X569,X570] : (sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(X569),'const.option.SOME_1'(s(X569,X570)))))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f8433,plain,(
  sP4('T_0',sK8) | $spl230),
  inference(cnf_transformation,[],[f8433_D])).
fof(f8433_D,plain,(
  sP4('T_0',sK8) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f15043,plain,(
  ( ! [X646,X645] : (sP4('T_0',sK8) | sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(X645),'const.option.SOME_1'(s(X645,X646)))))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1280])).
fof(f1280,plain,(
  ( ! [X33,X32] : (sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(X32),'const.option.SOME_1'(s(X32,X33)))))) ) | $spl4),
  inference(subsumption_resolution,[],[f1252,f334])).
fof(f1252,plain,(
  ( ! [X33,X32] : (p(s(bool,'F_0')) | sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(X32),'const.option.SOME_1'(s(X32,X33)))))) ) | $spl4),
  inference(superposition,[],[f881,f216])).
fof(f13683,plain,(
  ( ! [X202,X203] : (~sP4(X202,X203) | sP4('T_0',X202) | sP4('T_0',X203)) ) | $spl146),
  inference(resolution,[],[f6706,f6663])).
fof(f15093,plain,(
  ( ! [X476,X478,X475,X477] : (sP4('T_0',X475) | sP0(X476,X477,X478) | sP4(X476,X475)) ) | $spl404),
  inference(cnf_transformation,[],[f15093_D])).
fof(f15093_D,plain,(
  ( ! [X476,X478,X475,X477] : (sP4('T_0',X475) | sP0(X476,X477,X478) | sP4(X476,X475)) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f15035,plain,(
  ( ! [X629,X631,X628,X630] : (sP4('T_0',sK8) | sP4('T_0',X628) | sP4(X629,X628) | sP0(X629,X630,X631)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f3373])).
fof(f3373,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(sK8,X84) | sP4(X85,X84) | sP0(X85,X86,X87)) ) | $spl4),
  inference(resolution,[],[f1194,f244])).
fof(f244,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f1194,plain,(
  ( ! [X80,X81] : (sP4(sK8,X80) | sP4(X81,X80) | p(s(bool,X81))) ) | $spl4),
  inference(resolution,[],[f881,f260])).
fof(f15092,plain,(
  ( ! [X472,X474,X471,X473] : (sP4('T_0',X471) | sP0(X473,X472,X474) | sP4(X472,X471)) ) | $spl402),
  inference(cnf_transformation,[],[f15092_D])).
fof(f15092_D,plain,(
  ( ! [X472,X474,X471,X473] : (sP4('T_0',X471) | sP0(X473,X472,X474) | sP4(X472,X471)) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f15034,plain,(
  ( ! [X625,X627,X624,X626] : (sP4('T_0',sK8) | sP4('T_0',X624) | sP4(X625,X624) | sP0(X626,X625,X627)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f3372])).
fof(f3372,plain,(
  ( ! [X80,X83,X81,X82] : (sP4(sK8,X80) | sP4(X81,X80) | sP0(X82,X81,X83)) ) | $spl4),
  inference(resolution,[],[f1194,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f161])).
fof(f15091,plain,(
  ( ! [X469,X467,X468,X470] : (sP4('T_0',X467) | sP0(X469,X470,X468) | sP4(X468,X467)) ) | $spl400),
  inference(cnf_transformation,[],[f15091_D])).
fof(f15091_D,plain,(
  ( ! [X469,X467,X468,X470] : (sP4('T_0',X467) | sP0(X469,X470,X468) | sP4(X468,X467)) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f15033,plain,(
  ( ! [X621,X623,X620,X622] : (sP4('T_0',sK8) | sP4('T_0',X620) | sP4(X621,X620) | sP0(X622,X623,X621)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f3371])).
fof(f3371,plain,(
  ( ! [X78,X76,X79,X77] : (sP4(sK8,X76) | sP4(X77,X76) | sP0(X78,X79,X77)) ) | $spl4),
  inference(resolution,[],[f1194,f242])).
fof(f242,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f161])).
fof(f15090,plain,(
  ( ! [X466,X463,X465,X464] : (sP4('T_0',X463) | sP1(X465,X466,X464) | sP4(X464,X463)) ) | $spl398),
  inference(cnf_transformation,[],[f15090_D])).
fof(f15090_D,plain,(
  ( ! [X466,X463,X465,X464] : (sP4('T_0',X463) | sP1(X465,X466,X464) | sP4(X464,X463)) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f15032,plain,(
  ( ! [X617,X619,X616,X618] : (sP4('T_0',sK8) | sP4('T_0',X616) | sP4(X617,X616) | sP1(X618,X619,X617)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f3370])).
fof(f3370,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(sK8,X72) | sP4(X73,X72) | sP1(X74,X75,X73)) ) | $spl4),
  inference(resolution,[],[f1194,f238])).
fof(f238,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f159])).
fof(f15089,plain,(
  ( ! [X461,X459,X460,X462] : (sP4('T_0',X459) | sP2(X461,X462,X460) | sP4(X460,X459)) ) | $spl396),
  inference(cnf_transformation,[],[f15089_D])).
fof(f15089_D,plain,(
  ( ! [X461,X459,X460,X462] : (sP4('T_0',X459) | sP2(X461,X462,X460) | sP4(X460,X459)) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f15031,plain,(
  ( ! [X613,X615,X612,X614] : (sP4('T_0',sK8) | sP4('T_0',X612) | sP4(X613,X612) | sP2(X614,X615,X613)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f3367])).
fof(f3367,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(sK8,X60) | sP4(X61,X60) | sP2(X62,X63,X61)) ) | $spl4),
  inference(resolution,[],[f1194,f234])).
fof(f234,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f156])).
fof(f15088,plain,(
  ( ! [X455,X457,X456,X458] : (sP4('T_0',X455) | sP0(X456,X457,X458) | sP4(X455,X456)) ) | $spl394),
  inference(cnf_transformation,[],[f15088_D])).
fof(f15088_D,plain,(
  ( ! [X455,X457,X456,X458] : (sP4('T_0',X455) | sP0(X456,X457,X458) | sP4(X455,X456)) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f15026,plain,(
  ( ! [X600,X602,X601,X599] : (sP4('T_0',sK8) | sP4('T_0',X599) | sP4(X599,X600) | sP0(X600,X601,X602)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f2118])).
fof(f2118,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(sK8,X84) | sP4(X84,X85) | sP0(X85,X86,X87)) ) | $spl4),
  inference(resolution,[],[f1193,f244])).
fof(f1193,plain,(
  ( ! [X78,X79] : (sP4(sK8,X78) | sP4(X78,X79) | p(s(bool,X79))) ) | $spl4),
  inference(resolution,[],[f881,f257])).
fof(f15087,plain,(
  ( ! [X453,X451,X452,X454] : (sP4('T_0',X451) | sP0(X453,X452,X454) | sP4(X451,X452)) ) | $spl392),
  inference(cnf_transformation,[],[f15087_D])).
fof(f15087_D,plain,(
  ( ! [X453,X451,X452,X454] : (sP4('T_0',X451) | sP0(X453,X452,X454) | sP4(X451,X452)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f15025,plain,(
  ( ! [X596,X598,X597,X595] : (sP4('T_0',sK8) | sP4('T_0',X595) | sP4(X595,X596) | sP0(X597,X596,X598)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f2117])).
fof(f2117,plain,(
  ( ! [X80,X83,X81,X82] : (sP4(sK8,X80) | sP4(X80,X81) | sP0(X82,X81,X83)) ) | $spl4),
  inference(resolution,[],[f1193,f243])).
fof(f15086,plain,(
  ( ! [X447,X449,X448,X450] : (sP4('T_0',X447) | sP0(X449,X450,X448) | sP4(X447,X448)) ) | $spl390),
  inference(cnf_transformation,[],[f15086_D])).
fof(f15086_D,plain,(
  ( ! [X447,X449,X448,X450] : (sP4('T_0',X447) | sP0(X449,X450,X448) | sP4(X447,X448)) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f15024,plain,(
  ( ! [X592,X594,X593,X591] : (sP4('T_0',sK8) | sP4('T_0',X591) | sP4(X591,X592) | sP0(X593,X594,X592)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f2116])).
fof(f2116,plain,(
  ( ! [X78,X76,X79,X77] : (sP4(sK8,X76) | sP4(X76,X77) | sP0(X78,X79,X77)) ) | $spl4),
  inference(resolution,[],[f1193,f242])).
fof(f15085,plain,(
  ( ! [X445,X443,X444,X446] : (sP4('T_0',X443) | sP1(X445,X446,X444) | sP4(X443,X444)) ) | $spl388),
  inference(cnf_transformation,[],[f15085_D])).
fof(f15085_D,plain,(
  ( ! [X445,X443,X444,X446] : (sP4('T_0',X443) | sP1(X445,X446,X444) | sP4(X443,X444)) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f15023,plain,(
  ( ! [X588,X590,X589,X587] : (sP4('T_0',sK8) | sP4('T_0',X587) | sP4(X587,X588) | sP1(X589,X590,X588)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f2115])).
fof(f2115,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(sK8,X72) | sP4(X72,X73) | sP1(X74,X75,X73)) ) | $spl4),
  inference(resolution,[],[f1193,f238])).
fof(f15084,plain,(
  ( ! [X441,X440,X442,X439] : (sP4('T_0',X439) | sP2(X441,X442,X440) | sP4(X439,X440)) ) | $spl386),
  inference(cnf_transformation,[],[f15084_D])).
fof(f15084_D,plain,(
  ( ! [X441,X440,X442,X439] : (sP4('T_0',X439) | sP2(X441,X442,X440) | sP4(X439,X440)) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f15022,plain,(
  ( ! [X584,X586,X585,X583] : (sP4('T_0',sK8) | sP4('T_0',X583) | sP4(X583,X584) | sP2(X585,X586,X584)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f2112])).
fof(f2112,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(sK8,X60) | sP4(X60,X61) | sP2(X62,X63,X61)) ) | $spl4),
  inference(resolution,[],[f1193,f234])).
fof(f15083,plain,(
  ( ! [X416] : (sP4('T_0',X416) | s(bool,'T_0') = s(bool,X416)) ) | $spl384),
  inference(cnf_transformation,[],[f15083_D])).
fof(f15083_D,plain,(
  ( ! [X416] : (sP4('T_0',X416) | s(bool,'T_0') = s(bool,X416)) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f15021,plain,(
  ( ! [X582] : (sP4('T_0',sK8) | sP4('T_0',X582) | s(bool,'T_0') = s(bool,X582)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1212])).
fof(f1212,plain,(
  ( ! [X109] : (sP4(sK8,X109) | s(bool,'T_0') = s(bool,X109)) ) | $spl4),
  inference(resolution,[],[f881,f325])).
fof(f15082,plain,(
  ( ! [X415] : (sP4('T_0',X415) | s(bool,'F_0') != s(bool,X415)) ) | $spl382),
  inference(cnf_transformation,[],[f15082_D])).
fof(f15082_D,plain,(
  ( ! [X415] : (sP4('T_0',X415) | s(bool,'F_0') != s(bool,X415)) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f15020,plain,(
  ( ! [X581] : (sP4('T_0',sK8) | sP4('T_0',X581) | s(bool,'F_0') != s(bool,X581)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1210])).
fof(f1210,plain,(
  ( ! [X107] : (sP4(sK8,X107) | s(bool,'F_0') != s(bool,X107)) ) | $spl4),
  inference(resolution,[],[f881,f320])).
fof(f15019,plain,(
  ( ! [X580,X579] : (sP4('T_0',sK8) | sP4('T_0',X579) | sP4(X580,X579) | p(s(bool,X580))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1194])).
fof(f15018,plain,(
  ( ! [X578,X577] : (sP4('T_0',sK8) | sP4('T_0',X577) | sP4(X577,X578) | p(s(bool,X578))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1193])).
fof(f15012,plain,(
  ( ! [X569,X570] : (sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(X569),'const.option.SOME_1'(s(X569,X570))))) | sP4('T_0',sK8)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1399])).
fof(f1399,plain,(
  ( ! [X33,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X32),'const.option.SOME_1'(s(X32,X33)))),sK8)) ) | $spl4),
  inference(subsumption_resolution,[],[f1372,f334])).
fof(f1372,plain,(
  ( ! [X33,X32] : (p(s(bool,'F_0')) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X32),'const.option.SOME_1'(s(X32,X33)))),sK8)) ) | $spl4),
  inference(superposition,[],[f882,f216])).
fof(f14966,plain,(
  ( ! [X476,X478,X475,X477] : (sP4('T_0',X475) | sP4('T_0',sK8) | sP4(X476,X475) | sP0(X476,X477,X478)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4883])).
fof(f4883,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(X84,sK8) | sP4(X85,X84) | sP0(X85,X86,X87)) ) | $spl4),
  inference(resolution,[],[f1314,f244])).
fof(f1314,plain,(
  ( ! [X80,X81] : (sP4(X80,sK8) | sP4(X81,X80) | p(s(bool,X81))) ) | $spl4),
  inference(resolution,[],[f882,f260])).
fof(f14965,plain,(
  ( ! [X472,X474,X471,X473] : (sP4('T_0',X471) | sP4('T_0',sK8) | sP4(X472,X471) | sP0(X473,X472,X474)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4882])).
fof(f4882,plain,(
  ( ! [X80,X83,X81,X82] : (sP4(X80,sK8) | sP4(X81,X80) | sP0(X82,X81,X83)) ) | $spl4),
  inference(resolution,[],[f1314,f243])).
fof(f14964,plain,(
  ( ! [X469,X467,X468,X470] : (sP4('T_0',X467) | sP4('T_0',sK8) | sP4(X468,X467) | sP0(X469,X470,X468)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4881])).
fof(f4881,plain,(
  ( ! [X78,X76,X79,X77] : (sP4(X76,sK8) | sP4(X77,X76) | sP0(X78,X79,X77)) ) | $spl4),
  inference(resolution,[],[f1314,f242])).
fof(f14963,plain,(
  ( ! [X466,X463,X465,X464] : (sP4('T_0',X463) | sP4('T_0',sK8) | sP4(X464,X463) | sP1(X465,X466,X464)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4880])).
fof(f4880,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,sK8) | sP4(X73,X72) | sP1(X74,X75,X73)) ) | $spl4),
  inference(resolution,[],[f1314,f238])).
fof(f14962,plain,(
  ( ! [X461,X459,X460,X462] : (sP4('T_0',X459) | sP4('T_0',sK8) | sP4(X460,X459) | sP2(X461,X462,X460)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4877])).
fof(f4877,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,sK8) | sP4(X61,X60) | sP2(X62,X63,X61)) ) | $spl4),
  inference(resolution,[],[f1314,f234])).
fof(f14961,plain,(
  ( ! [X455,X457,X456,X458] : (sP4('T_0',X455) | sP4('T_0',sK8) | sP4(X455,X456) | sP0(X456,X457,X458)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4735])).
fof(f4735,plain,(
  ( ! [X87,X85,X86,X84] : (sP4(X84,sK8) | sP4(X84,X85) | sP0(X85,X86,X87)) ) | $spl4),
  inference(resolution,[],[f1313,f244])).
fof(f1313,plain,(
  ( ! [X78,X79] : (sP4(X78,sK8) | sP4(X78,X79) | p(s(bool,X79))) ) | $spl4),
  inference(resolution,[],[f882,f257])).
fof(f14960,plain,(
  ( ! [X453,X451,X452,X454] : (sP4('T_0',X451) | sP4('T_0',sK8) | sP4(X451,X452) | sP0(X453,X452,X454)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4734])).
fof(f4734,plain,(
  ( ! [X80,X83,X81,X82] : (sP4(X80,sK8) | sP4(X80,X81) | sP0(X82,X81,X83)) ) | $spl4),
  inference(resolution,[],[f1313,f243])).
fof(f14959,plain,(
  ( ! [X447,X449,X448,X450] : (sP4('T_0',X447) | sP4('T_0',sK8) | sP4(X447,X448) | sP0(X449,X450,X448)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4733])).
fof(f4733,plain,(
  ( ! [X78,X76,X79,X77] : (sP4(X76,sK8) | sP4(X76,X77) | sP0(X78,X79,X77)) ) | $spl4),
  inference(resolution,[],[f1313,f242])).
fof(f14958,plain,(
  ( ! [X445,X443,X444,X446] : (sP4('T_0',X443) | sP4('T_0',sK8) | sP4(X443,X444) | sP1(X445,X446,X444)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4732])).
fof(f4732,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,sK8) | sP4(X72,X73) | sP1(X74,X75,X73)) ) | $spl4),
  inference(resolution,[],[f1313,f238])).
fof(f14957,plain,(
  ( ! [X441,X440,X442,X439] : (sP4('T_0',X439) | sP4('T_0',sK8) | sP4(X439,X440) | sP2(X441,X442,X440)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f4729])).
fof(f4729,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,sK8) | sP4(X60,X61) | sP2(X62,X63,X61)) ) | $spl4),
  inference(resolution,[],[f1313,f234])).
fof(f14946,plain,(
  ( ! [X416] : (sP4('T_0',X416) | sP4('T_0',sK8) | s(bool,'T_0') = s(bool,X416)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1332])).
fof(f1332,plain,(
  ( ! [X109] : (sP4(X109,sK8) | s(bool,'T_0') = s(bool,X109)) ) | $spl4),
  inference(resolution,[],[f882,f325])).
fof(f14945,plain,(
  ( ! [X415] : (sP4('T_0',X415) | sP4('T_0',sK8) | s(bool,'F_0') != s(bool,X415)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1330])).
fof(f1330,plain,(
  ( ! [X107] : (sP4(X107,sK8) | s(bool,'F_0') != s(bool,X107)) ) | $spl4),
  inference(resolution,[],[f882,f320])).
fof(f14944,plain,(
  ( ! [X414,X413] : (sP4('T_0',X413) | sP4('T_0',sK8) | sP4(X414,X413) | p(s(bool,X414))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1314])).
fof(f14943,plain,(
  ( ! [X412,X411] : (sP4('T_0',X411) | sP4('T_0',sK8) | sP4(X411,X412) | p(s(bool,X412))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f13683,f1313])).
fof(f13563,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl377),
  inference(cnf_transformation,[],[f13563_D])).
fof(f13563_D,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) <=> ~$spl377),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl377])])).
fof(f14023,plain,(
  ( ! [X24] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X24))) ) | $spl3),
  inference(subsumption_resolution,[],[f13980,f320])).
fof(f13980,plain,(
  ( ! [X24] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X24)) | s(bool,'F_0') = s(bool,X24)) ) | $spl3),
  inference(superposition,[],[f13935,f358])).
fof(f13935,plain,(
  ( ! [X12] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) | ~p(s(bool,X12))) ) | $spl3),
  inference(forward_demodulation,[],[f7350,f7311])).
fof(f7311,plain,(
  s(bool,'F_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | $spl3),
  inference(resolution,[],[f367,f321])).
fof(f7350,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))))) ) | $spl3),
  inference(superposition,[],[f367,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X12,X13] : (s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X13))) <=> s(X6,X12) = s(X6,X13))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.option.SOME_|49||49|')).
fof(f14017,plain,(
  ( ! [X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X23))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f13979])).
fof(f13979,plain,(
  ( ! [X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X23)) | ~p(s(bool,X23))) ) | $spl3),
  inference(superposition,[],[f13935,f327])).
fof(f327,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f14018,plain,(
  ( ! [X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X22))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f13978])).
fof(f13978,plain,(
  ( ! [X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ~p(s(bool,X22)) | ~p(s(bool,X22))) ) | $spl3),
  inference(superposition,[],[f13935,f325])).
fof(f13746,plain,(
  ( ! [X207] : (~sP4(X207,'T_0') | sP4('T_0',X207)) ) | $spl380),
  inference(cnf_transformation,[],[f13746_D])).
fof(f13746_D,plain,(
  ( ! [X207] : (~sP4(X207,'T_0') | sP4('T_0',X207)) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f13736,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP4('T_0',X45) | ~sP4(X45,'T_0')) ) | $spl146),
  inference(duplicate_literal_removal,[],[f13720])).
fof(f13720,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP4('T_0',X45) | ~sP4(X45,'T_0') | ~p(s(bool,X44))) ) | $spl146),
  inference(superposition,[],[f6706,f325])).
fof(f13737,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP4('T_0',X43) | ~sP4(X43,'T_0')) ) | $spl146),
  inference(duplicate_literal_removal,[],[f13719])).
fof(f13719,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP4('T_0',X43) | ~sP4(X43,'T_0') | ~p(s(bool,X42))) ) | $spl146),
  inference(superposition,[],[f6706,f327])).
fof(f13754,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP4('T_0',X41) | ~sP4(X41,'T_0')) ) | $spl146),
  inference(subsumption_resolution,[],[f13718,f320])).
fof(f13718,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP4('T_0',X41) | ~sP4(X41,'T_0') | s(bool,'F_0') = s(bool,X40)) ) | $spl146),
  inference(superposition,[],[f6706,f358])).
fof(f8353,plain,(
  ( ! [X1] : (sP4('T_0',X1)) ) | $spl220),
  inference(cnf_transformation,[],[f8353_D])).
fof(f8353_D,plain,(
  ( ! [X1] : (sP4('T_0',X1)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f13750,plain,(
  ( ! [X237,X235,X236] : (sP4('T_0',X235) | p(s(bool,X236)) | p(s(bool,X237)) | ~sP5(X236,X237)) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f13703,f6734])).
fof(f6734,plain,(
  ( ! [X128] : (sP4('T_0',X128) | sP4(X128,'F_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f6663,f2221])).
fof(f2221,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP4(X0,'F_0')) ) | $spl4),
  inference(resolution,[],[f2159,f879])).
fof(f879,plain,(
  ( ! [X52] : (~sP4(sK8,X52) | ~p(s(bool,X52))) ) | $spl4),
  inference(resolution,[],[f369,f256])).
fof(f2159,plain,(
  ( ! [X198] : (sP4(sK8,X198) | sP4(X198,'F_0')) ) | $spl4),
  inference(resolution,[],[f1193,f334])).
fof(f13749,plain,(
  ( ! [X231,X229,X230] : (sP4('T_0',X229) | p(s(bool,X230)) | p(s(bool,X231)) | ~sP5(X230,X231)) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f13701,f6734])).
fof(f13692,plain,(
  ( ! [X220,X219] : (sP4('T_0',X219) | ~sP4(X219,'T_0') | ~p(s(bool,X220))) ) | $spl146),
  inference(resolution,[],[f6706,f318])).
fof(f318,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) & p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.AND_CLAUSES|split|0')).
fof(f13691,plain,(
  ( ! [X218,X217] : (sP4('T_0',X217) | ~sP4(X217,'T_0') | ~p(s(bool,X218))) ) | $spl146),
  inference(resolution,[],[f6706,f313])).
fof(f313,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'T_0'))) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.AND_CLAUSES|split|1')).
fof(f13690,plain,(
  ( ! [X216,X215] : (sP4('T_0',X215) | ~sP4(X215,'T_0') | ~p(s(bool,X216))) ) | $spl146),
  inference(resolution,[],[f6706,f303])).
fof(f303,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.IMP_CLAUSES|split|3')).
fof(f13689,plain,(
  ( ! [X213,X214] : (sP4('T_0',X213) | ~sP4(X213,'T_0') | p(s(bool,X214))) ) | $spl146),
  inference(resolution,[],[f6706,f302])).
fof(f302,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f13688,plain,(
  ( ! [X212,X211] : (sP4('T_0',X211) | ~sP4(X211,'T_0') | p(s(bool,X212))) ) | $spl146),
  inference(resolution,[],[f6706,f299])).
fof(f299,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) => p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.IMP_CLAUSES|split|0')).
fof(f13687,plain,(
  ( ! [X210,X209] : (sP4('T_0',X209) | ~sP4(X209,'T_0') | p(s(bool,X210))) ) | $spl146),
  inference(resolution,[],[f6706,f296])).
fof(f296,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.IMP_CLAUSES|split|1')).
fof(f13686,plain,(
  ( ! [X208,X207] : (sP4('T_0',X207) | ~sP4(X207,'T_0') | ~p(s(bool,X208))) ) | $spl146),
  inference(resolution,[],[f6706,f294])).
fof(f294,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f13740,plain,(
  ( ! [X0,X3,X1] : (sP4('T_0',X0) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | ~sP4(X0,X1)) ) | $spl378),
  inference(cnf_transformation,[],[f13740_D])).
fof(f13740_D,plain,(
  ( ! [X0,X3,X1] : (sP4('T_0',X0) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | ~sP4(X0,X1)) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f13676,plain,(
  ( ! [X185,X182,X184,X183] : (sP4('T_0',X182) | ~sP4(X182,X183) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl146),
  inference(resolution,[],[f6706,f360])).
fof(f13631,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4('T_0',X35) | ~sP4(X35,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl146),
  inference(resolution,[],[f6706,f225])).
fof(f13630,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4('T_0',X30) | ~sP4(X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl146),
  inference(resolution,[],[f6706,f224])).
fof(f13628,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4('T_0',X20) | ~sP4(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl146),
  inference(resolution,[],[f6706,f223])).
fof(f13626,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4('T_0',X10) | ~sP4(X10,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl146),
  inference(resolution,[],[f6706,f221])).
fof(f13624,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4('T_0',X0) | ~sP4(X0,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl146),
  inference(resolution,[],[f6706,f219])).
fof(f13573,plain,(
  $false | ($spl0 | $spl4 | $spl31)),
  inference(subsumption_resolution,[],[f1100,f7268])).
fof(f13564,plain,(
  ( ! [X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X25))) ) | $spl4),
  inference(subsumption_resolution,[],[f13518,f324])).
fof(f13518,plain,(
  ( ! [X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X25)) | s(bool,'T_0') = s(bool,X25)) ) | $spl4),
  inference(superposition,[],[f13474,f358])).
fof(f13474,plain,(
  ( ! [X12] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) | p(s(bool,X12))) ) | $spl4),
  inference(forward_demodulation,[],[f921,f900])).
fof(f921,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK8))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) | $spl4),
  inference(superposition,[],[f369,f351])).
fof(f13556,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X21))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f13514])).
fof(f13514,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X21)) | p(s(bool,X21))) ) | $spl4),
  inference(superposition,[],[f13474,f323])).
fof(f13557,plain,(
  ( ! [X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X20))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f13513])).
fof(f13513,plain,(
  ( ! [X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,X20)) | p(s(bool,X20))) ) | $spl4),
  inference(superposition,[],[f13474,f321])).
fof(f3611,plain,(
  ( ! [X20] : (sP4('F_0',X20)) ) | $spl104),
  inference(cnf_transformation,[],[f3611_D])).
fof(f3611_D,plain,(
  ( ! [X20] : (sP4('F_0',X20)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f13416,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('F_0',X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f13405,f3519])).
fof(f3519,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4('F_0',X76)) ) | $spl4),
  inference(resolution,[],[f3479,f255])).
fof(f3479,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP4('F_0',X0)) ) | $spl4),
  inference(resolution,[],[f3415,f879])).
fof(f3415,plain,(
  ( ! [X200] : (sP4(sK8,X200) | sP4('F_0',X200)) ) | $spl4),
  inference(resolution,[],[f1194,f334])).
fof(f13405,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4('F_0',X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl4),
  inference(resolution,[],[f10894,f231])).
fof(f10894,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4('F_0',X72) | sP1(X73,X74,X75)) ) | $spl4),
  inference(resolution,[],[f3521,f240])).
fof(f3521,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4('F_0',X80)) ) | $spl4),
  inference(resolution,[],[f3479,f260])).
fof(f13400,plain,(
  ( ! [X47,X45,X46,X44] : (sP4(X44,X45) | sP1(X45,X46,X47) | p(s(bool,X44)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10894,f3519])).
fof(f13398,plain,(
  ( ! [X39,X37,X38,X36] : (sP4(X36,X37) | sP1(X37,X38,X39) | p(s(bool,X36)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10894,f2966])).
fof(f2966,plain,(
  ( ! [X76,X77] : (~sP4(X76,X77) | p(s(bool,X77)) | sP4(X76,'F_0')) ) | $spl4),
  inference(resolution,[],[f2221,f255])).
fof(f13355,plain,(
  ( ! [X47,X45,X46,X44] : (sP4(X44,X45) | sP1(X46,X45,X47) | p(s(bool,X44)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10893,f3519])).
fof(f10893,plain,(
  ( ! [X70,X68,X71,X69] : (sP4(X68,X69) | sP4('F_0',X68) | sP1(X70,X69,X71)) ) | $spl4),
  inference(resolution,[],[f3521,f239])).
fof(f13353,plain,(
  ( ! [X39,X37,X38,X36] : (sP4(X36,X37) | sP1(X38,X37,X39) | p(s(bool,X36)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10893,f2966])).
fof(f13326,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('F_0',X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl4),
  inference(subsumption_resolution,[],[f13315,f3519])).
fof(f13315,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4('F_0',X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl4),
  inference(resolution,[],[f10891,f232])).
fof(f10891,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4('F_0',X60) | sP2(X61,X62,X63)) ) | $spl4),
  inference(resolution,[],[f3521,f236])).
fof(f13309,plain,(
  ( ! [X47,X45,X46,X44] : (sP4(X44,X45) | sP2(X45,X46,X47) | p(s(bool,X44)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10891,f3519])).
fof(f13307,plain,(
  ( ! [X39,X37,X38,X36] : (sP4(X36,X37) | sP2(X37,X38,X39) | p(s(bool,X36)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10891,f2966])).
fof(f13265,plain,(
  ( ! [X47,X45,X46,X44] : (sP4(X44,X45) | sP2(X46,X45,X47) | p(s(bool,X44)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10890,f3519])).
fof(f10890,plain,(
  ( ! [X59,X57,X58,X56] : (sP4(X56,X57) | sP4('F_0',X56) | sP2(X58,X57,X59)) ) | $spl4),
  inference(resolution,[],[f3521,f235])).
fof(f13263,plain,(
  ( ! [X39,X37,X38,X36] : (sP4(X36,X37) | sP2(X38,X37,X39) | p(s(bool,X36)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10890,f2966])).
fof(f13182,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('F_0',X9) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f13171,f3518])).
fof(f3518,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4('F_0',X74)) ) | $spl4),
  inference(resolution,[],[f3479,f255])).
fof(f13171,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4('F_0',X9) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl4),
  inference(resolution,[],[f10756,f231])).
fof(f10756,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4('F_0',X73) | sP1(X72,X74,X75)) ) | $spl4),
  inference(resolution,[],[f3520,f240])).
fof(f3520,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4('F_0',X78)) ) | $spl4),
  inference(resolution,[],[f3479,f257])).
fof(f13166,plain,(
  ( ! [X47,X45,X46,X44] : (sP4(X44,X45) | sP1(X44,X46,X47) | p(s(bool,X45)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10756,f3519])).
fof(f13164,plain,(
  ( ! [X39,X37,X38,X36] : (sP4(X36,X37) | sP1(X36,X38,X39) | p(s(bool,X37)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10756,f2966])).
fof(f13121,plain,(
  ( ! [X47,X45,X46,X44] : (sP4(X44,X45) | sP1(X46,X44,X47) | p(s(bool,X45)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10755,f3519])).
fof(f10755,plain,(
  ( ! [X70,X68,X71,X69] : (sP4(X68,X69) | sP4('F_0',X69) | sP1(X70,X68,X71)) ) | $spl4),
  inference(resolution,[],[f3520,f239])).
fof(f13119,plain,(
  ( ! [X39,X37,X38,X36] : (sP4(X36,X37) | sP1(X38,X36,X39) | p(s(bool,X37)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10755,f2966])).
fof(f13092,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('F_0',X13) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl4),
  inference(subsumption_resolution,[],[f13081,f3518])).
fof(f13081,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4('F_0',X13) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl4),
  inference(resolution,[],[f10753,f232])).
fof(f10753,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4('F_0',X61) | sP2(X60,X62,X63)) ) | $spl4),
  inference(resolution,[],[f3520,f236])).
fof(f13075,plain,(
  ( ! [X47,X45,X46,X44] : (sP4(X44,X45) | sP2(X44,X46,X47) | p(s(bool,X45)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10753,f3519])).
fof(f13073,plain,(
  ( ! [X39,X37,X38,X36] : (sP4(X36,X37) | sP2(X36,X38,X39) | p(s(bool,X37)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10753,f2966])).
fof(f13031,plain,(
  ( ! [X47,X45,X46,X44] : (sP4(X44,X45) | sP2(X46,X44,X47) | p(s(bool,X45)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10752,f3519])).
fof(f10752,plain,(
  ( ! [X59,X57,X58,X56] : (sP4(X56,X57) | sP4('F_0',X57) | sP2(X58,X56,X59)) ) | $spl4),
  inference(resolution,[],[f3520,f235])).
fof(f13029,plain,(
  ( ! [X39,X37,X38,X36] : (sP4(X36,X37) | sP2(X38,X36,X39) | p(s(bool,X37)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f10752,f2966])).
fof(f12936,plain,(
  ( ! [X2,X1] : (sP1('F_0',X1,X2)) ) | $spl374),
  inference(cnf_transformation,[],[f12936_D])).
fof(f12936_D,plain,(
  ( ! [X2,X1] : (sP1('F_0',X1,X2)) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f12927,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'F_0') | sP1('F_0',X1,X2)) ) | $spl4),
  inference(factoring,[],[f9683])).
fof(f9683,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(X72,'F_0') | sP1(X73,X74,X75)) ) | $spl4),
  inference(resolution,[],[f2968,f240])).
fof(f2968,plain,(
  ( ! [X80,X81] : (~p(s(bool,X81)) | sP4(X80,X81) | sP4(X80,'F_0')) ) | $spl4),
  inference(resolution,[],[f2221,f260])).
fof(f12934,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f12922,f2966])).
fof(f12922,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(X8,'F_0') | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl4),
  inference(resolution,[],[f9683,f231])).
fof(f12916,plain,(
  ( ! [X43,X41,X42,X40] : (sP4(X40,X41) | sP1(X41,X42,X43) | p(s(bool,X40)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9683,f3518])).
fof(f12914,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP1(X33,X34,X35) | p(s(bool,X32)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9683,f2965])).
fof(f2965,plain,(
  ( ! [X74,X75] : (~sP4(X75,X74) | p(s(bool,X75)) | sP4(X74,'F_0')) ) | $spl4),
  inference(resolution,[],[f2221,f255])).
fof(f12836,plain,(
  ( ! [X2,X1] : (sP1(X1,'F_0',X2)) ) | $spl372),
  inference(cnf_transformation,[],[f12836_D])).
fof(f12836_D,plain,(
  ( ! [X2,X1] : (sP1(X1,'F_0',X2)) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f12828,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'F_0') | sP1(X1,'F_0',X2)) ) | $spl4),
  inference(factoring,[],[f9682])).
fof(f9682,plain,(
  ( ! [X70,X68,X71,X69] : (sP4(X68,X69) | sP4(X68,'F_0') | sP1(X70,X69,X71)) ) | $spl4),
  inference(resolution,[],[f2968,f239])).
fof(f12817,plain,(
  ( ! [X43,X41,X42,X40] : (sP4(X40,X41) | sP1(X42,X41,X43) | p(s(bool,X40)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9682,f3518])).
fof(f12815,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP1(X34,X33,X35) | p(s(bool,X32)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9682,f2965])).
fof(f12779,plain,(
  ( ! [X2,X1] : (sP2('F_0',X1,X2)) ) | $spl370),
  inference(cnf_transformation,[],[f12779_D])).
fof(f12779_D,plain,(
  ( ! [X2,X1] : (sP2('F_0',X1,X2)) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f12770,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'F_0') | sP2('F_0',X1,X2)) ) | $spl4),
  inference(factoring,[],[f9680])).
fof(f9680,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(X60,'F_0') | sP2(X61,X62,X63)) ) | $spl4),
  inference(resolution,[],[f2968,f236])).
fof(f12778,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'F_0') | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl4),
  inference(subsumption_resolution,[],[f12766,f2966])).
fof(f12766,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(X12,'F_0') | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl4),
  inference(resolution,[],[f9680,f232])).
fof(f12759,plain,(
  ( ! [X43,X41,X42,X40] : (sP4(X40,X41) | sP2(X41,X42,X43) | p(s(bool,X40)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9680,f3518])).
fof(f12757,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP2(X33,X34,X35) | p(s(bool,X32)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9680,f2965])).
fof(f12680,plain,(
  ( ! [X2,X1] : (sP2(X1,'F_0',X2)) ) | $spl368),
  inference(cnf_transformation,[],[f12680_D])).
fof(f12680_D,plain,(
  ( ! [X2,X1] : (sP2(X1,'F_0',X2)) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f12673,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'F_0') | sP2(X1,'F_0',X2)) ) | $spl4),
  inference(factoring,[],[f9679])).
fof(f9679,plain,(
  ( ! [X59,X57,X58,X56] : (sP4(X56,X57) | sP4(X56,'F_0') | sP2(X58,X57,X59)) ) | $spl4),
  inference(resolution,[],[f2968,f235])).
fof(f12662,plain,(
  ( ! [X43,X41,X42,X40] : (sP4(X40,X41) | sP2(X42,X41,X43) | p(s(bool,X40)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9679,f3518])).
fof(f12660,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP2(X34,X33,X35) | p(s(bool,X32)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9679,f2965])).
fof(f12602,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X9,'F_0') | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f12590,f2965])).
fof(f12590,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,X9) | sP4(X9,'F_0') | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) ) | $spl4),
  inference(resolution,[],[f9458,f231])).
fof(f9458,plain,(
  ( ! [X74,X72,X75,X73] : (sP4(X72,X73) | sP4(X73,'F_0') | sP1(X72,X74,X75)) ) | $spl4),
  inference(resolution,[],[f2967,f240])).
fof(f2967,plain,(
  ( ! [X78,X79] : (~p(s(bool,X79)) | sP4(X79,X78) | sP4(X78,'F_0')) ) | $spl4),
  inference(resolution,[],[f2221,f257])).
fof(f12584,plain,(
  ( ! [X43,X41,X42,X40] : (sP4(X40,X41) | sP1(X40,X42,X43) | p(s(bool,X41)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9458,f3518])).
fof(f12582,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP1(X32,X34,X35) | p(s(bool,X33)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9458,f2965])).
fof(f12538,plain,(
  ( ! [X43,X41,X42,X40] : (sP4(X40,X41) | sP1(X42,X40,X43) | p(s(bool,X41)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9457,f3518])).
fof(f9457,plain,(
  ( ! [X70,X68,X71,X69] : (sP4(X68,X69) | sP4(X69,'F_0') | sP1(X70,X68,X71)) ) | $spl4),
  inference(resolution,[],[f2967,f239])).
fof(f12536,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP1(X34,X32,X35) | p(s(bool,X33)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9457,f2965])).
fof(f12509,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X13,'F_0') | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl4),
  inference(subsumption_resolution,[],[f12497,f2965])).
fof(f12497,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,X13) | sP4(X13,'F_0') | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) ) | $spl4),
  inference(resolution,[],[f9455,f232])).
fof(f9455,plain,(
  ( ! [X61,X62,X60,X63] : (sP4(X60,X61) | sP4(X61,'F_0') | sP2(X60,X62,X63)) ) | $spl4),
  inference(resolution,[],[f2967,f236])).
fof(f12490,plain,(
  ( ! [X43,X41,X42,X40] : (sP4(X40,X41) | sP2(X40,X42,X43) | p(s(bool,X41)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9455,f3518])).
fof(f12488,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP2(X32,X34,X35) | p(s(bool,X33)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9455,f2965])).
fof(f12403,plain,(
  ( ! [X43,X41,X42,X40] : (sP4(X40,X41) | sP2(X42,X40,X43) | p(s(bool,X41)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9454,f3518])).
fof(f9454,plain,(
  ( ! [X59,X57,X58,X56] : (sP4(X56,X57) | sP4(X57,'F_0') | sP2(X58,X56,X59)) ) | $spl4),
  inference(resolution,[],[f2967,f235])).
fof(f12401,plain,(
  ( ! [X35,X33,X34,X32] : (sP4(X32,X33) | sP2(X34,X32,X35) | p(s(bool,X33)) | sP4('F_0','F_0')) ) | $spl4),
  inference(resolution,[],[f9454,f2965])).
fof(f12063,plain,(
  ( ! [X18] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'const.option.IS_NONE_1'(s('type.option.option'(X18),'const.option.NONE_0')))) ) | $spl362),
  inference(cnf_transformation,[],[f12063_D])).
fof(f12063_D,plain,(
  ( ! [X18] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'const.option.IS_NONE_1'(s('type.option.option'(X18),'const.option.NONE_0')))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f12220,plain,(
  ( ! [X21,X20] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'const.option.IS_NONE_1'(s('type.option.option'(X20),'const.option.NONE_0'))) | p(s(bool,X21))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f8596,f7360])).
fof(f7360,plain,(
  ( ! [X61] : (~sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X61) | p(s(bool,X61))) ) | $spl3),
  inference(subsumption_resolution,[],[f7300,f309])).
fof(f7300,plain,(
  ( ! [X61] : (p(s(bool,'F_0')) | p(s(bool,X61)) | ~sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X61)) ) | $spl3),
  inference(resolution,[],[f367,f271])).
fof(f8596,plain,(
  ( ! [X78,X76,X77] : (sP5(X77,X78) | sP4(X77,'const.option.IS_NONE_1'(s('type.option.option'(X76),'const.option.NONE_0')))) ) | $spl4),
  inference(subsumption_resolution,[],[f8569,f361])).
fof(f8569,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,'T_0')) | sP5(X77,X78) | sP4(X77,'const.option.IS_NONE_1'(s('type.option.option'(X76),'const.option.NONE_0')))) ) | $spl4),
  inference(superposition,[],[f3860,f217])).
fof(f3860,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP5(X0,X2) | sP4(X0,X1)) ) | $spl4),
  inference(resolution,[],[f3385,f879])).
fof(f3385,plain,(
  ( ! [X125,X123,X124] : (sP4(sK8,X123) | sP4(X124,X123) | sP5(X124,X125)) ) | $spl4),
  inference(resolution,[],[f1194,f273])).
fof(f12058,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0'))) | p(s(bool,X0))) ) | $spl356),
  inference(cnf_transformation,[],[f12058_D])).
fof(f12058_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0'))) | p(s(bool,X0))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f12244,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'const.option.IS_NONE_1'(s('type.option.option'(X16),'const.option.NONE_0'))) | p(s(bool,X15)) | p(s(bool,X17))) ) | $spl4),
  inference(subsumption_resolution,[],[f12218,f309])).
fof(f12218,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'const.option.IS_NONE_1'(s('type.option.option'(X16),'const.option.NONE_0'))) | p(s(bool,'F_0')) | p(s(bool,X15)) | p(s(bool,X17))) ) | $spl4),
  inference(resolution,[],[f8596,f271])).
fof(f12243,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'const.option.IS_NONE_1'(s('type.option.option'(X13),'const.option.NONE_0'))) | p(s(bool,X12)) | p(s(bool,X14))) ) | $spl4),
  inference(subsumption_resolution,[],[f12217,f309])).
fof(f12217,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'const.option.IS_NONE_1'(s('type.option.option'(X13),'const.option.NONE_0'))) | p(s(bool,'F_0')) | p(s(bool,X12)) | p(s(bool,X14))) ) | $spl4),
  inference(resolution,[],[f8596,f275])).
fof(f12216,plain,(
  ( ! [X10,X11,X9] : (sP4(X9,'const.option.IS_NONE_1'(s('type.option.option'(X10),'const.option.NONE_0'))) | p(s(bool,X9)) | sP4(X11,'F_0')) ) | $spl4),
  inference(resolution,[],[f8596,f3058])).
fof(f3058,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4(X90,'F_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f2973,f309])).
fof(f2973,plain,(
  ( ! [X90,X91] : (sP4(X90,'F_0') | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl4),
  inference(resolution,[],[f2221,f271])).
fof(f12060,plain,(
  ( ! [X10,X9] : (sP4(X9,'const.option.IS_NONE_1'(s('type.option.option'(X10),'const.option.NONE_0'))) | sP4(X9,'F_0')) ) | $spl360),
  inference(cnf_transformation,[],[f12060_D])).
fof(f12060_D,plain,(
  ( ! [X10,X9] : (sP4(X9,'const.option.IS_NONE_1'(s('type.option.option'(X10),'const.option.NONE_0'))) | sP4(X9,'F_0')) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f12215,plain,(
  ( ! [X6,X8,X7] : (sP4(X6,'const.option.IS_NONE_1'(s('type.option.option'(X7),'const.option.NONE_0'))) | p(s(bool,X8)) | sP4(X6,'F_0')) ) | $spl4),
  inference(resolution,[],[f8596,f3059])).
fof(f3059,plain,(
  ( ! [X92,X93] : (~sP5(X92,X93) | p(s(bool,X93)) | sP4(X92,'F_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f2974,f309])).
fof(f2974,plain,(
  ( ! [X92,X93] : (sP4(X92,'F_0') | p(s(bool,'F_0')) | p(s(bool,X93)) | ~sP5(X92,X93)) ) | $spl4),
  inference(resolution,[],[f2221,f271])).
fof(f12214,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'const.option.IS_NONE_1'(s('type.option.option'(X4),'const.option.NONE_0'))) | p(s(bool,X3)) | sP4('F_0',X5)) ) | $spl4),
  inference(resolution,[],[f8596,f3612])).
fof(f3612,plain,(
  ( ! [X90,X91] : (~sP5(X91,X90) | p(s(bool,X91)) | sP4('F_0',X90)) ) | $spl4),
  inference(subsumption_resolution,[],[f3526,f309])).
fof(f3526,plain,(
  ( ! [X90,X91] : (sP4('F_0',X90) | p(s(bool,'F_0')) | p(s(bool,X91)) | ~sP5(X91,X90)) ) | $spl4),
  inference(resolution,[],[f3479,f271])).
fof(f12059,plain,(
  ( ! [X4,X3] : (sP4(X3,'const.option.IS_NONE_1'(s('type.option.option'(X4),'const.option.NONE_0'))) | sP4('F_0',X3)) ) | $spl358),
  inference(cnf_transformation,[],[f12059_D])).
fof(f12059_D,plain,(
  ( ! [X4,X3] : (sP4(X3,'const.option.IS_NONE_1'(s('type.option.option'(X4),'const.option.NONE_0'))) | sP4('F_0',X3)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f12213,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0'))) | p(s(bool,X2)) | sP4('F_0',X0)) ) | $spl4),
  inference(resolution,[],[f8596,f3613])).
fof(f3613,plain,(
  ( ! [X92,X93] : (~sP5(X92,X93) | p(s(bool,X93)) | sP4('F_0',X92)) ) | $spl4),
  inference(subsumption_resolution,[],[f3527,f309])).
fof(f3527,plain,(
  ( ! [X92,X93] : (sP4('F_0',X92) | p(s(bool,'F_0')) | p(s(bool,X93)) | ~sP5(X92,X93)) ) | $spl4),
  inference(resolution,[],[f3479,f271])).
fof(f12024,plain,(
  ( ! [X30,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP1(X30,X31,X32)) ) | $spl354),
  inference(cnf_transformation,[],[f12024_D])).
fof(f12024_D,plain,(
  ( ! [X30,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP1(X30,X31,X32)) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f12184,plain,(
  ( ! [X37,X35,X36,X34] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X34) | sP1(X34,X35,X36) | p(s(bool,X37))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f8484,f7360])).
fof(f8484,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP5(X88,X89) | sP4(X88,X90) | sP1(X90,X91,X92)) ) | $spl4),
  inference(resolution,[],[f3860,f240])).
fof(f12209,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X26,X27,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f12182,f309])).
fof(f12182,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f8484,f271])).
fof(f12208,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X21,X22,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f12181,f309])).
fof(f12181,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f8484,f275])).
fof(f12180,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP1(X16,X17,X18) | p(s(bool,X15)) | sP4(X19,'F_0')) ) | $spl4),
  inference(resolution,[],[f8484,f3058])).
fof(f12021,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | sP1(X16,X17,X18)) ) | $spl352),
  inference(cnf_transformation,[],[f12021_D])).
fof(f12021_D,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | sP1(X16,X17,X18)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f12179,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP1(X11,X12,X13) | p(s(bool,X14)) | sP4(X10,'F_0')) ) | $spl4),
  inference(resolution,[],[f8484,f3059])).
fof(f12178,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP1(X6,X7,X8) | p(s(bool,X5)) | sP4('F_0',X9)) ) | $spl4),
  inference(resolution,[],[f8484,f3612])).
fof(f12020,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X5) | sP1(X6,X7,X8)) ) | $spl350),
  inference(cnf_transformation,[],[f12020_D])).
fof(f12020_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X5) | sP1(X6,X7,X8)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f12177,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP1(X1,X2,X3) | p(s(bool,X4)) | sP4('F_0',X0)) ) | $spl4),
  inference(resolution,[],[f8484,f3613])).
fof(f11985,plain,(
  ( ! [X30,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP1(X31,X30,X32)) ) | $spl346),
  inference(cnf_transformation,[],[f11985_D])).
fof(f11985_D,plain,(
  ( ! [X30,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP1(X31,X30,X32)) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f12149,plain,(
  ( ! [X37,X35,X36,X34] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X34) | sP1(X35,X34,X36) | p(s(bool,X37))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f8483,f7360])).
fof(f8483,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP5(X83,X84) | sP4(X83,X85) | sP1(X86,X85,X87)) ) | $spl4),
  inference(resolution,[],[f3860,f239])).
fof(f12174,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X27,X26,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f12147,f309])).
fof(f12147,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f8483,f271])).
fof(f12173,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X22,X21,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f12146,f309])).
fof(f12146,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f8483,f275])).
fof(f12145,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP1(X17,X16,X18) | p(s(bool,X15)) | sP4(X19,'F_0')) ) | $spl4),
  inference(resolution,[],[f8483,f3058])).
fof(f11982,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | sP1(X17,X16,X18)) ) | $spl344),
  inference(cnf_transformation,[],[f11982_D])).
fof(f11982_D,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | sP1(X17,X16,X18)) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f12144,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP1(X12,X11,X13) | p(s(bool,X14)) | sP4(X10,'F_0')) ) | $spl4),
  inference(resolution,[],[f8483,f3059])).
fof(f12143,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP1(X7,X6,X8) | p(s(bool,X5)) | sP4('F_0',X9)) ) | $spl4),
  inference(resolution,[],[f8483,f3612])).
fof(f11981,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X5) | sP1(X7,X6,X8)) ) | $spl342),
  inference(cnf_transformation,[],[f11981_D])).
fof(f11981_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X5) | sP1(X7,X6,X8)) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f12142,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP1(X2,X1,X3) | p(s(bool,X4)) | sP4('F_0',X0)) ) | $spl4),
  inference(resolution,[],[f8483,f3613])).
fof(f11945,plain,(
  ( ! [X30,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP2(X30,X31,X32)) ) | $spl338),
  inference(cnf_transformation,[],[f11945_D])).
fof(f11945_D,plain,(
  ( ! [X30,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP2(X30,X31,X32)) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f12113,plain,(
  ( ! [X37,X35,X36,X34] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X34) | sP2(X34,X35,X36) | p(s(bool,X37))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f8481,f7360])).
fof(f8481,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP5(X73,X74) | sP4(X73,X75) | sP2(X75,X76,X77)) ) | $spl4),
  inference(resolution,[],[f3860,f236])).
fof(f12138,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X26,X27,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f12111,f309])).
fof(f12111,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f8481,f271])).
fof(f12137,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X21,X22,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f12110,f309])).
fof(f12110,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f8481,f275])).
fof(f12109,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP2(X16,X17,X18) | p(s(bool,X15)) | sP4(X19,'F_0')) ) | $spl4),
  inference(resolution,[],[f8481,f3058])).
fof(f11942,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | sP2(X16,X17,X18)) ) | $spl336),
  inference(cnf_transformation,[],[f11942_D])).
fof(f11942_D,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | sP2(X16,X17,X18)) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f12108,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP2(X11,X12,X13) | p(s(bool,X14)) | sP4(X10,'F_0')) ) | $spl4),
  inference(resolution,[],[f8481,f3059])).
fof(f12107,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP2(X6,X7,X8) | p(s(bool,X5)) | sP4('F_0',X9)) ) | $spl4),
  inference(resolution,[],[f8481,f3612])).
fof(f11941,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X5) | sP2(X6,X7,X8)) ) | $spl334),
  inference(cnf_transformation,[],[f11941_D])).
fof(f11941_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X5) | sP2(X6,X7,X8)) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f12106,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP2(X1,X2,X3) | p(s(bool,X4)) | sP4('F_0',X0)) ) | $spl4),
  inference(resolution,[],[f8481,f3613])).
fof(f11907,plain,(
  ( ! [X30,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP2(X31,X30,X32)) ) | $spl330),
  inference(cnf_transformation,[],[f11907_D])).
fof(f11907_D,plain,(
  ( ! [X30,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP2(X31,X30,X32)) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f12078,plain,(
  ( ! [X37,X35,X36,X34] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X34) | sP2(X35,X34,X36) | p(s(bool,X37))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f8480,f7360])).
fof(f8480,plain,(
  ( ! [X70,X68,X72,X71,X69] : (sP5(X68,X69) | sP4(X68,X70) | sP2(X71,X70,X72)) ) | $spl4),
  inference(resolution,[],[f3860,f235])).
fof(f12103,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X27,X26,X28) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f12076,f309])).
fof(f12076,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X25)) | p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f8480,f271])).
fof(f12102,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X21,X23) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f12075,f309])).
fof(f12075,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X20)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f8480,f275])).
fof(f12074,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP2(X17,X16,X18) | p(s(bool,X15)) | sP4(X19,'F_0')) ) | $spl4),
  inference(resolution,[],[f8480,f3058])).
fof(f11904,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | sP2(X17,X16,X18)) ) | $spl328),
  inference(cnf_transformation,[],[f11904_D])).
fof(f11904_D,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | sP2(X17,X16,X18)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f12073,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP2(X12,X11,X13) | p(s(bool,X14)) | sP4(X10,'F_0')) ) | $spl4),
  inference(resolution,[],[f8480,f3059])).
fof(f12072,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP2(X7,X6,X8) | p(s(bool,X5)) | sP4('F_0',X9)) ) | $spl4),
  inference(resolution,[],[f8480,f3612])).
fof(f11903,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X5) | sP2(X7,X6,X8)) ) | $spl326),
  inference(cnf_transformation,[],[f11903_D])).
fof(f11903_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X5) | sP2(X7,X6,X8)) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f12071,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP2(X2,X1,X3) | p(s(bool,X4)) | sP4('F_0',X0)) ) | $spl4),
  inference(resolution,[],[f8480,f3613])).
fof(f12034,plain,(
  ( ! [X19,X18] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'const.option.IS_NONE_1'(s('type.option.option'(X18),'const.option.NONE_0'))) | p(s(bool,X19))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f7927,f7359])).
fof(f7927,plain,(
  ( ! [X78,X76,X77] : (sP5(X77,X78) | sP4(X78,'const.option.IS_NONE_1'(s('type.option.option'(X76),'const.option.NONE_0')))) ) | $spl4),
  inference(subsumption_resolution,[],[f7896,f361])).
fof(f7896,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,'T_0')) | sP5(X77,X78) | sP4(X78,'const.option.IS_NONE_1'(s('type.option.option'(X76),'const.option.NONE_0')))) ) | $spl4),
  inference(superposition,[],[f3826,f217])).
fof(f3826,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP5(X2,X0) | sP4(X0,X1)) ) | $spl4),
  inference(resolution,[],[f3384,f879])).
fof(f3384,plain,(
  ( ! [X121,X122,X120] : (sP4(sK8,X120) | sP4(X121,X120) | sP5(X122,X121)) ) | $spl4),
  inference(resolution,[],[f1194,f272])).
fof(f12062,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'const.option.IS_NONE_1'(s('type.option.option'(X16),'const.option.NONE_0'))) | p(s(bool,X17)) | p(s(bool,X15))) ) | $spl4),
  inference(subsumption_resolution,[],[f12033,f309])).
fof(f12033,plain,(
  ( ! [X17,X15,X16] : (sP4(X15,'const.option.IS_NONE_1'(s('type.option.option'(X16),'const.option.NONE_0'))) | p(s(bool,'F_0')) | p(s(bool,X17)) | p(s(bool,X15))) ) | $spl4),
  inference(resolution,[],[f7927,f271])).
fof(f12061,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'const.option.IS_NONE_1'(s('type.option.option'(X13),'const.option.NONE_0'))) | p(s(bool,X14)) | p(s(bool,X12))) ) | $spl4),
  inference(subsumption_resolution,[],[f12032,f309])).
fof(f12032,plain,(
  ( ! [X14,X12,X13] : (sP4(X12,'const.option.IS_NONE_1'(s('type.option.option'(X13),'const.option.NONE_0'))) | p(s(bool,'F_0')) | p(s(bool,X14)) | p(s(bool,X12))) ) | $spl4),
  inference(resolution,[],[f7927,f275])).
fof(f12031,plain,(
  ( ! [X10,X11,X9] : (sP4(X9,'const.option.IS_NONE_1'(s('type.option.option'(X10),'const.option.NONE_0'))) | p(s(bool,X11)) | sP4(X9,'F_0')) ) | $spl4),
  inference(resolution,[],[f7927,f3058])).
fof(f12030,plain,(
  ( ! [X6,X8,X7] : (sP4(X6,'const.option.IS_NONE_1'(s('type.option.option'(X7),'const.option.NONE_0'))) | p(s(bool,X6)) | sP4(X8,'F_0')) ) | $spl4),
  inference(resolution,[],[f7927,f3059])).
fof(f12029,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'const.option.IS_NONE_1'(s('type.option.option'(X4),'const.option.NONE_0'))) | p(s(bool,X5)) | sP4('F_0',X3)) ) | $spl4),
  inference(resolution,[],[f7927,f3612])).
fof(f12028,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(X1),'const.option.NONE_0'))) | p(s(bool,X0)) | sP4('F_0',X2)) ) | $spl4),
  inference(resolution,[],[f7927,f3613])).
fof(f11994,plain,(
  ( ! [X30,X33,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP1(X30,X31,X32) | p(s(bool,X33))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f7812,f7359])).
fof(f7812,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP5(X88,X89) | sP4(X89,X90) | sP1(X90,X91,X92)) ) | $spl4),
  inference(resolution,[],[f3826,f240])).
fof(f12023,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X26,X27,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl4),
  inference(subsumption_resolution,[],[f11993,f309])).
fof(f11993,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl4),
  inference(resolution,[],[f7812,f271])).
fof(f12022,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X21,X22,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl4),
  inference(subsumption_resolution,[],[f11992,f309])).
fof(f11992,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl4),
  inference(resolution,[],[f7812,f275])).
fof(f11991,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP1(X16,X17,X18) | p(s(bool,X19)) | sP4(X15,'F_0')) ) | $spl4),
  inference(resolution,[],[f7812,f3058])).
fof(f11990,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP1(X11,X12,X13) | p(s(bool,X10)) | sP4(X14,'F_0')) ) | $spl4),
  inference(resolution,[],[f7812,f3059])).
fof(f11989,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP1(X6,X7,X8) | p(s(bool,X9)) | sP4('F_0',X5)) ) | $spl4),
  inference(resolution,[],[f7812,f3612])).
fof(f11988,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP1(X1,X2,X3) | p(s(bool,X0)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f7812,f3613])).
fof(f11955,plain,(
  ( ! [X30,X33,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP1(X31,X30,X32) | p(s(bool,X33))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f7811,f7359])).
fof(f7811,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP5(X83,X84) | sP4(X84,X85) | sP1(X86,X85,X87)) ) | $spl4),
  inference(resolution,[],[f3826,f239])).
fof(f11984,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X27,X26,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl4),
  inference(subsumption_resolution,[],[f11954,f309])).
fof(f11954,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl4),
  inference(resolution,[],[f7811,f271])).
fof(f11983,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X22,X21,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl4),
  inference(subsumption_resolution,[],[f11953,f309])).
fof(f11953,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl4),
  inference(resolution,[],[f7811,f275])).
fof(f11952,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP1(X17,X16,X18) | p(s(bool,X19)) | sP4(X15,'F_0')) ) | $spl4),
  inference(resolution,[],[f7811,f3058])).
fof(f11951,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP1(X12,X11,X13) | p(s(bool,X10)) | sP4(X14,'F_0')) ) | $spl4),
  inference(resolution,[],[f7811,f3059])).
fof(f11950,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP1(X7,X6,X8) | p(s(bool,X9)) | sP4('F_0',X5)) ) | $spl4),
  inference(resolution,[],[f7811,f3612])).
fof(f11949,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP1(X2,X1,X3) | p(s(bool,X0)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f7811,f3613])).
fof(f11915,plain,(
  ( ! [X30,X33,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP2(X30,X31,X32) | p(s(bool,X33))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f7809,f7359])).
fof(f7809,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP5(X73,X74) | sP4(X74,X75) | sP2(X75,X76,X77)) ) | $spl4),
  inference(resolution,[],[f3826,f236])).
fof(f11944,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X26,X27,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl4),
  inference(subsumption_resolution,[],[f11914,f309])).
fof(f11914,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X26,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl4),
  inference(resolution,[],[f7809,f271])).
fof(f11943,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X21,X22,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl4),
  inference(subsumption_resolution,[],[f11913,f309])).
fof(f11913,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X21,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl4),
  inference(resolution,[],[f7809,f275])).
fof(f11912,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP2(X16,X17,X18) | p(s(bool,X19)) | sP4(X15,'F_0')) ) | $spl4),
  inference(resolution,[],[f7809,f3058])).
fof(f11911,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP2(X11,X12,X13) | p(s(bool,X10)) | sP4(X14,'F_0')) ) | $spl4),
  inference(resolution,[],[f7809,f3059])).
fof(f11910,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP2(X6,X7,X8) | p(s(bool,X9)) | sP4('F_0',X5)) ) | $spl4),
  inference(resolution,[],[f7809,f3612])).
fof(f11909,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP2(X1,X2,X3) | p(s(bool,X0)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f7809,f3613])).
fof(f11877,plain,(
  ( ! [X30,X33,X31,X32] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X30) | sP2(X31,X30,X32) | p(s(bool,X33))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f7808,f7359])).
fof(f7808,plain,(
  ( ! [X70,X68,X72,X71,X69] : (sP5(X68,X69) | sP4(X69,X70) | sP2(X71,X70,X72)) ) | $spl4),
  inference(resolution,[],[f3826,f235])).
fof(f11906,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X27,X26,X28) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl4),
  inference(subsumption_resolution,[],[f11876,f309])).
fof(f11876,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X27,X26,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X25))) ) | $spl4),
  inference(resolution,[],[f7808,f271])).
fof(f11905,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X21,X23) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl4),
  inference(subsumption_resolution,[],[f11875,f309])).
fof(f11875,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X21,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X20))) ) | $spl4),
  inference(resolution,[],[f7808,f275])).
fof(f11874,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP2(X17,X16,X18) | p(s(bool,X19)) | sP4(X15,'F_0')) ) | $spl4),
  inference(resolution,[],[f7808,f3058])).
fof(f11873,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP2(X12,X11,X13) | p(s(bool,X10)) | sP4(X14,'F_0')) ) | $spl4),
  inference(resolution,[],[f7808,f3059])).
fof(f11872,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP2(X7,X6,X8) | p(s(bool,X9)) | sP4('F_0',X5)) ) | $spl4),
  inference(resolution,[],[f7808,f3612])).
fof(f11871,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP2(X2,X1,X3) | p(s(bool,X0)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f7808,f3613])).
fof(f11683,plain,(
  ( ! [X18] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X18),'const.option.NONE_0')),'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl322),
  inference(cnf_transformation,[],[f11683_D])).
fof(f11683_D,plain,(
  ( ! [X18] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X18),'const.option.NONE_0')),'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f11839,plain,(
  ( ! [X21,X20] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X20),'const.option.NONE_0')),'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X21))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f7082,f7360])).
fof(f7082,plain,(
  ( ! [X78,X76,X77] : (sP5(X77,X78) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X76),'const.option.NONE_0')),X77)) ) | $spl4),
  inference(subsumption_resolution,[],[f7054,f361])).
fof(f7054,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,'T_0')) | sP5(X77,X78) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X76),'const.option.NONE_0')),X77)) ) | $spl4),
  inference(superposition,[],[f3270,f217])).
fof(f3270,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP5(X1,X2) | sP4(X0,X1)) ) | $spl4),
  inference(resolution,[],[f2130,f879])).
fof(f2130,plain,(
  ( ! [X125,X123,X124] : (sP4(sK8,X123) | sP4(X123,X124) | sP5(X124,X125)) ) | $spl4),
  inference(resolution,[],[f1193,f273])).
fof(f11678,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')),X1) | p(s(bool,X1))) ) | $spl316),
  inference(cnf_transformation,[],[f11678_D])).
fof(f11678_D,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')),X1) | p(s(bool,X1))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f11863,plain,(
  ( ! [X17,X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X15),'const.option.NONE_0')),X16) | p(s(bool,X16)) | p(s(bool,X17))) ) | $spl4),
  inference(subsumption_resolution,[],[f11837,f309])).
fof(f11837,plain,(
  ( ! [X17,X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X15),'const.option.NONE_0')),X16) | p(s(bool,'F_0')) | p(s(bool,X16)) | p(s(bool,X17))) ) | $spl4),
  inference(resolution,[],[f7082,f271])).
fof(f11862,plain,(
  ( ! [X14,X12,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X12),'const.option.NONE_0')),X13) | p(s(bool,X13)) | p(s(bool,X14))) ) | $spl4),
  inference(subsumption_resolution,[],[f11836,f309])).
fof(f11836,plain,(
  ( ! [X14,X12,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X12),'const.option.NONE_0')),X13) | p(s(bool,'F_0')) | p(s(bool,X13)) | p(s(bool,X14))) ) | $spl4),
  inference(resolution,[],[f7082,f275])).
fof(f11835,plain,(
  ( ! [X10,X11,X9] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X9),'const.option.NONE_0')),X10) | p(s(bool,X10)) | sP4(X11,'F_0')) ) | $spl4),
  inference(resolution,[],[f7082,f3058])).
fof(f11680,plain,(
  ( ! [X10,X9] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X9),'const.option.NONE_0')),X10) | sP4(X10,'F_0')) ) | $spl320),
  inference(cnf_transformation,[],[f11680_D])).
fof(f11680_D,plain,(
  ( ! [X10,X9] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X9),'const.option.NONE_0')),X10) | sP4(X10,'F_0')) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f11834,plain,(
  ( ! [X6,X8,X7] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X6),'const.option.NONE_0')),X7) | p(s(bool,X8)) | sP4(X7,'F_0')) ) | $spl4),
  inference(resolution,[],[f7082,f3059])).
fof(f11833,plain,(
  ( ! [X4,X5,X3] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X3),'const.option.NONE_0')),X4) | p(s(bool,X4)) | sP4('F_0',X5)) ) | $spl4),
  inference(resolution,[],[f7082,f3612])).
fof(f11679,plain,(
  ( ! [X4,X3] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X3),'const.option.NONE_0')),X4) | sP4('F_0',X4)) ) | $spl318),
  inference(cnf_transformation,[],[f11679_D])).
fof(f11679_D,plain,(
  ( ! [X4,X3] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X3),'const.option.NONE_0')),X4) | sP4('F_0',X4)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f11832,plain,(
  ( ! [X2,X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')),X1) | p(s(bool,X2)) | sP4('F_0',X1)) ) | $spl4),
  inference(resolution,[],[f7082,f3613])).
fof(f11644,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP1(X30,X31,X32)) ) | $spl314),
  inference(cnf_transformation,[],[f11644_D])).
fof(f11644_D,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP1(X30,X31,X32)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f11803,plain,(
  ( ! [X37,X35,X36,X34] : (sP4(X34,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP1(X34,X35,X36) | p(s(bool,X37))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6970,f7360])).
fof(f6970,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP5(X88,X89) | sP4(X90,X88) | sP1(X90,X91,X92)) ) | $spl4),
  inference(resolution,[],[f3270,f240])).
fof(f11828,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X25,X27,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f11801,f309])).
fof(f11801,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f6970,f271])).
fof(f11827,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X20,X22,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f11800,f309])).
fof(f11800,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f6970,f275])).
fof(f11799,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP1(X15,X17,X18) | p(s(bool,X16)) | sP4(X19,'F_0')) ) | $spl4),
  inference(resolution,[],[f6970,f3058])).
fof(f11641,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | sP1(X15,X17,X18)) ) | $spl312),
  inference(cnf_transformation,[],[f11641_D])).
fof(f11641_D,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | sP1(X15,X17,X18)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f11798,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP1(X10,X12,X13) | p(s(bool,X14)) | sP4(X11,'F_0')) ) | $spl4),
  inference(resolution,[],[f6970,f3059])).
fof(f11797,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP1(X5,X7,X8) | p(s(bool,X6)) | sP4('F_0',X9)) ) | $spl4),
  inference(resolution,[],[f6970,f3612])).
fof(f11640,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X6) | sP1(X5,X7,X8)) ) | $spl310),
  inference(cnf_transformation,[],[f11640_D])).
fof(f11640_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X6) | sP1(X5,X7,X8)) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f11796,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP1(X0,X2,X3) | p(s(bool,X4)) | sP4('F_0',X1)) ) | $spl4),
  inference(resolution,[],[f6970,f3613])).
fof(f11605,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP1(X31,X30,X32)) ) | $spl306),
  inference(cnf_transformation,[],[f11605_D])).
fof(f11605_D,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP1(X31,X30,X32)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f11768,plain,(
  ( ! [X37,X35,X36,X34] : (sP4(X34,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP1(X35,X34,X36) | p(s(bool,X37))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6969,f7360])).
fof(f6969,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP5(X83,X84) | sP4(X85,X83) | sP1(X86,X85,X87)) ) | $spl4),
  inference(resolution,[],[f3270,f239])).
fof(f11793,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X27,X25,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f11766,f309])).
fof(f11766,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f6969,f271])).
fof(f11792,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X22,X20,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f11765,f309])).
fof(f11765,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f6969,f275])).
fof(f11764,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP1(X17,X15,X18) | p(s(bool,X16)) | sP4(X19,'F_0')) ) | $spl4),
  inference(resolution,[],[f6969,f3058])).
fof(f11602,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | sP1(X17,X15,X18)) ) | $spl304),
  inference(cnf_transformation,[],[f11602_D])).
fof(f11602_D,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | sP1(X17,X15,X18)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f11763,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP1(X12,X10,X13) | p(s(bool,X14)) | sP4(X11,'F_0')) ) | $spl4),
  inference(resolution,[],[f6969,f3059])).
fof(f11762,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP1(X7,X5,X8) | p(s(bool,X6)) | sP4('F_0',X9)) ) | $spl4),
  inference(resolution,[],[f6969,f3612])).
fof(f11601,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X6) | sP1(X7,X5,X8)) ) | $spl302),
  inference(cnf_transformation,[],[f11601_D])).
fof(f11601_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X6) | sP1(X7,X5,X8)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f11761,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP1(X2,X0,X3) | p(s(bool,X4)) | sP4('F_0',X1)) ) | $spl4),
  inference(resolution,[],[f6969,f3613])).
fof(f11565,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP2(X30,X31,X32)) ) | $spl298),
  inference(cnf_transformation,[],[f11565_D])).
fof(f11565_D,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP2(X30,X31,X32)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f11732,plain,(
  ( ! [X37,X35,X36,X34] : (sP4(X34,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP2(X34,X35,X36) | p(s(bool,X37))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6967,f7360])).
fof(f6967,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP5(X73,X74) | sP4(X75,X73) | sP2(X75,X76,X77)) ) | $spl4),
  inference(resolution,[],[f3270,f236])).
fof(f11757,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X25,X27,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f11730,f309])).
fof(f11730,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f6967,f271])).
fof(f11756,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X20,X22,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f11729,f309])).
fof(f11729,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f6967,f275])).
fof(f11728,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP2(X15,X17,X18) | p(s(bool,X16)) | sP4(X19,'F_0')) ) | $spl4),
  inference(resolution,[],[f6967,f3058])).
fof(f11562,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | sP2(X15,X17,X18)) ) | $spl296),
  inference(cnf_transformation,[],[f11562_D])).
fof(f11562_D,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | sP2(X15,X17,X18)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f11727,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP2(X10,X12,X13) | p(s(bool,X14)) | sP4(X11,'F_0')) ) | $spl4),
  inference(resolution,[],[f6967,f3059])).
fof(f11726,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP2(X5,X7,X8) | p(s(bool,X6)) | sP4('F_0',X9)) ) | $spl4),
  inference(resolution,[],[f6967,f3612])).
fof(f11561,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X6) | sP2(X5,X7,X8)) ) | $spl294),
  inference(cnf_transformation,[],[f11561_D])).
fof(f11561_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X6) | sP2(X5,X7,X8)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f11725,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP2(X0,X2,X3) | p(s(bool,X4)) | sP4('F_0',X1)) ) | $spl4),
  inference(resolution,[],[f6967,f3613])).
fof(f11526,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP2(X31,X30,X32)) ) | $spl290),
  inference(cnf_transformation,[],[f11526_D])).
fof(f11526_D,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP2(X31,X30,X32)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f11698,plain,(
  ( ! [X37,X35,X36,X34] : (sP4(X34,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP2(X35,X34,X36) | p(s(bool,X37))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6966,f7360])).
fof(f6966,plain,(
  ( ! [X70,X68,X72,X71,X69] : (sP5(X68,X69) | sP4(X70,X68) | sP2(X71,X70,X72)) ) | $spl4),
  inference(resolution,[],[f3270,f235])).
fof(f11723,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X27,X25,X28) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f11696,f309])).
fof(f11696,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X26)) | p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f6966,f271])).
fof(f11722,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X20,X23) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl4),
  inference(subsumption_resolution,[],[f11695,f309])).
fof(f11695,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X21)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f6966,f275])).
fof(f11694,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP2(X17,X15,X18) | p(s(bool,X16)) | sP4(X19,'F_0')) ) | $spl4),
  inference(resolution,[],[f6966,f3058])).
fof(f11523,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | sP2(X17,X15,X18)) ) | $spl288),
  inference(cnf_transformation,[],[f11523_D])).
fof(f11523_D,plain,(
  ( ! [X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | sP2(X17,X15,X18)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f11693,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP2(X12,X10,X13) | p(s(bool,X14)) | sP4(X11,'F_0')) ) | $spl4),
  inference(resolution,[],[f6966,f3059])).
fof(f11692,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP2(X7,X5,X8) | p(s(bool,X6)) | sP4('F_0',X9)) ) | $spl4),
  inference(resolution,[],[f6966,f3612])).
fof(f11522,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X6) | sP2(X7,X5,X8)) ) | $spl286),
  inference(cnf_transformation,[],[f11522_D])).
fof(f11522_D,plain,(
  ( ! [X6,X8,X7,X5] : (sP4(X5,X6) | sP4('F_0',X6) | sP2(X7,X5,X8)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f11691,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP2(X2,X0,X3) | p(s(bool,X4)) | sP4('F_0',X1)) ) | $spl4),
  inference(resolution,[],[f6966,f3613])).
fof(f11654,plain,(
  ( ! [X19,X18] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X18),'const.option.NONE_0')),'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X19))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6658,f7359])).
fof(f6658,plain,(
  ( ! [X78,X76,X77] : (sP5(X77,X78) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X76),'const.option.NONE_0')),X78)) ) | $spl4),
  inference(subsumption_resolution,[],[f6626,f361])).
fof(f6626,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,'T_0')) | sP5(X77,X78) | sP4('const.option.IS_NONE_1'(s('type.option.option'(X76),'const.option.NONE_0')),X78)) ) | $spl4),
  inference(superposition,[],[f3241,f217])).
fof(f3241,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP5(X2,X1) | sP4(X0,X1)) ) | $spl4),
  inference(resolution,[],[f2129,f879])).
fof(f2129,plain,(
  ( ! [X121,X122,X120] : (sP4(sK8,X120) | sP4(X120,X121) | sP5(X122,X121)) ) | $spl4),
  inference(resolution,[],[f1193,f272])).
fof(f11682,plain,(
  ( ! [X17,X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X15),'const.option.NONE_0')),X16) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl4),
  inference(subsumption_resolution,[],[f11653,f309])).
fof(f11653,plain,(
  ( ! [X17,X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X15),'const.option.NONE_0')),X16) | p(s(bool,'F_0')) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl4),
  inference(resolution,[],[f6658,f271])).
fof(f11681,plain,(
  ( ! [X14,X12,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X12),'const.option.NONE_0')),X13) | p(s(bool,X14)) | p(s(bool,X13))) ) | $spl4),
  inference(subsumption_resolution,[],[f11652,f309])).
fof(f11652,plain,(
  ( ! [X14,X12,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X12),'const.option.NONE_0')),X13) | p(s(bool,'F_0')) | p(s(bool,X14)) | p(s(bool,X13))) ) | $spl4),
  inference(resolution,[],[f6658,f275])).
fof(f11651,plain,(
  ( ! [X10,X11,X9] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X9),'const.option.NONE_0')),X10) | p(s(bool,X11)) | sP4(X10,'F_0')) ) | $spl4),
  inference(resolution,[],[f6658,f3058])).
fof(f11650,plain,(
  ( ! [X6,X8,X7] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X6),'const.option.NONE_0')),X7) | p(s(bool,X7)) | sP4(X8,'F_0')) ) | $spl4),
  inference(resolution,[],[f6658,f3059])).
fof(f11649,plain,(
  ( ! [X4,X5,X3] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X3),'const.option.NONE_0')),X4) | p(s(bool,X5)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f6658,f3612])).
fof(f11648,plain,(
  ( ! [X2,X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(X0),'const.option.NONE_0')),X1) | p(s(bool,X1)) | sP4('F_0',X2)) ) | $spl4),
  inference(resolution,[],[f6658,f3613])).
fof(f11614,plain,(
  ( ! [X30,X33,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP1(X30,X31,X32) | p(s(bool,X33))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6543,f7359])).
fof(f6543,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP5(X88,X89) | sP4(X90,X89) | sP1(X90,X91,X92)) ) | $spl4),
  inference(resolution,[],[f3241,f240])).
fof(f11643,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X25,X27,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl4),
  inference(subsumption_resolution,[],[f11613,f309])).
fof(f11613,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl4),
  inference(resolution,[],[f6543,f271])).
fof(f11642,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X20,X22,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl4),
  inference(subsumption_resolution,[],[f11612,f309])).
fof(f11612,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl4),
  inference(resolution,[],[f6543,f275])).
fof(f11611,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP1(X15,X17,X18) | p(s(bool,X19)) | sP4(X16,'F_0')) ) | $spl4),
  inference(resolution,[],[f6543,f3058])).
fof(f11610,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP1(X10,X12,X13) | p(s(bool,X11)) | sP4(X14,'F_0')) ) | $spl4),
  inference(resolution,[],[f6543,f3059])).
fof(f11609,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP1(X5,X7,X8) | p(s(bool,X9)) | sP4('F_0',X6)) ) | $spl4),
  inference(resolution,[],[f6543,f3612])).
fof(f11608,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP1(X0,X2,X3) | p(s(bool,X1)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f6543,f3613])).
fof(f11575,plain,(
  ( ! [X30,X33,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP1(X31,X30,X32) | p(s(bool,X33))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6542,f7359])).
fof(f6542,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP5(X83,X84) | sP4(X85,X84) | sP1(X86,X85,X87)) ) | $spl4),
  inference(resolution,[],[f3241,f239])).
fof(f11604,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X27,X25,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl4),
  inference(subsumption_resolution,[],[f11574,f309])).
fof(f11574,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP1(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl4),
  inference(resolution,[],[f6542,f271])).
fof(f11603,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X22,X20,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl4),
  inference(subsumption_resolution,[],[f11573,f309])).
fof(f11573,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP1(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl4),
  inference(resolution,[],[f6542,f275])).
fof(f11572,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP1(X17,X15,X18) | p(s(bool,X19)) | sP4(X16,'F_0')) ) | $spl4),
  inference(resolution,[],[f6542,f3058])).
fof(f11571,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP1(X12,X10,X13) | p(s(bool,X11)) | sP4(X14,'F_0')) ) | $spl4),
  inference(resolution,[],[f6542,f3059])).
fof(f11570,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP1(X7,X5,X8) | p(s(bool,X9)) | sP4('F_0',X6)) ) | $spl4),
  inference(resolution,[],[f6542,f3612])).
fof(f11569,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP1(X2,X0,X3) | p(s(bool,X1)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f6542,f3613])).
fof(f11535,plain,(
  ( ! [X30,X33,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP2(X30,X31,X32) | p(s(bool,X33))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6540,f7359])).
fof(f6540,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP5(X73,X74) | sP4(X75,X74) | sP2(X75,X76,X77)) ) | $spl4),
  inference(resolution,[],[f3241,f236])).
fof(f11564,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X25,X27,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl4),
  inference(subsumption_resolution,[],[f11534,f309])).
fof(f11534,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X25,X27,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl4),
  inference(resolution,[],[f6540,f271])).
fof(f11563,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X20,X22,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl4),
  inference(subsumption_resolution,[],[f11533,f309])).
fof(f11533,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X20,X22,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl4),
  inference(resolution,[],[f6540,f275])).
fof(f11532,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP2(X15,X17,X18) | p(s(bool,X19)) | sP4(X16,'F_0')) ) | $spl4),
  inference(resolution,[],[f6540,f3058])).
fof(f11531,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP2(X10,X12,X13) | p(s(bool,X11)) | sP4(X14,'F_0')) ) | $spl4),
  inference(resolution,[],[f6540,f3059])).
fof(f11530,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP2(X5,X7,X8) | p(s(bool,X9)) | sP4('F_0',X6)) ) | $spl4),
  inference(resolution,[],[f6540,f3612])).
fof(f11529,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP2(X0,X2,X3) | p(s(bool,X1)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f6540,f3613])).
fof(f11496,plain,(
  ( ! [X30,X33,X31,X32] : (sP4(X30,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP2(X31,X30,X32) | p(s(bool,X33))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f6539,f7359])).
fof(f6539,plain,(
  ( ! [X70,X68,X72,X71,X69] : (sP5(X68,X69) | sP4(X70,X69) | sP2(X71,X70,X72)) ) | $spl4),
  inference(resolution,[],[f3241,f235])).
fof(f11525,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X27,X25,X28) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl4),
  inference(subsumption_resolution,[],[f11495,f309])).
fof(f11495,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP2(X27,X25,X28) | p(s(bool,'F_0')) | p(s(bool,X29)) | p(s(bool,X26))) ) | $spl4),
  inference(resolution,[],[f6539,f271])).
fof(f11524,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X20,X23) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl4),
  inference(subsumption_resolution,[],[f11494,f309])).
fof(f11494,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP2(X22,X20,X23) | p(s(bool,'F_0')) | p(s(bool,X24)) | p(s(bool,X21))) ) | $spl4),
  inference(resolution,[],[f6539,f275])).
fof(f11493,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP2(X17,X15,X18) | p(s(bool,X19)) | sP4(X16,'F_0')) ) | $spl4),
  inference(resolution,[],[f6539,f3058])).
fof(f11492,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP2(X12,X10,X13) | p(s(bool,X11)) | sP4(X14,'F_0')) ) | $spl4),
  inference(resolution,[],[f6539,f3059])).
fof(f11491,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP2(X7,X5,X8) | p(s(bool,X9)) | sP4('F_0',X6)) ) | $spl4),
  inference(resolution,[],[f6539,f3612])).
fof(f11490,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP2(X2,X0,X3) | p(s(bool,X1)) | sP4('F_0',X4)) ) | $spl4),
  inference(resolution,[],[f6539,f3613])).
fof(f11437,plain,(
  $false | ($spl3 | $spl26)),
  inference(subsumption_resolution,[],[f11391,f11318])).
fof(f11318,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl3 | $spl26)),
  inference(subsumption_resolution,[],[f11317,f323])).
fof(f11317,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1))) ) | ($spl3 | $spl26)),
  inference(forward_demodulation,[],[f1044,f7311])).
fof(f11391,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl26)),
  inference(resolution,[],[f11318,f338])).
fof(f11436,plain,(
  $false | ($spl3 | $spl26)),
  inference(subsumption_resolution,[],[f11390,f11318])).
fof(f11390,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl26)),
  inference(resolution,[],[f11318,f336])).
fof(f11435,plain,(
  $false | ($spl3 | $spl26)),
  inference(subsumption_resolution,[],[f11384,f11318])).
fof(f11384,plain,(
  ( ! [X149] : (~p(s(bool,X149))) ) | ($spl3 | $spl26)),
  inference(resolution,[],[f11318,f291])).
fof(f11389,plain,(
  $false | ($spl3 | $spl26)),
  inference(resolution,[],[f11318,f334])).
fof(f11304,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,X67)) | sP4('F_0',X68) | sP4(X69,'T_0') | sP4(X67,X69)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3613,f8109])).
fof(f8109,plain,(
  ( ! [X138,X136,X137] : (sP4(X136,'T_0') | sP5(X137,X138) | sP4(X138,X136)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f7931,f3826])).
fof(f11302,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,X61)) | sP4('F_0',X62) | sP4(X63,'T_0') | sP4(X63,X61)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3613,f8106])).
fof(f8106,plain,(
  ( ! [X130,X131,X129] : (sP4(X129,'T_0') | sP5(X130,X131) | sP4(X129,X131)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f7931,f3241])).
fof(f11299,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X54)) | sP4('F_0',X55) | sP4(X54,X56) | sP4('T_0',X56)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3613,f7855])).
fof(f7855,plain,(
  ( ! [X261,X263,X262] : (sP5(X261,X262) | sP4(X262,X263) | sP4('T_0',X263)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3826,f6663])).
fof(f11297,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | sP4('F_0',X49) | sP4('T_0',X50) | sP4(X50,X48)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3613,f6735])).
fof(f6735,plain,(
  ( ! [X130,X131,X129] : (sP4('T_0',X129) | sP5(X130,X131) | sP4(X129,X131)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f6663,f3241])).
fof(f11268,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,X73)) | sP4('F_0',X74) | sP4(X73,X75) | sP4(X75,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3612,f8528])).
fof(f8528,plain,(
  ( ! [X265,X264,X266] : (sP5(X264,X265) | sP4(X264,X266) | sP4(X266,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3860,f7931])).
fof(f11267,plain,(
  ( ! [X70,X72,X71] : (p(s(bool,X70)) | sP4('F_0',X71) | sP4(X70,X72) | sP4('T_0',X72)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3612,f8527])).
fof(f8527,plain,(
  ( ! [X261,X263,X262] : (sP5(X261,X262) | sP4(X261,X263) | sP4('T_0',X263)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3860,f6663])).
fof(f11265,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,X64)) | sP4('F_0',X65) | sP4(X66,'T_0') | sP4(X66,X64)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3612,f8107])).
fof(f8107,plain,(
  ( ! [X134,X132,X133] : (sP4(X132,'T_0') | sP5(X133,X134) | sP4(X132,X133)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f7931,f3270])).
fof(f11260,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | sP4('F_0',X52) | sP4(X53,X51) | sP4('T_0',X53)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3612,f7013])).
fof(f7013,plain,(
  ( ! [X261,X263,X262] : (sP5(X261,X262) | sP4(X263,X261) | sP4('T_0',X263)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3270,f6663])).
fof(f7236,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | $spl166),
  inference(cnf_transformation,[],[f7236_D])).
fof(f7236_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f7238,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | $spl168),
  inference(cnf_transformation,[],[f7238_D])).
fof(f7238_D,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f11085,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | $spl3),
  inference(forward_demodulation,[],[f11040,f7364])).
fof(f11040,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | $spl3),
  inference(superposition,[],[f357,f7364])).
fof(f11078,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | $spl3),
  inference(forward_demodulation,[],[f11015,f7364])).
fof(f11015,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | $spl3),
  inference(superposition,[],[f7364,f357])).
fof(f10957,plain,(
  ( ! [X237,X235,X236] : (sP4(X235,'F_0') | sP4('F_0',X235) | p(s(bool,X236)) | p(s(bool,X237)) | ~sP5(X236,X237)) ) | $spl4),
  inference(resolution,[],[f3521,f271])).
fof(f10955,plain,(
  ( ! [X231,X229,X230] : (sP4(X229,'F_0') | sP4('F_0',X229) | p(s(bool,X230)) | p(s(bool,X231)) | ~sP5(X230,X231)) ) | $spl4),
  inference(resolution,[],[f3521,f275])).
fof(f10994,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4('F_0',X0)) ) | $spl278),
  inference(cnf_transformation,[],[f10994_D])).
fof(f10994_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP4('F_0',X0)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f10930,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4('F_0',X182) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl4),
  inference(resolution,[],[f3521,f360])).
fof(f448,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X3)) | p(s(bool,X4))) ) | $spl12),
  inference(cnf_transformation,[],[f448_D])).
fof(f448_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X3)) | p(s(bool,X4))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f10997,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4('F_0',X25)) ) | $spl282),
  inference(cnf_transformation,[],[f10997_D])).
fof(f10997_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4('F_0',X25)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f10929,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,X179) | sP4('F_0',X178) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl4),
  inference(resolution,[],[f3521,f360])).
fof(f10885,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4('F_0',X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f3521,f225])).
fof(f10884,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4('F_0',X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f3521,f224])).
fof(f10996,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,X26) | sP4('F_0',X25) | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl4),
  inference(subsumption_resolution,[],[f10883,f360])).
fof(f10883,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP4('F_0',X25) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl4),
  inference(resolution,[],[f3521,f224])).
fof(f10882,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4('F_0',X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f3521,f223])).
fof(f10995,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X6) != s(bool,X9) | sP4('F_0',X5)) ) | $spl280),
  inference(cnf_transformation,[],[f10995_D])).
fof(f10995_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X6) != s(bool,X9) | sP4('F_0',X5)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f10881,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP4('F_0',X15) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f3521,f223])).
fof(f10880,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4('F_0',X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f3521,f221])).
fof(f10879,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP4('F_0',X5) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X6) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f3521,f219])).
fof(f10878,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl4),
  inference(resolution,[],[f3521,f219])).
fof(f10859,plain,(
  ( ! [X237,X235,X236] : (sP4('F_0',X235) | p(s(bool,X236)) | p(s(bool,X237)) | ~sP5(X236,X237)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f10819])).
fof(f10819,plain,(
  ( ! [X237,X235,X236] : (sP4('F_0',X235) | sP4('F_0',X235) | p(s(bool,X236)) | p(s(bool,X237)) | ~sP5(X236,X237)) ) | $spl4),
  inference(resolution,[],[f3520,f271])).
fof(f10861,plain,(
  ( ! [X231,X229,X230] : (sP4('F_0',X229) | p(s(bool,X230)) | p(s(bool,X231)) | ~sP5(X230,X231)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f10817])).
fof(f10817,plain,(
  ( ! [X231,X229,X230] : (sP4('F_0',X229) | sP4('F_0',X229) | p(s(bool,X230)) | p(s(bool,X231)) | ~sP5(X230,X231)) ) | $spl4),
  inference(resolution,[],[f3520,f275])).
fof(f10865,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4('F_0',X1)) ) | $spl272),
  inference(cnf_transformation,[],[f10865_D])).
fof(f10865_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP4('F_0',X1)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f10792,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4('F_0',X183) | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl4),
  inference(resolution,[],[f3520,f360])).
fof(f10868,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4('F_0',X26)) ) | $spl276),
  inference(cnf_transformation,[],[f10868_D])).
fof(f10868_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4('F_0',X26)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f10791,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,X179) | sP4('F_0',X179) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl4),
  inference(resolution,[],[f3520,f360])).
fof(f10747,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4('F_0',X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f3520,f225])).
fof(f10746,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4('F_0',X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f3520,f224])).
fof(f10867,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,X26) | sP4('F_0',X26) | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl4),
  inference(subsumption_resolution,[],[f10745,f360])).
fof(f10745,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP4('F_0',X26) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl4),
  inference(resolution,[],[f3520,f224])).
fof(f10744,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4('F_0',X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f3520,f223])).
fof(f10866,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X5) != s(bool,X9) | sP4('F_0',X6)) ) | $spl274),
  inference(cnf_transformation,[],[f10866_D])).
fof(f10866_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X5) != s(bool,X9) | sP4('F_0',X6)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f10743,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP4('F_0',X16) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X15) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f3520,f223])).
fof(f10742,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4('F_0',X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f3520,f221])).
fof(f10741,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP4('F_0',X6) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X5) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f3520,f219])).
fof(f10740,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4('F_0',X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl4),
  inference(resolution,[],[f3520,f219])).
fof(f10722,plain,(
  ( ! [X14,X15] : (sP4('F_0',X14) | sP4(X14,X15) | ~p(s(bool,X15)) | sP4('T_0',sK8)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f10481,f1312])).
fof(f10481,plain,(
  ( ! [X204,X205] : (sP4('F_0',X204) | sP4(X204,X205) | sP4(X205,'T_0')) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f10379,f334])).
fof(f10379,plain,(
  ( ! [X204,X205] : (p(s(bool,'F_0')) | sP4('F_0',X204) | sP4(X204,X205) | sP4(X205,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3519,f9727])).
fof(f9727,plain,(
  ( ! [X204,X205] : (sP4(X204,X205) | sP4(X204,'F_0') | sP4(X205,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f2968,f7931])).
fof(f10569,plain,(
  sP4(sK8,'T_0') | $spl266),
  inference(cnf_transformation,[],[f10569_D])).
fof(f10569_D,plain,(
  sP4(sK8,'T_0') <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f10720,plain,(
  ( ! [X10,X11] : (sP4('F_0',X10) | sP4(X10,X11) | ~p(s(bool,X11)) | sP4(sK8,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f10481,f1192])).
fof(f10695,plain,(
  ( ! [X23,X22] : (sP4(X22,X23) | sP4(X23,'T_0') | p(s(bool,X22)) | sP4('F_0','F_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f10481,f3519])).
fof(f10693,plain,(
  ( ! [X19,X18] : (sP4(X18,X19) | sP4(X19,'T_0') | p(s(bool,X18)) | sP4('F_0','F_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f10481,f2966])).
fof(f10665,plain,(
  ( ! [X12,X13] : (sP4('F_0',X12) | sP4(X12,X13) | ~p(s(bool,X13)) | sP4('T_0',sK8)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f10480,f1311])).
fof(f1311,plain,(
  ( ! [X74,X75] : (~sP4(X74,X75) | ~p(s(bool,X75)) | sP4(X74,sK8)) ) | $spl4),
  inference(resolution,[],[f882,f256])).
fof(f10480,plain,(
  ( ! [X202,X203] : (sP4('F_0',X202) | sP4(X202,X203) | sP4('T_0',X203)) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f10378,f334])).
fof(f10378,plain,(
  ( ! [X202,X203] : (p(s(bool,'F_0')) | sP4('F_0',X202) | sP4(X202,X203) | sP4('T_0',X203)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3519,f9726])).
fof(f9726,plain,(
  ( ! [X202,X203] : (sP4(X202,X203) | sP4(X202,'F_0') | sP4('T_0',X203)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f2968,f6663])).
fof(f10663,plain,(
  ( ! [X8,X9] : (sP4('F_0',X8) | sP4(X8,X9) | ~p(s(bool,X9)) | sP4(sK8,'T_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f10480,f1191])).
fof(f1191,plain,(
  ( ! [X74,X75] : (~sP4(X74,X75) | ~p(s(bool,X75)) | sP4(sK8,X74)) ) | $spl4),
  inference(resolution,[],[f881,f256])).
fof(f10639,plain,(
  ( ! [X23,X22] : (sP4(X22,X23) | sP4('T_0',X23) | p(s(bool,X22)) | sP4('F_0','F_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f10480,f3519])).
fof(f10637,plain,(
  ( ! [X19,X18] : (sP4(X18,X19) | sP4('T_0',X19) | p(s(bool,X18)) | sP4('F_0','F_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f10480,f2966])).
fof(f10609,plain,(
  ( ! [X14,X15] : (sP4('F_0',X14) | sP4(X15,X14) | ~p(s(bool,X15)) | sP4('T_0',sK8)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f10479,f1312])).
fof(f10479,plain,(
  ( ! [X200,X201] : (sP4('F_0',X200) | sP4(X201,X200) | sP4(X201,'T_0')) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f10377,f334])).
fof(f10377,plain,(
  ( ! [X200,X201] : (p(s(bool,'F_0')) | sP4('F_0',X200) | sP4(X201,X200) | sP4(X201,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3519,f9502])).
fof(f9502,plain,(
  ( ! [X204,X205] : (sP4(X204,X205) | sP4(X205,'F_0') | sP4(X204,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f2967,f7931])).
fof(f10607,plain,(
  ( ! [X10,X11] : (sP4('F_0',X10) | sP4(X11,X10) | ~p(s(bool,X11)) | sP4(sK8,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f10479,f1192])).
fof(f10582,plain,(
  ( ! [X23,X22] : (sP4(X22,X23) | sP4(X22,'T_0') | p(s(bool,X23)) | sP4('F_0','F_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f10479,f3519])).
fof(f10580,plain,(
  ( ! [X19,X18] : (sP4(X18,X19) | sP4(X18,'T_0') | p(s(bool,X19)) | sP4('F_0','F_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f10479,f2966])).
fof(f10550,plain,(
  ( ! [X12,X13] : (sP4('F_0',X12) | sP4(X13,X12) | ~p(s(bool,X13)) | sP4('T_0',sK8)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f10478,f1311])).
fof(f10478,plain,(
  ( ! [X198,X199] : (sP4('F_0',X198) | sP4(X199,X198) | sP4('T_0',X199)) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f10376,f334])).
fof(f10376,plain,(
  ( ! [X198,X199] : (p(s(bool,'F_0')) | sP4('F_0',X198) | sP4(X199,X198) | sP4('T_0',X199)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3519,f9501])).
fof(f9501,plain,(
  ( ! [X202,X203] : (sP4(X202,X203) | sP4(X203,'F_0') | sP4('T_0',X202)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f2967,f6663])).
fof(f10548,plain,(
  ( ! [X8,X9] : (sP4('F_0',X8) | sP4(X9,X8) | ~p(s(bool,X9)) | sP4(sK8,'T_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f10478,f1191])).
fof(f10524,plain,(
  ( ! [X23,X22] : (sP4(X22,X23) | sP4('T_0',X22) | p(s(bool,X23)) | sP4('F_0','F_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f10478,f3519])).
fof(f10522,plain,(
  ( ! [X19,X18] : (sP4(X18,X19) | sP4('T_0',X18) | p(s(bool,X19)) | sP4('F_0','F_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f10478,f2966])).
fof(f10178,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | sP4('F_0','F_0') | sP4(X204,X205) | sP4(X205,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3518,f9727])).
fof(f10177,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | sP4('F_0','F_0') | sP4(X202,X203) | sP4('T_0',X203)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3518,f9726])).
fof(f10176,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | sP4('F_0','F_0') | sP4(X201,X200) | sP4(X201,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3518,f9502])).
fof(f10175,plain,(
  ( ! [X198,X199] : (p(s(bool,X198)) | sP4('F_0','F_0') | sP4(X199,X198) | sP4('T_0',X199)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3518,f9501])).
fof(f10091,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,X67)) | sP4(X68,'F_0') | sP4(X69,'T_0') | sP4(X67,X69)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3059,f8109])).
fof(f10089,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,X61)) | sP4(X62,'F_0') | sP4(X63,'T_0') | sP4(X63,X61)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3059,f8106])).
fof(f10086,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X54)) | sP4(X55,'F_0') | sP4(X54,X56) | sP4('T_0',X56)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3059,f7855])).
fof(f10084,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | sP4(X49,'F_0') | sP4('T_0',X50) | sP4(X50,X48)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3059,f6735])).
fof(f10055,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,X73)) | sP4(X74,'F_0') | sP4(X73,X75) | sP4(X75,'T_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3058,f8528])).
fof(f10054,plain,(
  ( ! [X70,X72,X71] : (p(s(bool,X70)) | sP4(X71,'F_0') | sP4(X70,X72) | sP4('T_0',X72)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3058,f8527])).
fof(f10052,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,X64)) | sP4(X65,'F_0') | sP4(X66,'T_0') | sP4(X66,X64)) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f3058,f8107])).
fof(f10047,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | sP4(X52,'F_0') | sP4(X53,X51) | sP4('T_0',X53)) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f3058,f7013])).
fof(f10001,plain,(
  ( ! [X0] : (~sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) | $spl264),
  inference(cnf_transformation,[],[f10001_D])).
fof(f10001_D,plain,(
  ( ! [X0] : (~sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f9986,plain,(
  ( ! [X0] : (~sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f7360,f357])).
fof(f9903,plain,(
  ( ! [X41] : (sP4(X41,'T_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X41)) ) | $spl262),
  inference(cnf_transformation,[],[f9903_D])).
fof(f9903_D,plain,(
  ( ! [X41] : (sP4(X41,'T_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X41)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f9982,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X45) | sP4(X45,'T_0')) ) | ($spl3 | $spl4 | $spl204)),
  inference(resolution,[],[f7360,f8528])).
fof(f9901,plain,(
  ( ! [X33] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X33) | sP4('T_0',X33)) ) | $spl258),
  inference(cnf_transformation,[],[f9901_D])).
fof(f9901_D,plain,(
  ( ! [X33] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X33) | sP4('T_0',X33)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f9981,plain,(
  ( ! [X43,X42] : (p(s(bool,X42)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X43) | sP4('T_0',X43)) ) | ($spl3 | $spl4 | $spl146)),
  inference(resolution,[],[f7360,f8527])).
fof(f9902,plain,(
  ( ! [X37] : (sP4(X37,'T_0') | sP4(X37,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl260),
  inference(cnf_transformation,[],[f9902_D])).
fof(f9902_D,plain,(
  ( ! [X37] : (sP4(X37,'T_0') | sP4(X37,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f9979,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | sP4(X39,'T_0') | sP4(X39,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl4 | $spl204)),
  inference(resolution,[],[f7360,f8107])).
fof(f9900,plain,(
  ( ! [X29] : (sP4('T_0',X29) | sP4(X29,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl256),
  inference(cnf_transformation,[],[f9900_D])).
fof(f9900_D,plain,(
  ( ! [X29] : (sP4('T_0',X29) | sP4(X29,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f9974,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | sP4(X31,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP4('T_0',X31)) ) | ($spl3 | $spl4 | $spl146)),
  inference(resolution,[],[f7360,f7013])).
fof(f9929,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4(X17,'T_0') | p(s(bool,X16)) | sP4('F_0','F_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f9727,f2965])).
fof(f9905,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f9886,f1834])).
fof(f9886,plain,(
  ( ! [X0] : (~sP5(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f7359,f357])).
fof(f9881,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | sP4(X41,'T_0') | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X41)) ) | ($spl3 | $spl4 | $spl204)),
  inference(resolution,[],[f7359,f8109])).
fof(f9879,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | sP4(X37,'T_0') | sP4(X37,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl4 | $spl204)),
  inference(resolution,[],[f7359,f8106])).
fof(f9876,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X33) | sP4('T_0',X33)) ) | ($spl3 | $spl4 | $spl146)),
  inference(resolution,[],[f7359,f7855])).
fof(f9874,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | sP4('T_0',X29) | sP4(X29,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | ($spl3 | $spl4 | $spl146)),
  inference(resolution,[],[f7359,f6735])).
fof(f9828,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4('T_0',X17) | p(s(bool,X16)) | sP4('F_0','F_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f9726,f2965])).
fof(f9786,plain,(
  ( ! [X237,X235,X236] : (sP4(X235,'F_0') | p(s(bool,X236)) | p(s(bool,X237)) | ~sP5(X236,X237)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f9746])).
fof(f9746,plain,(
  ( ! [X237,X235,X236] : (sP4(X235,'F_0') | sP4(X235,'F_0') | p(s(bool,X236)) | p(s(bool,X237)) | ~sP5(X236,X237)) ) | $spl4),
  inference(resolution,[],[f2968,f271])).
fof(f9788,plain,(
  ( ! [X231,X229,X230] : (sP4(X229,'F_0') | p(s(bool,X230)) | p(s(bool,X231)) | ~sP5(X230,X231)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f9744])).
fof(f9744,plain,(
  ( ! [X231,X229,X230] : (sP4(X229,'F_0') | sP4(X229,'F_0') | p(s(bool,X230)) | p(s(bool,X231)) | ~sP5(X230,X231)) ) | $spl4),
  inference(resolution,[],[f2968,f275])).
fof(f9719,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X182,'F_0') | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X183) != s(bool,X185)) ) | $spl4),
  inference(resolution,[],[f2968,f360])).
fof(f9795,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(X25,'F_0')) ) | $spl254),
  inference(cnf_transformation,[],[f9795_D])).
fof(f9795_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(X25,'F_0')) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f9718,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,X179) | sP4(X178,'F_0') | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl4),
  inference(resolution,[],[f2968,f360])).
fof(f9674,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X35,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f2968,f225])).
fof(f9673,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X30,'F_0') | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f2968,f224])).
fof(f9794,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,X26) | sP4(X25,'F_0') | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl4),
  inference(subsumption_resolution,[],[f9672,f360])).
fof(f9672,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP4(X25,'F_0') | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl4),
  inference(resolution,[],[f2968,f224])).
fof(f9671,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X20,'F_0') | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f2968,f223])).
fof(f9793,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X6) != s(bool,X9) | sP4(X5,'F_0')) ) | $spl252),
  inference(cnf_transformation,[],[f9793_D])).
fof(f9793_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X6) != s(bool,X9) | sP4(X5,'F_0')) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f9670,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X15,'F_0') | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f2968,f223])).
fof(f9669,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X10,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f2968,f221])).
fof(f9668,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP4(X5,'F_0') | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X6) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f2968,f219])).
fof(f9667,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X0,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl4),
  inference(resolution,[],[f2968,f219])).
fof(f9643,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4(X16,'T_0') | p(s(bool,X17)) | sP4('F_0','F_0')) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f9502,f2965])).
fof(f9595,plain,(
  ( ! [X17,X16] : (sP4(X16,X17) | sP4('T_0',X16) | p(s(bool,X17)) | sP4('F_0','F_0')) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f9501,f2965])).
fof(f9521,plain,(
  ( ! [X237,X235,X236] : (sP4('F_0',X235) | sP4(X235,'F_0') | p(s(bool,X236)) | p(s(bool,X237)) | ~sP5(X236,X237)) ) | $spl4),
  inference(resolution,[],[f2967,f271])).
fof(f9519,plain,(
  ( ! [X231,X229,X230] : (sP4('F_0',X229) | sP4(X229,'F_0') | p(s(bool,X230)) | p(s(bool,X231)) | ~sP5(X230,X231)) ) | $spl4),
  inference(resolution,[],[f2967,f275])).
fof(f9494,plain,(
  ( ! [X185,X182,X184,X183] : (sP4(X182,X183) | sP4(X183,'F_0') | p(s(bool,X184)) | ~p(s(bool,X185)) | s(bool,X182) != s(bool,X185)) ) | $spl4),
  inference(resolution,[],[f2967,f360])).
fof(f9561,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(X26,'F_0')) ) | $spl246),
  inference(cnf_transformation,[],[f9561_D])).
fof(f9561_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26) | sP4(X26,'F_0')) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f9493,plain,(
  ( ! [X180,X178,X181,X179] : (sP4(X178,X179) | sP4(X179,'F_0') | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X181)) ) | $spl4),
  inference(resolution,[],[f2967,f360])).
fof(f9449,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,X36) | sP4(X36,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f2967,f225])).
fof(f9448,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,X31) | sP4(X31,'F_0') | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f2967,f224])).
fof(f9560,plain,(
  ( ! [X26,X29,X27,X25] : (sP4(X25,X26) | sP4(X26,'F_0') | ~p(s(bool,X27)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl4),
  inference(subsumption_resolution,[],[f9447,f360])).
fof(f9447,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP4(X25,X26) | sP4(X26,'F_0') | ~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X27) != s(bool,X29)) ) | $spl4),
  inference(resolution,[],[f2967,f224])).
fof(f9446,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,X21) | sP4(X21,'F_0') | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f2967,f223])).
fof(f9559,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X5) != s(bool,X9) | sP4(X6,'F_0')) ) | $spl244),
  inference(cnf_transformation,[],[f9559_D])).
fof(f9559_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,X6) | s(bool,X5) != s(bool,X9) | sP4(X6,'F_0')) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f9445,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,X16) | sP4(X16,'F_0') | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X15) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f2967,f223])).
fof(f9444,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,X11) | sP4(X11,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f2967,f221])).
fof(f9443,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,X6) | sP4(X6,'F_0') | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X5) != s(bool,X9) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f2967,f219])).
fof(f9442,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1) | sP4(X1,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl4),
  inference(resolution,[],[f2967,f219])).
fof(f2374,plain,(
  ( ! [X182] : (s(bool,'F_0') != s(bool,X182)) ) | $spl78),
  inference(cnf_transformation,[],[f2374_D])).
fof(f2374_D,plain,(
  ( ! [X182] : (s(bool,'F_0') != s(bool,X182)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8930,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl3),
  inference(subsumption_resolution,[],[f8706,f223])).
fof(f8706,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl3),
  inference(superposition,[],[f225,f7311])).
fof(f8929,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl3),
  inference(subsumption_resolution,[],[f8705,f360])).
fof(f8705,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl3),
  inference(superposition,[],[f225,f7311])).
fof(f8918,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') != s(bool,X56) | ~p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X54) != s(bool,X55)) ) | $spl3),
  inference(forward_demodulation,[],[f8917,f7311])).
fof(f8917,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X56) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X54) != s(bool,X55)) ) | $spl3),
  inference(subsumption_resolution,[],[f8696,f309])).
fof(f8696,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'F_0')) | ~p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X56) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X54) != s(bool,X55)) ) | $spl3),
  inference(superposition,[],[f223,f7311])).
fof(f8909,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl3),
  inference(subsumption_resolution,[],[f8686,f360])).
fof(f8686,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl3),
  inference(superposition,[],[f221,f7311])).
fof(f8908,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f8685,f223])).
fof(f8685,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl3),
  inference(superposition,[],[f221,f7311])).
fof(f8905,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X17) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X16)) ) | $spl3),
  inference(forward_demodulation,[],[f8904,f7311])).
fof(f8904,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X15) != s(bool,X16)) ) | $spl3),
  inference(subsumption_resolution,[],[f8683,f309])).
fof(f8683,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'F_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X15) != s(bool,X16)) ) | $spl3),
  inference(superposition,[],[f219,f7311])).
fof(f8652,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4(X4,'T_0') | p(s(bool,X3)) | p(s(bool,X5))) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f8626,f309])).
fof(f8626,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4(X4,'T_0') | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X5))) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f8528,f271])).
fof(f8651,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4(X1,'T_0') | p(s(bool,X0)) | p(s(bool,X2))) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f8625,f309])).
fof(f8625,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4(X1,'T_0') | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X2))) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f8528,f275])).
fof(f8624,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X4) | p(s(bool,X3)) | p(s(bool,X5))) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f8598,f309])).
fof(f8598,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X4) | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X5))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f8527,f271])).
fof(f8623,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X1) | p(s(bool,X0)) | p(s(bool,X2))) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f8597,f309])).
fof(f8597,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X1) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X2))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f8527,f275])).
fof(f8591,plain,(
  ( ! [X306,X307] : (sP5(X306,X307) | sP4(X306,'F_0')) ) | $spl240),
  inference(cnf_transformation,[],[f8591_D])).
fof(f8591_D,plain,(
  ( ! [X306,X307] : (sP5(X306,X307) | sP4(X306,'F_0')) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f8547,plain,(
  ( ! [X315,X316,X314,X317] : (sP5(X314,X315) | sP4(X314,'F_0') | p(s(bool,X316)) | p(s(bool,X317)) | ~sP5(X316,X317)) ) | $spl4),
  inference(resolution,[],[f3860,f271])).
fof(f8545,plain,(
  ( ! [X308,X306,X309,X307] : (sP5(X306,X307) | sP4(X306,'F_0') | p(s(bool,X308)) | p(s(bool,X309)) | ~sP5(X308,X309)) ) | $spl4),
  inference(resolution,[],[f3860,f275])).
fof(f8583,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X0,X2)) ) | $spl234),
  inference(cnf_transformation,[],[f8583_D])).
fof(f8583_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X0,X2)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f8520,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X234,X236) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl4),
  inference(resolution,[],[f3860,f360])).
fof(f8586,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31) | sP4(X30,X32)) ) | $spl238),
  inference(cnf_transformation,[],[f8586_D])).
fof(f8586_D,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31) | sP4(X30,X32)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f8519,plain,(
  ( ! [X231,X229,X233,X230,X232] : (sP5(X229,X230) | sP4(X229,X231) | ~p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X232) != s(bool,X233)) ) | $spl4),
  inference(resolution,[],[f3860,f360])).
fof(f8475,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X42,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl4),
  inference(resolution,[],[f3860,f225])).
fof(f8474,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X36,X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl4),
  inference(resolution,[],[f3860,f224])).
fof(f8585,plain,(
  ( ! [X30,X35,X33,X31,X32] : (sP5(X30,X31) | sP4(X30,X32) | ~p(s(bool,X33)) | p(s(bool,X35)) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(subsumption_resolution,[],[f8473,f360])).
fof(f8473,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP5(X30,X31) | sP4(X30,X32) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(resolution,[],[f3860,f224])).
fof(f8472,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X24,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl4),
  inference(resolution,[],[f3860,f223])).
fof(f8584,plain,(
  ( ! [X6,X8,X7,X11] : (sP5(X6,X7) | s(bool,X8) != s(bool,X11) | sP4(X6,X8)) ) | $spl236),
  inference(cnf_transformation,[],[f8584_D])).
fof(f8584_D,plain,(
  ( ! [X6,X8,X7,X11] : (sP5(X6,X7) | s(bool,X8) != s(bool,X11) | sP4(X6,X8)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f8471,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP5(X18,X19) | sP4(X18,X20) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) != s(bool,X23) | s(bool,X21) != s(bool,X22)) ) | $spl4),
  inference(resolution,[],[f3860,f223])).
fof(f8470,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X12,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl4),
  inference(resolution,[],[f3860,f221])).
fof(f8469,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP5(X6,X7) | sP4(X6,X8) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f3860,f219])).
fof(f8468,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X0,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl4),
  inference(resolution,[],[f3860,f219])).
fof(f8467,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'T_0') | sP4(X4,X3) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f8450,f309])).
fof(f8450,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'T_0') | sP4(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f8109,f271])).
fof(f8465,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'T_0') | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f8449,f309])).
fof(f8449,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'T_0') | sP4(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f8109,f275])).
fof(f8434,plain,(
  $false | ($spl4 | $spl228)),
  inference(subsumption_resolution,[],[f8426,f340])).
fof(f8426,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl228)),
  inference(resolution,[],[f8417,f879])).
fof(f8425,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | sP4('T_0',sK8)) ) | ($spl4 | $spl228)),
  inference(resolution,[],[f8417,f1312])).
fof(f8416,plain,(
  ( ! [X1] : (sP5('T_0',X1)) ) | $spl226),
  inference(cnf_transformation,[],[f8416_D])).
fof(f8416_D,plain,(
  ( ! [X1] : (sP5('T_0',X1)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f8417,plain,(
  ( ! [X0] : (sP4(X0,'T_0')) ) | $spl228),
  inference(cnf_transformation,[],[f8417_D])).
fof(f8417_D,plain,(
  ( ! [X0] : (sP4(X0,'T_0')) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f8413,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0') | sP5('T_0',X1)) ) | ($spl4 | $spl204)),
  inference(factoring,[],[f8107])).
fof(f8415,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'T_0') | sP4(X3,X4) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f8399,f309])).
fof(f8399,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'T_0') | sP4(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f8107,f271])).
fof(f8414,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'T_0') | sP4(X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f8398,f309])).
fof(f8398,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'T_0') | sP4(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f8107,f275])).
fof(f8387,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'T_0') | sP4(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f8370,f309])).
fof(f8370,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,'T_0') | sP4(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f8106,f271])).
fof(f8385,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'T_0') | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl4 | $spl204)),
  inference(subsumption_resolution,[],[f8369,f309])).
fof(f8369,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'T_0') | sP4(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl4 | $spl204)),
  inference(resolution,[],[f8106,f275])).
fof(f8354,plain,(
  ( ! [X0] : (sP5(X0,'T_0')) ) | $spl222),
  inference(cnf_transformation,[],[f8354_D])).
fof(f8354_D,plain,(
  ( ! [X0] : (sP5(X0,'T_0')) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f8349,plain,(
  ( ! [X0,X1] : (sP5(X0,'T_0') | sP4('T_0',X1)) ) | ($spl4 | $spl146)),
  inference(factoring,[],[f7855])).
fof(f8352,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X4) | p(s(bool,X5)) | p(s(bool,X3))) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f8325,f309])).
fof(f8325,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X3))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f7855,f271])).
fof(f8350,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X1) | p(s(bool,X2)) | p(s(bool,X0))) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f8324,f309])).
fof(f8324,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X0))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f7855,f275])).
fof(f8301,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'T_0') | $spl216),
  inference(cnf_transformation,[],[f8301_D])).
fof(f8301_D,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'T_0') <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f8293,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'T_0') | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl3 | $spl204)),
  inference(superposition,[],[f8123,f357])).
fof(f8123,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'T_0') | ($spl3 | $spl204)),
  inference(resolution,[],[f7931,f367])).
fof(f8257,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0))) ) | $spl214),
  inference(cnf_transformation,[],[f8257_D])).
fof(f8257_D,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f8244,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f7294,f357])).
fof(f7294,plain,(
  ( ! [X55] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X55) | ~p(s(bool,X55))) ) | $spl3),
  inference(resolution,[],[f367,f260])).
fof(f8191,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0))) ) | $spl212),
  inference(cnf_transformation,[],[f8191_D])).
fof(f8191_D,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f8178,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f7293,f357])).
fof(f7293,plain,(
  ( ! [X54] : (sP4(X54,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ~p(s(bool,X54))) ) | $spl3),
  inference(resolution,[],[f367,f257])).
fof(f8159,plain,(
  ( ! [X2,X0] : (sP4(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl208),
  inference(cnf_transformation,[],[f8159_D])).
fof(f8159_D,plain,(
  ( ! [X2,X0] : (sP4(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f8104,plain,(
  ( ! [X127,X125,X126] : (sP4(X125,'T_0') | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl204),
  inference(resolution,[],[f7931,f360])).
fof(f8062,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,'T_0') | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl204),
  inference(resolution,[],[f7931,f225])).
fof(f8160,plain,(
  ( ! [X4,X7] : (sP4(X4,'T_0') | s(bool,X4) != s(bool,X7)) ) | $spl210),
  inference(cnf_transformation,[],[f8160_D])).
fof(f8160_D,plain,(
  ( ! [X4,X7] : (sP4(X4,'T_0') | s(bool,X4) != s(bool,X7)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f8061,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,'T_0') | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl204),
  inference(resolution,[],[f7931,f225])).
fof(f8059,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,'T_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl204),
  inference(resolution,[],[f7931,f224])).
fof(f8058,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'T_0') | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl204),
  inference(resolution,[],[f7931,f223])).
fof(f8057,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'T_0') | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl204),
  inference(resolution,[],[f7931,f221])).
fof(f8056,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'T_0') | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl204),
  inference(resolution,[],[f7931,f221])).
fof(f8055,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl204),
  inference(resolution,[],[f7931,f219])).
fof(f8054,plain,(
  ( ! [X0] : (~sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) | $spl206),
  inference(cnf_transformation,[],[f8054_D])).
fof(f8054_D,plain,(
  ( ! [X0] : (~sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f8016,plain,(
  ( ! [X0] : (~sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f7292,f357])).
fof(f7292,plain,(
  ( ! [X53] : (~sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X53) | p(s(bool,X53))) ) | $spl3),
  inference(resolution,[],[f367,f255])).
fof(f7931,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP4(X0,'T_0')) ) | $spl204),
  inference(cnf_transformation,[],[f7931_D])).
fof(f7931_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP4(X0,'T_0')) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7932,plain,(
  ( ! [X2,X3] : (sP4(X2,'T_0') | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f7929,f309])).
fof(f7929,plain,(
  ( ! [X2,X3] : (sP4(X2,'T_0') | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f7866,f271])).
fof(f7866,plain,(
  ( ! [X292,X291] : (sP5(X291,X292) | sP4(X292,'T_0')) ) | $spl4),
  inference(resolution,[],[f3826,f339])).
fof(f339,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.bool.TRUTH')).
fof(f7930,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f7928,f309])).
fof(f7928,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0') | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl4),
  inference(resolution,[],[f7866,f275])).
fof(f7919,plain,(
  ( ! [X267,X266] : (sP5(X266,X267) | sP4(X267,'T_0')) ) | $spl200),
  inference(cnf_transformation,[],[f7919_D])).
fof(f7919_D,plain,(
  ( ! [X267,X266] : (sP5(X266,X267) | sP4(X267,'T_0')) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f7906,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | sP5(X61,X62) | sP4(X62,'T_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f7891])).
fof(f7891,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | sP5(X61,X62) | sP4(X62,'T_0') | ~p(s(bool,X60))) ) | $spl4),
  inference(superposition,[],[f3826,f325])).
fof(f7907,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP5(X58,X59) | sP4(X59,'T_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f7890])).
fof(f7890,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP5(X58,X59) | sP4(X59,'T_0') | ~p(s(bool,X57))) ) | $spl4),
  inference(superposition,[],[f3826,f327])).
fof(f7926,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP5(X55,X56) | sP4(X56,'T_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f7889,f320])).
fof(f7889,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP5(X55,X56) | sP4(X56,'T_0') | s(bool,'F_0') = s(bool,X54)) ) | $spl4),
  inference(superposition,[],[f3826,f358])).
fof(f7922,plain,(
  ( ! [X304,X303] : (sP5(X303,X304) | sP4(X304,'F_0')) ) | $spl202),
  inference(cnf_transformation,[],[f7922_D])).
fof(f7922_D,plain,(
  ( ! [X304,X303] : (sP5(X303,X304) | sP4(X304,'F_0')) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f7874,plain,(
  ( ! [X313,X312,X314,X311] : (sP5(X311,X312) | sP4(X312,'F_0') | p(s(bool,X313)) | p(s(bool,X314)) | ~sP5(X313,X314)) ) | $spl4),
  inference(resolution,[],[f3826,f271])).
fof(f7872,plain,(
  ( ! [X304,X306,X303,X305] : (sP5(X303,X304) | sP4(X304,'F_0') | p(s(bool,X305)) | p(s(bool,X306)) | ~sP5(X305,X306)) ) | $spl4),
  inference(resolution,[],[f3826,f275])).
fof(f7863,plain,(
  ( ! [X285,X284,X286] : (sP5(X284,X285) | sP4(X285,'T_0') | ~p(s(bool,X286))) ) | $spl4),
  inference(resolution,[],[f3826,f318])).
fof(f7862,plain,(
  ( ! [X282,X281,X283] : (sP5(X281,X282) | sP4(X282,'T_0') | ~p(s(bool,X283))) ) | $spl4),
  inference(resolution,[],[f3826,f313])).
fof(f7861,plain,(
  ( ! [X280,X279,X278] : (sP5(X278,X279) | sP4(X279,'T_0') | ~p(s(bool,X280))) ) | $spl4),
  inference(resolution,[],[f3826,f303])).
fof(f7860,plain,(
  ( ! [X277,X275,X276] : (sP5(X275,X276) | sP4(X276,'T_0') | p(s(bool,X277))) ) | $spl4),
  inference(resolution,[],[f3826,f302])).
fof(f7859,plain,(
  ( ! [X273,X272,X274] : (sP5(X272,X273) | sP4(X273,'T_0') | p(s(bool,X274))) ) | $spl4),
  inference(resolution,[],[f3826,f299])).
fof(f7858,plain,(
  ( ! [X269,X271,X270] : (sP5(X269,X270) | sP4(X270,'T_0') | p(s(bool,X271))) ) | $spl4),
  inference(resolution,[],[f3826,f296])).
fof(f7857,plain,(
  ( ! [X267,X268,X266] : (sP5(X266,X267) | sP4(X267,'T_0') | ~p(s(bool,X268))) ) | $spl4),
  inference(resolution,[],[f3826,f294])).
fof(f7910,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X1,X2)) ) | $spl194),
  inference(cnf_transformation,[],[f7910_D])).
fof(f7910_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X1,X2)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7848,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X235,X236) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl4),
  inference(resolution,[],[f3826,f360])).
fof(f7913,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31) | sP4(X31,X32)) ) | $spl198),
  inference(cnf_transformation,[],[f7913_D])).
fof(f7913_D,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31) | sP4(X31,X32)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f7847,plain,(
  ( ! [X231,X229,X233,X230,X232] : (sP5(X229,X230) | sP4(X230,X231) | ~p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X232) != s(bool,X233)) ) | $spl4),
  inference(resolution,[],[f3826,f360])).
fof(f7803,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X43,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl4),
  inference(resolution,[],[f3826,f225])).
fof(f7802,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X37,X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl4),
  inference(resolution,[],[f3826,f224])).
fof(f7912,plain,(
  ( ! [X30,X35,X33,X31,X32] : (sP5(X30,X31) | sP4(X31,X32) | ~p(s(bool,X33)) | p(s(bool,X35)) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(subsumption_resolution,[],[f7801,f360])).
fof(f7801,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP5(X30,X31) | sP4(X31,X32) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(resolution,[],[f3826,f224])).
fof(f7800,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X25,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl4),
  inference(resolution,[],[f3826,f223])).
fof(f7911,plain,(
  ( ! [X6,X8,X7,X11] : (sP5(X6,X7) | s(bool,X8) != s(bool,X11) | sP4(X7,X8)) ) | $spl196),
  inference(cnf_transformation,[],[f7911_D])).
fof(f7911_D,plain,(
  ( ! [X6,X8,X7,X11] : (sP5(X6,X7) | s(bool,X8) != s(bool,X11) | sP4(X7,X8)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f7799,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP5(X18,X19) | sP4(X19,X20) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) != s(bool,X23) | s(bool,X21) != s(bool,X22)) ) | $spl4),
  inference(resolution,[],[f3826,f223])).
fof(f7798,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X13,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl4),
  inference(resolution,[],[f3826,f221])).
fof(f7797,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP5(X6,X7) | sP4(X7,X8) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f3826,f219])).
fof(f7796,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X1,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl4),
  inference(resolution,[],[f3826,f219])).
fof(f7795,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(X0,sK8)) ) | $spl192),
  inference(cnf_transformation,[],[f7795_D])).
fof(f7795_D,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(X0,sK8)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f7785,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(X0,sK8) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f7327,f357])).
fof(f7327,plain,(
  ( ! [X83] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X83) | sP4(X83,sK8)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f367,f1314])).
fof(f7774,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) | $spl190),
  inference(cnf_transformation,[],[f7774_D])).
fof(f7774_D,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7736,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f7291,f357])).
fof(f7291,plain,(
  ( ! [X52] : (~sP4(X52,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X52))) ) | $spl3),
  inference(resolution,[],[f367,f255])).
fof(f7649,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(X0,sK8)) ) | $spl188),
  inference(cnf_transformation,[],[f7649_D])).
fof(f7649_D,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(X0,sK8)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7639,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(X0,sK8) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f7326,f357])).
fof(f7326,plain,(
  ( ! [X82] : (sP4(X82,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP4(X82,sK8)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f367,f1313])).
fof(f7629,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(sK8,X0)) ) | $spl186),
  inference(cnf_transformation,[],[f7629_D])).
fof(f7629_D,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(sK8,X0)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f7619,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP4(sK8,X0) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f7325,f357])).
fof(f7325,plain,(
  ( ! [X81] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X81) | sP4(sK8,X81)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f367,f1194])).
fof(f7609,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(sK8,X0)) ) | $spl184),
  inference(cnf_transformation,[],[f7609_D])).
fof(f7609_D,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(sK8,X0)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7599,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | sP4(sK8,X0) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f7324,f357])).
fof(f7324,plain,(
  ( ! [X80] : (sP4(X80,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | sP4(sK8,X80)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f367,f1193])).
fof(f7589,plain,(
  ( ! [X0,X1] : (sP1('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1)) ) | $spl182),
  inference(cnf_transformation,[],[f7589_D])).
fof(f7589_D,plain,(
  ( ! [X0,X1] : (sP1('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f7579,plain,(
  ( ! [X0,X1] : (sP1('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f663,f357])).
fof(f663,plain,(
  ( ! [X41,X40] : (sP1('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X40,X41)) ) | $spl3),
  inference(resolution,[],[f367,f240])).
fof(f7573,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1)) ) | $spl180),
  inference(cnf_transformation,[],[f7573_D])).
fof(f7573_D,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f7563,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f662,f357])).
fof(f662,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X39)) ) | $spl3),
  inference(resolution,[],[f367,f239])).
fof(f7557,plain,(
  ( ! [X0,X1] : (sP2('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1)) ) | $spl178),
  inference(cnf_transformation,[],[f7557_D])).
fof(f7557_D,plain,(
  ( ! [X0,X1] : (sP2('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f7547,plain,(
  ( ! [X0,X1] : (sP2('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f660,f357])).
fof(f660,plain,(
  ( ! [X35,X34] : (sP2('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X34,X35)) ) | $spl3),
  inference(resolution,[],[f367,f236])).
fof(f7541,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1)) ) | $spl176),
  inference(cnf_transformation,[],[f7541_D])).
fof(f7541_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f7532,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl3),
  inference(superposition,[],[f659,f357])).
fof(f659,plain,(
  ( ! [X33,X32] : (sP2(X32,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X33)) ) | $spl3),
  inference(resolution,[],[f367,f235])).
fof(f7430,plain,(
  sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl174),
  inference(cnf_transformation,[],[f7430_D])).
fof(f7430_D,plain,(
  sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7422,plain,(
  sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl3 | $spl146)),
  inference(superposition,[],[f7328,f357])).
fof(f7328,plain,(
  sP4('T_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ($spl3 | $spl146)),
  inference(resolution,[],[f367,f6663])).
fof(f7413,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),sK8) | $spl172),
  inference(cnf_transformation,[],[f7413_D])).
fof(f7413_D,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),sK8) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7401,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),sK8) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl3 | $spl4)),
  inference(superposition,[],[f7323,f357])).
fof(f7323,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),sK8) | ($spl3 | $spl4)),
  inference(resolution,[],[f367,f882])).
fof(f7391,plain,(
  sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl170),
  inference(cnf_transformation,[],[f7391_D])).
fof(f7391_D,plain,(
  sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7377,plain,(
  sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl3 | $spl4)),
  inference(superposition,[],[f7322,f357])).
fof(f7322,plain,(
  sP4(sK8,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ($spl3 | $spl4)),
  inference(resolution,[],[f367,f881])).
fof(f729,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl22),
  inference(cnf_transformation,[],[f729_D])).
fof(f729_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f7321,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X79) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f360])).
fof(f7276,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f367,f225])).
fof(f7275,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f224])).
fof(f7273,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f223])).
fof(f730,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl24),
  inference(cnf_transformation,[],[f730_D])).
fof(f730_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f7272,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X9) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f367,f223])).
fof(f7271,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f7270,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f367,f219])).
fof(f7269,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f219])).
fof(f7246,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7245,f1444])).
fof(f1444,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl4),
  inference(backward_demodulation,[],[f900,f898])).
fof(f898,plain,(
  s(bool,'F_0') != s(bool,sK8) | $spl4),
  inference(resolution,[],[f369,f320])).
fof(f7234,plain,(
  s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | $spl1),
  inference(forward_demodulation,[],[f7171,f1105])).
fof(f7171,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),sK7) | s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9('const.option.SOME_1'(s(sK6,sK9(sK7,sK6))),sK6)))) | $spl1),
  inference(superposition,[],[f1105,f357])).
fof(f7162,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP5(X0,X1)) ) | $spl164),
  inference(cnf_transformation,[],[f7162_D])).
fof(f7162_D,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP5(X0,X1)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f7151,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP5(X0,X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f3693,f357])).
fof(f3693,plain,(
  ( ! [X15,X16] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X15) | sP5(X15,X16)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f3692,f361])).
fof(f3692,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X15) | sP5(X15,X16)) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f3670,f900])).
fof(f3670,plain,(
  ( ! [X15,X16] : (~p(s(bool,sK8)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X15) | sP5(X15,X16)) ) | ($spl2 | $spl4)),
  inference(resolution,[],[f988,f2130])).
fof(f988,plain,(
  ( ! [X53] : (~sP4(X53,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ~p(s(bool,X53))) ) | $spl2),
  inference(resolution,[],[f366,f256])).
fof(f7137,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP5(X1,X0)) ) | $spl162),
  inference(cnf_transformation,[],[f7137_D])).
fof(f7137_D,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP5(X1,X0)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f7126,plain,(
  ( ! [X0,X1] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | sP5(X1,X0) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f3691,f357])).
fof(f3691,plain,(
  ( ! [X14,X13] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X13) | sP5(X14,X13)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f3690,f361])).
fof(f3690,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X13) | sP5(X14,X13)) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f3669,f900])).
fof(f3669,plain,(
  ( ! [X14,X13] : (~p(s(bool,sK8)) | sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X13) | sP5(X14,X13)) ) | ($spl2 | $spl4)),
  inference(resolution,[],[f988,f2129])).
fof(f7111,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X3) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f7085,f309])).
fof(f7085,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,X4) | sP4('T_0',X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f7013,f271])).
fof(f7110,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f7084,f309])).
fof(f7084,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1) | sP4('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f7013,f275])).
fof(f7077,plain,(
  ( ! [X304,X303] : (sP5(X303,X304) | sP4('F_0',X303)) ) | $spl160),
  inference(cnf_transformation,[],[f7077_D])).
fof(f7077_D,plain,(
  ( ! [X304,X303] : (sP5(X303,X304) | sP4('F_0',X303)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f7032,plain,(
  ( ! [X313,X312,X314,X311] : (sP5(X311,X312) | sP4('F_0',X311) | p(s(bool,X313)) | p(s(bool,X314)) | ~sP5(X313,X314)) ) | $spl4),
  inference(resolution,[],[f3270,f271])).
fof(f7030,plain,(
  ( ! [X304,X306,X303,X305] : (sP5(X303,X304) | sP4('F_0',X303) | p(s(bool,X305)) | p(s(bool,X306)) | ~sP5(X305,X306)) ) | $spl4),
  inference(resolution,[],[f3270,f275])).
fof(f7069,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X2,X0)) ) | $spl154),
  inference(cnf_transformation,[],[f7069_D])).
fof(f7069_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X2,X0)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f7006,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X236,X234) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl4),
  inference(resolution,[],[f3270,f360])).
fof(f7072,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31) | sP4(X32,X30)) ) | $spl158),
  inference(cnf_transformation,[],[f7072_D])).
fof(f7072_D,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31) | sP4(X32,X30)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f7005,plain,(
  ( ! [X231,X229,X233,X230,X232] : (sP5(X229,X230) | sP4(X231,X229) | ~p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X232) != s(bool,X233)) ) | $spl4),
  inference(resolution,[],[f3270,f360])).
fof(f6961,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X44,X42) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl4),
  inference(resolution,[],[f3270,f225])).
fof(f6960,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X38,X36) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl4),
  inference(resolution,[],[f3270,f224])).
fof(f7071,plain,(
  ( ! [X30,X35,X33,X31,X32] : (sP5(X30,X31) | sP4(X32,X30) | ~p(s(bool,X33)) | p(s(bool,X35)) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(subsumption_resolution,[],[f6959,f360])).
fof(f6959,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP5(X30,X31) | sP4(X32,X30) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(resolution,[],[f3270,f224])).
fof(f6958,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X26,X24) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl4),
  inference(resolution,[],[f3270,f223])).
fof(f7070,plain,(
  ( ! [X6,X8,X7,X11] : (sP5(X6,X7) | s(bool,X8) != s(bool,X11) | sP4(X8,X6)) ) | $spl156),
  inference(cnf_transformation,[],[f7070_D])).
fof(f7070_D,plain,(
  ( ! [X6,X8,X7,X11] : (sP5(X6,X7) | s(bool,X8) != s(bool,X11) | sP4(X8,X6)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f6957,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP5(X18,X19) | sP4(X20,X18) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) != s(bool,X23) | s(bool,X21) != s(bool,X22)) ) | $spl4),
  inference(resolution,[],[f3270,f223])).
fof(f6956,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X14,X12) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl4),
  inference(resolution,[],[f3270,f221])).
fof(f6955,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP5(X6,X7) | sP4(X8,X6) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f3270,f219])).
fof(f6954,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X2,X0) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl4),
  inference(resolution,[],[f3270,f219])).
fof(f6937,plain,(
  ( ! [X4,X5,X3] : (sP4('T_0',X3) | sP4(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f6919,f309])).
fof(f6919,plain,(
  ( ! [X4,X5,X3] : (sP4('T_0',X3) | sP4(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f6735,f271])).
fof(f6935,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0) | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl4 | $spl146)),
  inference(subsumption_resolution,[],[f6918,f309])).
fof(f6918,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0) | sP4(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl4 | $spl146)),
  inference(resolution,[],[f6735,f275])).
fof(f6733,plain,(
  ( ! [X127,X125,X126] : (sP4('T_0',X125) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl146),
  inference(resolution,[],[f6663,f360])).
fof(f6691,plain,(
  ( ! [X30,X28,X31,X29] : (sP4('T_0',X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl146),
  inference(resolution,[],[f6663,f225])).
fof(f6787,plain,(
  ( ! [X4,X7] : (sP4('T_0',X4) | s(bool,X4) != s(bool,X7)) ) | $spl150),
  inference(cnf_transformation,[],[f6787_D])).
fof(f6787_D,plain,(
  ( ! [X4,X7] : (sP4('T_0',X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f6690,plain,(
  ( ! [X26,X24,X27,X25] : (sP4('T_0',X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl146),
  inference(resolution,[],[f6663,f225])).
fof(f6688,plain,(
  ( ! [X19,X17,X18,X16] : (sP4('T_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl146),
  inference(resolution,[],[f6663,f224])).
fof(f6687,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('T_0',X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl146),
  inference(resolution,[],[f6663,f223])).
fof(f6686,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('T_0',X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl146),
  inference(resolution,[],[f6663,f221])).
fof(f6685,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('T_0',X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl146),
  inference(resolution,[],[f6663,f221])).
fof(f6684,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl146),
  inference(resolution,[],[f6663,f219])).
fof(f6664,plain,(
  ( ! [X2,X3] : (sP4('T_0',X2) | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f6661,f309])).
fof(f6661,plain,(
  ( ! [X2,X3] : (sP4('T_0',X2) | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f6596,f271])).
fof(f6596,plain,(
  ( ! [X288,X289] : (sP5(X288,X289) | sP4('T_0',X289)) ) | $spl4),
  inference(resolution,[],[f3241,f339])).
fof(f6662,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f6660,f309])).
fof(f6660,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl4),
  inference(resolution,[],[f6596,f275])).
fof(f6650,plain,(
  ( ! [X263,X264] : (sP5(X263,X264) | sP4('T_0',X264)) ) | $spl142),
  inference(cnf_transformation,[],[f6650_D])).
fof(f6650_D,plain,(
  ( ! [X263,X264] : (sP5(X263,X264) | sP4('T_0',X264)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f6637,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | sP5(X61,X62) | sP4('T_0',X62)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f6621])).
fof(f6621,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | sP5(X61,X62) | sP4('T_0',X62) | ~p(s(bool,X60))) ) | $spl4),
  inference(superposition,[],[f3241,f325])).
fof(f6638,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP5(X58,X59) | sP4('T_0',X59)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f6620])).
fof(f6620,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP5(X58,X59) | sP4('T_0',X59) | ~p(s(bool,X57))) ) | $spl4),
  inference(superposition,[],[f3241,f327])).
fof(f6657,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP5(X55,X56) | sP4('T_0',X56)) ) | $spl4),
  inference(subsumption_resolution,[],[f6619,f320])).
fof(f6619,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP5(X55,X56) | sP4('T_0',X56) | s(bool,'F_0') = s(bool,X54)) ) | $spl4),
  inference(superposition,[],[f3241,f358])).
fof(f6653,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4('F_0',X301)) ) | $spl144),
  inference(cnf_transformation,[],[f6653_D])).
fof(f6653_D,plain,(
  ( ! [X300,X301] : (sP5(X300,X301) | sP4('F_0',X301)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f6604,plain,(
  ( ! [X308,X310,X309,X311] : (sP5(X308,X309) | sP4('F_0',X309) | p(s(bool,X310)) | p(s(bool,X311)) | ~sP5(X310,X311)) ) | $spl4),
  inference(resolution,[],[f3241,f271])).
fof(f6602,plain,(
  ( ! [X300,X302,X301,X303] : (sP5(X300,X301) | sP4('F_0',X301) | p(s(bool,X302)) | p(s(bool,X303)) | ~sP5(X302,X303)) ) | $spl4),
  inference(resolution,[],[f3241,f275])).
fof(f6593,plain,(
  ( ! [X282,X281,X283] : (sP5(X281,X282) | sP4('T_0',X282) | ~p(s(bool,X283))) ) | $spl4),
  inference(resolution,[],[f3241,f318])).
fof(f6592,plain,(
  ( ! [X280,X279,X278] : (sP5(X278,X279) | sP4('T_0',X279) | ~p(s(bool,X280))) ) | $spl4),
  inference(resolution,[],[f3241,f313])).
fof(f6591,plain,(
  ( ! [X277,X275,X276] : (sP5(X275,X276) | sP4('T_0',X276) | ~p(s(bool,X277))) ) | $spl4),
  inference(resolution,[],[f3241,f303])).
fof(f6590,plain,(
  ( ! [X273,X272,X274] : (sP5(X272,X273) | sP4('T_0',X273) | p(s(bool,X274))) ) | $spl4),
  inference(resolution,[],[f3241,f302])).
fof(f6589,plain,(
  ( ! [X269,X271,X270] : (sP5(X269,X270) | sP4('T_0',X270) | p(s(bool,X271))) ) | $spl4),
  inference(resolution,[],[f3241,f299])).
fof(f6588,plain,(
  ( ! [X267,X268,X266] : (sP5(X266,X267) | sP4('T_0',X267) | p(s(bool,X268))) ) | $spl4),
  inference(resolution,[],[f3241,f296])).
fof(f6587,plain,(
  ( ! [X263,X265,X264] : (sP5(X263,X264) | sP4('T_0',X264) | ~p(s(bool,X265))) ) | $spl4),
  inference(resolution,[],[f3241,f294])).
fof(f6641,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X2,X1)) ) | $spl136),
  inference(cnf_transformation,[],[f6641_D])).
fof(f6641_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP5(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP4(X2,X1)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f6579,plain,(
  ( ! [X237,X235,X238,X236,X234] : (sP5(X234,X235) | sP4(X236,X235) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X236) != s(bool,X238)) ) | $spl4),
  inference(resolution,[],[f3241,f360])).
fof(f6644,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31) | sP4(X32,X31)) ) | $spl140),
  inference(cnf_transformation,[],[f6644_D])).
fof(f6644_D,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31) | sP4(X32,X31)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f6578,plain,(
  ( ! [X231,X229,X233,X230,X232] : (sP5(X229,X230) | sP4(X231,X230) | ~p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X232) != s(bool,X233)) ) | $spl4),
  inference(resolution,[],[f3241,f360])).
fof(f6534,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP5(X42,X43) | sP4(X44,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl4),
  inference(resolution,[],[f3241,f225])).
fof(f6533,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP5(X36,X37) | sP4(X38,X37) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) ) | $spl4),
  inference(resolution,[],[f3241,f224])).
fof(f6643,plain,(
  ( ! [X30,X35,X33,X31,X32] : (sP5(X30,X31) | sP4(X32,X31) | ~p(s(bool,X33)) | p(s(bool,X35)) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(subsumption_resolution,[],[f6532,f360])).
fof(f6532,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP5(X30,X31) | sP4(X32,X31) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(resolution,[],[f3241,f224])).
fof(f6531,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP5(X24,X25) | sP4(X26,X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X26) != s(bool,X28)) ) | $spl4),
  inference(resolution,[],[f3241,f223])).
fof(f6642,plain,(
  ( ! [X6,X8,X7,X11] : (sP5(X6,X7) | s(bool,X8) != s(bool,X11) | sP4(X8,X7)) ) | $spl138),
  inference(cnf_transformation,[],[f6642_D])).
fof(f6642_D,plain,(
  ( ! [X6,X8,X7,X11] : (sP5(X6,X7) | s(bool,X8) != s(bool,X11) | sP4(X8,X7)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f6530,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP5(X18,X19) | sP4(X20,X19) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) != s(bool,X23) | s(bool,X21) != s(bool,X22)) ) | $spl4),
  inference(resolution,[],[f3241,f223])).
fof(f6529,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP5(X12,X13) | sP4(X14,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl4),
  inference(resolution,[],[f3241,f221])).
fof(f6528,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP5(X6,X7) | sP4(X8,X7) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f3241,f219])).
fof(f6527,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1) | sP4(X2,X1) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl4),
  inference(resolution,[],[f3241,f219])).
fof(f5663,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,X5) | ~sP1(X7,X5,X6) | ~sP2(X6,X5,X7) | ~p(s(bool,X5))) ) | $spl130),
  inference(cnf_transformation,[],[f5663_D])).
fof(f5663_D,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,X5) | ~sP1(X7,X5,X6) | ~sP2(X6,X5,X7) | ~p(s(bool,X5))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f6475,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl4),
  inference(subsumption_resolution,[],[f6460,f1312])).
fof(f6460,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | sP4(X5,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl4),
  inference(resolution,[],[f4883,f230])).
fof(f230,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f6210,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl4),
  inference(subsumption_resolution,[],[f6195,f1311])).
fof(f6195,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | sP4(X4,X5) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl4),
  inference(resolution,[],[f4735,f230])).
fof(f5932,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl4),
  inference(subsumption_resolution,[],[f5915,f1192])).
fof(f5915,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | sP4(X5,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl4),
  inference(resolution,[],[f3373,f230])).
fof(f5662,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl4),
  inference(subsumption_resolution,[],[f5645,f1191])).
fof(f5645,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | sP4(X4,X5) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) ) | $spl4),
  inference(resolution,[],[f2118,f230])).
fof(f1670,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21)) ) | $spl50),
  inference(cnf_transformation,[],[f1670_D])).
fof(f1670_D,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f5246,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97)) ) | $spl2),
  inference(forward_demodulation,[],[f5245,f1008])).
fof(f5245,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5043,f360])).
fof(f5043,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(superposition,[],[f225,f1008])).
fof(f5242,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl2),
  inference(subsumption_resolution,[],[f5039,f223])).
fof(f5039,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl2),
  inference(superposition,[],[f225,f1008])).
fof(f5241,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl2),
  inference(subsumption_resolution,[],[f5038,f360])).
fof(f5038,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl2),
  inference(superposition,[],[f225,f1008])).
fof(f5230,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X38) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl2),
  inference(forward_demodulation,[],[f5229,f1008])).
fof(f5229,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5023,f360])).
fof(f5023,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(superposition,[],[f221,f1008])).
fof(f5226,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl2),
  inference(subsumption_resolution,[],[f5019,f360])).
fof(f5019,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl2),
  inference(superposition,[],[f221,f1008])).
fof(f5225,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl2),
  inference(subsumption_resolution,[],[f5018,f223])).
fof(f5018,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl2),
  inference(superposition,[],[f221,f1008])).
fof(f4976,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,sK8) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) | $spl126),
  inference(cnf_transformation,[],[f4976_D])).
fof(f4976_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,sK8) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f4913,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(X174,sK8) | sP4(X175,X174) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl4),
  inference(resolution,[],[f1314,f360])).
fof(f4871,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,sK8) | sP4(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f1314,f225])).
fof(f4977,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,sK8) | s(bool,X6) != s(bool,X9) | sP4(X6,X5)) ) | $spl128),
  inference(cnf_transformation,[],[f4977_D])).
fof(f4977_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,sK8) | s(bool,X6) != s(bool,X9) | sP4(X6,X5)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f4870,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,sK8) | sP4(X31,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) ) | $spl4),
  inference(resolution,[],[f1314,f225])).
fof(f4868,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,sK8) | sP4(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f1314,f224])).
fof(f4867,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,sK8) | sP4(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f1314,f223])).
fof(f4866,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,sK8) | sP4(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f1314,f221])).
fof(f4865,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,sK8) | sP4(X6,X5) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f1314,f221])).
fof(f4864,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl4),
  inference(resolution,[],[f1314,f219])).
fof(f4829,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,sK8) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) | $spl122),
  inference(cnf_transformation,[],[f4829_D])).
fof(f4829_D,plain,(
  ( ! [X0,X3,X1] : (sP4(X0,sK8) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f4765,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(X174,sK8) | sP4(X174,X175) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl4),
  inference(resolution,[],[f1313,f360])).
fof(f4723,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(X35,sK8) | sP4(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f1313,f225])).
fof(f4830,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,sK8) | s(bool,X6) != s(bool,X9) | sP4(X5,X6)) ) | $spl124),
  inference(cnf_transformation,[],[f4830_D])).
fof(f4830_D,plain,(
  ( ! [X6,X5,X9] : (sP4(X5,sK8) | s(bool,X6) != s(bool,X9) | sP4(X5,X6)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f4722,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(X30,sK8) | sP4(X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) ) | $spl4),
  inference(resolution,[],[f1313,f225])).
fof(f4720,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(X20,sK8) | sP4(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f1313,f224])).
fof(f4719,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(X15,sK8) | sP4(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f1313,f223])).
fof(f4718,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(X10,sK8) | sP4(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f1313,f221])).
fof(f4717,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(X5,sK8) | sP4(X5,X6) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f1313,f221])).
fof(f4716,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl4),
  inference(resolution,[],[f1313,f219])).
fof(f4405,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | sP4(X1,X0)) ) | $spl120),
  inference(cnf_transformation,[],[f4405_D])).
fof(f4405_D,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | sP4(X1,X0)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f4448,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X4,X3) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl4),
  inference(subsumption_resolution,[],[f4433,f309])).
fof(f4433,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl4),
  inference(resolution,[],[f4270,f271])).
fof(f4270,plain,(
  ( ! [X68,X69,X67] : (sP4(X67,sK8) | sP4(X68,X67) | sP5(X68,X69)) ) | $spl4),
  inference(subsumption_resolution,[],[f4269,f361])).
fof(f4269,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'T_0')) | sP4(X67,sK8) | sP4(X68,X67) | sP5(X68,X69)) ) | $spl4),
  inference(forward_demodulation,[],[f4235,f900])).
fof(f4235,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,sK8)) | sP4(X67,sK8) | sP4(X68,X67) | sP5(X68,X69)) ) | $spl4),
  inference(resolution,[],[f1312,f3385])).
fof(f4447,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f4432,f309])).
fof(f4432,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f4270,f275])).
fof(f4406,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X4,X3) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl4),
  inference(subsumption_resolution,[],[f4390,f309])).
fof(f4390,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl4),
  inference(resolution,[],[f4268,f271])).
fof(f4268,plain,(
  ( ! [X66,X64,X65] : (sP4(X64,sK8) | sP4(X65,X64) | sP5(X66,X65)) ) | $spl4),
  inference(subsumption_resolution,[],[f4267,f361])).
fof(f4267,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | sP4(X64,sK8) | sP4(X65,X64) | sP5(X66,X65)) ) | $spl4),
  inference(forward_demodulation,[],[f4234,f900])).
fof(f4234,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,sK8)) | sP4(X64,sK8) | sP4(X65,X64) | sP5(X66,X65)) ) | $spl4),
  inference(resolution,[],[f1312,f3384])).
fof(f4404,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl4),
  inference(subsumption_resolution,[],[f4389,f309])).
fof(f4389,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl4),
  inference(resolution,[],[f4268,f275])).
fof(f4320,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | sP4(X0,X1)) ) | $spl118),
  inference(cnf_transformation,[],[f4320_D])).
fof(f4320_D,plain,(
  ( ! [X0,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | sP4(X0,X1)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4363,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X3,X4) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl4),
  inference(subsumption_resolution,[],[f4348,f309])).
fof(f4348,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl4),
  inference(resolution,[],[f4264,f271])).
fof(f4264,plain,(
  ( ! [X59,X60,X58] : (sP4(X58,sK8) | sP4(X58,X59) | sP5(X59,X60)) ) | $spl4),
  inference(subsumption_resolution,[],[f4263,f361])).
fof(f4263,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | sP4(X58,sK8) | sP4(X58,X59) | sP5(X59,X60)) ) | $spl4),
  inference(forward_demodulation,[],[f4231,f900])).
fof(f4231,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,sK8)) | sP4(X58,sK8) | sP4(X58,X59) | sP5(X59,X60)) ) | $spl4),
  inference(resolution,[],[f1312,f2130])).
fof(f4362,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f4347,f309])).
fof(f4347,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f4264,f275])).
fof(f4321,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl4),
  inference(subsumption_resolution,[],[f4305,f309])).
fof(f4305,plain,(
  ( ! [X4,X5,X3] : (sP4(X3,sK8) | sP4(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl4),
  inference(resolution,[],[f4262,f271])).
fof(f4262,plain,(
  ( ! [X57,X56,X55] : (sP4(X55,sK8) | sP4(X55,X56) | sP5(X57,X56)) ) | $spl4),
  inference(subsumption_resolution,[],[f4261,f361])).
fof(f4261,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | sP4(X55,sK8) | sP4(X55,X56) | sP5(X57,X56)) ) | $spl4),
  inference(forward_demodulation,[],[f4230,f900])).
fof(f4230,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,sK8)) | sP4(X55,sK8) | sP4(X55,X56) | sP5(X57,X56)) ) | $spl4),
  inference(resolution,[],[f1312,f2129])).
fof(f4319,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl4),
  inference(subsumption_resolution,[],[f4304,f309])).
fof(f4304,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,sK8) | sP4(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl4),
  inference(resolution,[],[f4262,f275])).
fof(f3807,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) | $spl116),
  inference(cnf_transformation,[],[f3807_D])).
fof(f3807_D,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f3796,plain,(
  ( ! [X0] : (sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f990,f357])).
fof(f990,plain,(
  ( ! [X55] : (sP4(X55,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | p(s(bool,X55))) ) | $spl2),
  inference(resolution,[],[f366,f260])).
fof(f3761,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) | $spl114),
  inference(cnf_transformation,[],[f3761_D])).
fof(f3761_D,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f3750,plain,(
  ( ! [X0] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f989,f357])).
fof(f989,plain,(
  ( ! [X54] : (sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X54) | p(s(bool,X54))) ) | $spl2),
  inference(resolution,[],[f366,f257])).
fof(f3740,plain,(
  sP4('F_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl112),
  inference(cnf_transformation,[],[f3740_D])).
fof(f3740_D,plain,(
  sP4('F_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f3730,plain,(
  sP4('F_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl2 | $spl4)),
  inference(superposition,[],[f3697,f357])).
fof(f3697,plain,(
  sP4('F_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f3696,f361])).
fof(f3696,plain,(
  ~p(s(bool,'T_0')) | sP4('F_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f3673,f900])).
fof(f3673,plain,(
  ~p(s(bool,sK8)) | sP4('F_0','const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | ($spl2 | $spl4)),
  inference(resolution,[],[f988,f3415])).
fof(f3698,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0))) ) | $spl110),
  inference(cnf_transformation,[],[f3698_D])).
fof(f3698_D,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f3674,plain,(
  ( ! [X0] : (~sP4(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f988,f357])).
fof(f3657,plain,(
  ( ! [X0] : (~sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0))) ) | $spl108),
  inference(cnf_transformation,[],[f3657_D])).
fof(f3657_D,plain,(
  ( ! [X0] : (~sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f3647,plain,(
  ( ! [X0] : (~sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | ~p(s(bool,X0)) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f987,f357])).
fof(f987,plain,(
  ( ! [X52] : (~sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X52) | ~p(s(bool,X52))) ) | $spl2),
  inference(resolution,[],[f366,f256])).
fof(f3604,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP4('F_0','T_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f3589])).
fof(f3589,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP4('F_0','T_0') | ~p(s(bool,X28))) ) | $spl4),
  inference(superposition,[],[f3479,f325])).
fof(f3605,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP4('F_0','T_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f3588])).
fof(f3588,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP4('F_0','T_0') | ~p(s(bool,X27))) ) | $spl4),
  inference(superposition,[],[f3479,f327])).
fof(f3624,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP4('F_0','T_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f3587,f320])).
fof(f3587,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP4('F_0','T_0') | s(bool,'F_0') = s(bool,X26)) ) | $spl4),
  inference(superposition,[],[f3479,f358])).
fof(f3571,plain,(
  ( ! [X151,X150] : (sP4('F_0','F_0') | p(s(bool,X150)) | p(s(bool,X151)) | ~sP5(X150,X151)) ) | $spl4),
  inference(resolution,[],[f3479,f271])).
fof(f3569,plain,(
  ( ! [X146,X147] : (sP4('F_0','F_0') | p(s(bool,X146)) | p(s(bool,X147)) | ~sP5(X146,X147)) ) | $spl4),
  inference(resolution,[],[f3479,f275])).
fof(f3560,plain,(
  ( ! [X145] : (sP4('F_0','T_0') | ~p(s(bool,X145))) ) | $spl4),
  inference(resolution,[],[f3479,f318])).
fof(f3559,plain,(
  ( ! [X144] : (sP4('F_0','T_0') | ~p(s(bool,X144))) ) | $spl4),
  inference(resolution,[],[f3479,f313])).
fof(f3558,plain,(
  ( ! [X143] : (sP4('F_0','T_0') | ~p(s(bool,X143))) ) | $spl4),
  inference(resolution,[],[f3479,f303])).
fof(f3557,plain,(
  ( ! [X142] : (sP4('F_0','T_0') | p(s(bool,X142))) ) | $spl4),
  inference(resolution,[],[f3479,f302])).
fof(f3556,plain,(
  ( ! [X141] : (sP4('F_0','T_0') | p(s(bool,X141))) ) | $spl4),
  inference(resolution,[],[f3479,f299])).
fof(f3555,plain,(
  ( ! [X140] : (sP4('F_0','T_0') | p(s(bool,X140))) ) | $spl4),
  inference(resolution,[],[f3479,f296])).
fof(f3554,plain,(
  ( ! [X139] : (sP4('F_0','T_0') | ~p(s(bool,X139))) ) | $spl4),
  inference(resolution,[],[f3479,f294])).
fof(f3548,plain,(
  ( ! [X132,X130,X131] : (sP4('F_0',X130) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl4),
  inference(resolution,[],[f3479,f360])).
fof(f3547,plain,(
  ( ! [X127,X128,X129] : (sP4('F_0',X127) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl4),
  inference(resolution,[],[f3479,f360])).
fof(f3503,plain,(
  ( ! [X30,X28,X31,X29] : (sP4('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f3479,f225])).
fof(f3502,plain,(
  ( ! [X26,X24,X27,X25] : (sP4('F_0',X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl4),
  inference(resolution,[],[f3479,f224])).
fof(f3610,plain,(
  ( ! [X23,X21,X20] : (sP4('F_0',X20) | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl4),
  inference(subsumption_resolution,[],[f3501,f360])).
fof(f3501,plain,(
  ( ! [X23,X21,X22,X20] : (sP4('F_0',X20) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f3479,f224])).
fof(f3500,plain,(
  ( ! [X19,X17,X18,X16] : (sP4('F_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f3479,f223])).
fof(f3609,plain,(
  ( ! [X4,X7] : (sP4('F_0',X4) | s(bool,X4) != s(bool,X7)) ) | $spl102),
  inference(cnf_transformation,[],[f3609_D])).
fof(f3609_D,plain,(
  ( ! [X4,X7] : (sP4('F_0',X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f3499,plain,(
  ( ! [X14,X12,X15,X13] : (sP4('F_0',X12) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f3479,f223])).
fof(f3498,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl4),
  inference(resolution,[],[f3479,f221])).
fof(f3497,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('F_0',X4) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl4),
  inference(resolution,[],[f3479,f219])).
fof(f3496,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f3479,f219])).
fof(f3451,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4(sK8,X51) | sP4('F_0',X51)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f3438])).
fof(f3438,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4(sK8,X51) | sP4('F_0',X51) | p(s(bool,X50))) ) | $spl4),
  inference(superposition,[],[f1194,f321])).
fof(f3452,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4(sK8,X49) | sP4('F_0',X49)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f3437])).
fof(f3437,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4(sK8,X49) | sP4('F_0',X49) | p(s(bool,X48))) ) | $spl4),
  inference(superposition,[],[f1194,f323])).
fof(f3476,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4(sK8,X47) | sP4('F_0',X47)) ) | $spl4),
  inference(subsumption_resolution,[],[f3436,f324])).
fof(f3436,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4(sK8,X47) | sP4('F_0',X47) | s(bool,'T_0') = s(bool,X46)) ) | $spl4),
  inference(superposition,[],[f1194,f358])).
fof(f3414,plain,(
  ( ! [X198,X199] : (sP4(sK8,X198) | sP4('F_0',X198) | p(s(bool,X199))) ) | $spl4),
  inference(resolution,[],[f1194,f333])).
fof(f3413,plain,(
  ( ! [X196,X197] : (sP4(sK8,X196) | sP4('F_0',X196) | p(s(bool,X197))) ) | $spl4),
  inference(resolution,[],[f1194,f316])).
fof(f3412,plain,(
  ( ! [X194,X195] : (sP4(sK8,X194) | sP4('F_0',X194) | p(s(bool,X195))) ) | $spl4),
  inference(resolution,[],[f1194,f309])).
fof(f3469,plain,(
  ( ! [X192,X193] : (sP4(sK8,X192) | sP4('F_0',X192) | p(s(bool,X193))) ) | $spl4),
  inference(subsumption_resolution,[],[f3411,f361])).
fof(f3411,plain,(
  ( ! [X192,X193] : (sP4(sK8,X192) | sP4('F_0',X192) | p(s(bool,X193)) | ~p(s(bool,'T_0'))) ) | $spl4),
  inference(resolution,[],[f1194,f295])).
fof(f3410,plain,(
  ( ! [X191,X190] : (sP4(sK8,X190) | sP4('F_0',X190) | ~p(s(bool,X191))) ) | $spl4),
  inference(resolution,[],[f1194,f291])).
fof(f3409,plain,(
  ( ! [X189,X187,X188] : (sP4(sK8,X187) | sP4('F_0',X187) | sP5(X188,X189)) ) | $spl4),
  inference(resolution,[],[f1194,f279])).
fof(f3408,plain,(
  ( ! [X185,X186,X184] : (sP4(sK8,X184) | sP4('F_0',X184) | sP5(X185,X186)) ) | $spl4),
  inference(resolution,[],[f1194,f274])).
fof(f3464,plain,(
  ( ! [X0,X3,X1] : (sP4(sK8,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) | $spl94),
  inference(cnf_transformation,[],[f3464_D])).
fof(f3464_D,plain,(
  ( ! [X0,X3,X1] : (sP4(sK8,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X1,X0)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3403,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(sK8,X174) | sP4(X175,X174) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl4),
  inference(resolution,[],[f1194,f360])).
fof(f3361,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(sK8,X35) | sP4(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f1194,f225])).
fof(f3465,plain,(
  ( ! [X6,X5,X9] : (sP4(sK8,X5) | s(bool,X6) != s(bool,X9) | sP4(X6,X5)) ) | $spl96),
  inference(cnf_transformation,[],[f3465_D])).
fof(f3465_D,plain,(
  ( ! [X6,X5,X9] : (sP4(sK8,X5) | s(bool,X6) != s(bool,X9) | sP4(X6,X5)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f3360,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(sK8,X30) | sP4(X31,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) ) | $spl4),
  inference(resolution,[],[f1194,f225])).
fof(f3358,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(sK8,X20) | sP4(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f1194,f224])).
fof(f3357,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(sK8,X15) | sP4(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f1194,f223])).
fof(f3356,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(sK8,X10) | sP4(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f1194,f221])).
fof(f3355,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(sK8,X5) | sP4(X6,X5) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f1194,f221])).
fof(f3354,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(sK8,X0) | sP4(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl4),
  inference(resolution,[],[f1194,f219])).
fof(f3222,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'F_0') | $spl92),
  inference(cnf_transformation,[],[f3222_D])).
fof(f3222_D,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'F_0') <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3212,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),'F_0') | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6)))) | ($spl2 | $spl4)),
  inference(superposition,[],[f3018,f357])).
fof(f3018,plain,(
  sP4('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),'F_0') | ($spl2 | $spl4)),
  inference(resolution,[],[f2221,f366])).
fof(f3135,plain,(
  $false | $spl24),
  inference(equality_resolution,[],[f730])).
fof(f3050,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP4('T_0','F_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f3035])).
fof(f3035,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | sP4('T_0','F_0') | ~p(s(bool,X28))) ) | $spl4),
  inference(superposition,[],[f2221,f325])).
fof(f3051,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP4('T_0','F_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f3034])).
fof(f3034,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP4('T_0','F_0') | ~p(s(bool,X27))) ) | $spl4),
  inference(superposition,[],[f2221,f327])).
fof(f3073,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP4('T_0','F_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f3033,f320])).
fof(f3033,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP4('T_0','F_0') | s(bool,'F_0') = s(bool,X26)) ) | $spl4),
  inference(superposition,[],[f2221,f358])).
fof(f3017,plain,(
  ( ! [X149,X148] : (sP4('F_0','F_0') | p(s(bool,X148)) | p(s(bool,X149)) | ~sP5(X148,X149)) ) | $spl4),
  inference(resolution,[],[f2221,f271])).
fof(f3015,plain,(
  ( ! [X144,X145] : (sP4('F_0','F_0') | p(s(bool,X144)) | p(s(bool,X145)) | ~sP5(X144,X145)) ) | $spl4),
  inference(resolution,[],[f2221,f275])).
fof(f3006,plain,(
  ( ! [X143] : (sP4('T_0','F_0') | ~p(s(bool,X143))) ) | $spl4),
  inference(resolution,[],[f2221,f318])).
fof(f3005,plain,(
  ( ! [X142] : (sP4('T_0','F_0') | ~p(s(bool,X142))) ) | $spl4),
  inference(resolution,[],[f2221,f313])).
fof(f3004,plain,(
  ( ! [X141] : (sP4('T_0','F_0') | ~p(s(bool,X141))) ) | $spl4),
  inference(resolution,[],[f2221,f303])).
fof(f3003,plain,(
  ( ! [X140] : (sP4('T_0','F_0') | p(s(bool,X140))) ) | $spl4),
  inference(resolution,[],[f2221,f302])).
fof(f3002,plain,(
  ( ! [X139] : (sP4('T_0','F_0') | p(s(bool,X139))) ) | $spl4),
  inference(resolution,[],[f2221,f299])).
fof(f3001,plain,(
  ( ! [X138] : (sP4('T_0','F_0') | p(s(bool,X138))) ) | $spl4),
  inference(resolution,[],[f2221,f296])).
fof(f3000,plain,(
  ( ! [X137] : (sP4('T_0','F_0') | ~p(s(bool,X137))) ) | $spl4),
  inference(resolution,[],[f2221,f294])).
fof(f2995,plain,(
  ( ! [X132,X130,X131] : (sP4(X130,'F_0') | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X130) != s(bool,X132)) ) | $spl4),
  inference(resolution,[],[f2221,f360])).
fof(f2994,plain,(
  ( ! [X127,X128,X129] : (sP4(X127,'F_0') | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X128) != s(bool,X129)) ) | $spl4),
  inference(resolution,[],[f2221,f360])).
fof(f2950,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f2221,f225])).
fof(f2949,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,'F_0') | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) ) | $spl4),
  inference(resolution,[],[f2221,f224])).
fof(f3056,plain,(
  ( ! [X23,X21,X20] : (sP4(X20,'F_0') | ~p(s(bool,X21)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl4),
  inference(subsumption_resolution,[],[f2948,f360])).
fof(f2948,plain,(
  ( ! [X23,X21,X22,X20] : (sP4(X20,'F_0') | ~p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X21) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f2221,f224])).
fof(f2947,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f2221,f223])).
fof(f3055,plain,(
  ( ! [X4,X7] : (sP4(X4,'F_0') | s(bool,X4) != s(bool,X7)) ) | $spl82),
  inference(cnf_transformation,[],[f3055_D])).
fof(f3055_D,plain,(
  ( ! [X4,X7] : (sP4(X4,'F_0') | s(bool,X4) != s(bool,X7)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2946,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,'F_0') | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X12) != s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f2221,f223])).
fof(f2945,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl4),
  inference(resolution,[],[f2221,f221])).
fof(f2944,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'F_0') | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl4),
  inference(resolution,[],[f2221,f219])).
fof(f2943,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f2221,f219])).
fof(f2934,plain,(
  $false | $spl78),
  inference(trivial_inequality_removal,[],[f2921])).
fof(f2921,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl78),
  inference(superposition,[],[f2374,f216])).
fof(f2933,plain,(
  $false | $spl78),
  inference(trivial_inequality_removal,[],[f2923])).
fof(f2923,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl78),
  inference(superposition,[],[f2374,f359])).
fof(f359,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.NONE_0')))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X6),'const.option.NONE_0')))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.option.IS_SOME_DEF|split|1')).
fof(f2932,plain,(
  $false | $spl78),
  inference(equality_resolution,[],[f2374])).
fof(f2872,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2860])).
fof(f2860,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl50),
  inference(superposition,[],[f1670,f217])).
fof(f2871,plain,(
  $false | ($spl4 | $spl50)),
  inference(trivial_inequality_removal,[],[f2862])).
fof(f2862,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl50)),
  inference(superposition,[],[f1670,f900])).
fof(f2870,plain,(
  $false | $spl50),
  inference(equality_resolution,[],[f1670])).
fof(f2814,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2743,f2518])).
fof(f2518,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(equality_resolution,[],[f449])).
fof(f2743,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f2518,f337])).
fof(f337,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f203])).
fof(f2813,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2742,f2518])).
fof(f2742,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f2518,f335])).
fof(f335,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f202])).
fof(f2812,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2738,f2518])).
fof(f2738,plain,(
  ( ! [X131] : (p(s(bool,X131))) ) | $spl14),
  inference(resolution,[],[f2518,f302])).
fof(f2811,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2737,f2518])).
fof(f2737,plain,(
  ( ! [X130] : (p(s(bool,X130))) ) | $spl14),
  inference(resolution,[],[f2518,f299])).
fof(f2810,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2736,f2518])).
fof(f2736,plain,(
  ( ! [X129] : (p(s(bool,X129))) ) | $spl14),
  inference(resolution,[],[f2518,f296])).
fof(f2809,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2734,f2518])).
fof(f2734,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f2518,f293])).
fof(f293,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f181])).
fof(f2744,plain,(
  $false | $spl14),
  inference(resolution,[],[f2518,f339])).
fof(f2745,plain,(
  $false | $spl14),
  inference(resolution,[],[f2518,f340])).
fof(f2746,plain,(
  $false | $spl14),
  inference(resolution,[],[f2518,f361])).
fof(f2753,plain,(
  $false | ($spl2 | $spl14)),
  inference(resolution,[],[f2518,f366])).
fof(f2673,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2602,f451])).
fof(f2602,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f451,f337])).
fof(f2672,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2601,f451])).
fof(f2601,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f451,f335])).
fof(f2671,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2597,f451])).
fof(f2597,plain,(
  ( ! [X131] : (p(s(bool,X131))) ) | $spl16),
  inference(resolution,[],[f451,f302])).
fof(f2670,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2596,f451])).
fof(f2596,plain,(
  ( ! [X130] : (p(s(bool,X130))) ) | $spl16),
  inference(resolution,[],[f451,f299])).
fof(f2669,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2595,f451])).
fof(f2595,plain,(
  ( ! [X129] : (p(s(bool,X129))) ) | $spl16),
  inference(resolution,[],[f451,f296])).
fof(f2668,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2593,f451])).
fof(f2593,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f451,f293])).
fof(f2603,plain,(
  $false | $spl16),
  inference(resolution,[],[f451,f339])).
fof(f2604,plain,(
  $false | $spl16),
  inference(resolution,[],[f451,f340])).
fof(f2605,plain,(
  $false | $spl16),
  inference(resolution,[],[f451,f361])).
fof(f2612,plain,(
  $false | ($spl2 | $spl16)),
  inference(resolution,[],[f451,f366])).
fof(f2524,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl14),
  inference(subsumption_resolution,[],[f2469,f449])).
fof(f2469,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl14),
  inference(superposition,[],[f449,f358])).
fof(f2523,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl14),
  inference(subsumption_resolution,[],[f2468,f449])).
fof(f2468,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl14),
  inference(superposition,[],[f449,f358])).
fof(f2519,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2467])).
fof(f2467,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl14),
  inference(superposition,[],[f449,f327])).
fof(f2520,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2466])).
fof(f2466,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl14),
  inference(superposition,[],[f449,f325])).
fof(f2445,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f2432])).
fof(f2432,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl8),
  inference(superposition,[],[f2334,f216])).
fof(f2334,plain,(
  ( ! [X198] : (s(bool,'F_0') != s(bool,X198)) ) | $spl8),
  inference(resolution,[],[f446,f334])).
fof(f2444,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f2434])).
fof(f2434,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl8),
  inference(superposition,[],[f2334,f359])).
fof(f2443,plain,(
  $false | $spl8),
  inference(equality_resolution,[],[f2334])).
fof(f2368,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | s(bool,'F_0') != s(bool,X51)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2356])).
fof(f2356,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | s(bool,'F_0') != s(bool,X51) | p(s(bool,X50))) ) | $spl8),
  inference(superposition,[],[f446,f321])).
fof(f2369,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') != s(bool,X49)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2355])).
fof(f2355,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') != s(bool,X49) | p(s(bool,X48))) ) | $spl8),
  inference(superposition,[],[f446,f323])).
fof(f2381,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,'F_0') != s(bool,X47)) ) | $spl8),
  inference(subsumption_resolution,[],[f2354,f324])).
fof(f2354,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,'F_0') != s(bool,X47) | s(bool,'T_0') = s(bool,X46)) ) | $spl8),
  inference(superposition,[],[f446,f358])).
fof(f2333,plain,(
  ( ! [X196,X197] : (s(bool,'F_0') != s(bool,X196) | p(s(bool,X197))) ) | $spl8),
  inference(resolution,[],[f446,f333])).
fof(f2332,plain,(
  ( ! [X194,X195] : (s(bool,'F_0') != s(bool,X194) | p(s(bool,X195))) ) | $spl8),
  inference(resolution,[],[f446,f316])).
fof(f2331,plain,(
  ( ! [X192,X193] : (s(bool,'F_0') != s(bool,X192) | p(s(bool,X193))) ) | $spl8),
  inference(resolution,[],[f446,f309])).
fof(f2375,plain,(
  ( ! [X191,X190] : (s(bool,'F_0') != s(bool,X190) | p(s(bool,X191))) ) | $spl8),
  inference(subsumption_resolution,[],[f2330,f361])).
fof(f2330,plain,(
  ( ! [X191,X190] : (s(bool,'F_0') != s(bool,X190) | p(s(bool,X191)) | ~p(s(bool,'T_0'))) ) | $spl8),
  inference(resolution,[],[f446,f295])).
fof(f2329,plain,(
  ( ! [X189,X188] : (s(bool,'F_0') != s(bool,X188) | ~p(s(bool,X189))) ) | $spl8),
  inference(resolution,[],[f446,f291])).
fof(f2328,plain,(
  ( ! [X187,X185,X186] : (s(bool,'F_0') != s(bool,X185) | sP5(X186,X187)) ) | $spl8),
  inference(resolution,[],[f446,f279])).
fof(f2327,plain,(
  ( ! [X182,X184,X183] : (s(bool,'F_0') != s(bool,X182) | sP5(X183,X184)) ) | $spl8),
  inference(resolution,[],[f446,f274])).
fof(f2273,plain,(
  $false | ($spl4 | $spl74)),
  inference(subsumption_resolution,[],[f2272,f340])).
fof(f2272,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl74)),
  inference(forward_demodulation,[],[f2254,f900])).
fof(f2254,plain,(
  ~p(s(bool,sK8)) | $spl74),
  inference(duplicate_literal_removal,[],[f2253])).
fof(f2253,plain,(
  ~p(s(bool,sK8)) | ~p(s(bool,sK8)) | $spl74),
  inference(resolution,[],[f2238,f256])).
fof(f2269,plain,(
  $false | ($spl4 | $spl74)),
  inference(subsumption_resolution,[],[f2268,f340])).
fof(f2268,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl74)),
  inference(forward_demodulation,[],[f2257,f900])).
fof(f2257,plain,(
  ~p(s(bool,sK8)) | $spl74),
  inference(duplicate_literal_removal,[],[f2250])).
fof(f2250,plain,(
  ~p(s(bool,sK8)) | ~p(s(bool,sK8)) | $spl74),
  inference(resolution,[],[f2238,f263])).
fof(f2267,plain,(
  $false | ($spl4 | $spl74)),
  inference(subsumption_resolution,[],[f2266,f340])).
fof(f2266,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl74)),
  inference(forward_demodulation,[],[f2265,f900])).
fof(f2265,plain,(
  ~p(s(bool,sK8)) | ($spl4 | $spl74)),
  inference(subsumption_resolution,[],[f2249,f879])).
fof(f2249,plain,(
  ~p(s(bool,sK8)) | sP4(sK8,sK8) | ($spl4 | $spl74)),
  inference(resolution,[],[f2238,f1191])).
fof(f2264,plain,(
  $false | ($spl4 | $spl74)),
  inference(subsumption_resolution,[],[f2263,f340])).
fof(f2263,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl74)),
  inference(forward_demodulation,[],[f2262,f900])).
fof(f2262,plain,(
  ~p(s(bool,sK8)) | ($spl4 | $spl74)),
  inference(subsumption_resolution,[],[f2248,f879])).
fof(f2248,plain,(
  ~p(s(bool,sK8)) | sP4(sK8,sK8) | ($spl4 | $spl74)),
  inference(resolution,[],[f2238,f1192])).
fof(f2261,plain,(
  $false | ($spl4 | $spl74)),
  inference(subsumption_resolution,[],[f2260,f340])).
fof(f2260,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl74)),
  inference(forward_demodulation,[],[f2247,f900])).
fof(f2247,plain,(
  ~p(s(bool,sK8)) | ($spl4 | $spl74)),
  inference(resolution,[],[f2238,f880])).
fof(f2259,plain,(
  $false | ($spl4 | $spl74)),
  inference(subsumption_resolution,[],[f2258,f340])).
fof(f2258,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl74)),
  inference(forward_demodulation,[],[f2246,f900])).
fof(f2246,plain,(
  ~p(s(bool,sK8)) | ($spl4 | $spl74)),
  inference(resolution,[],[f2238,f879])).
fof(f2238,plain,(
  sP4(sK8,sK8) | $spl74),
  inference(cnf_transformation,[],[f2238_D])).
fof(f2238_D,plain,(
  sP4(sK8,sK8) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2223,plain,(
  ( ! [X2] : (sP4(X2,'F_0') | ~p(s(bool,X2)) | sP4(sK8,sK8)) ) | $spl4),
  inference(resolution,[],[f2159,f1191])).
fof(f2194,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4(sK8,X51) | sP4(X51,'F_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f2182])).
fof(f2182,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | sP4(sK8,X51) | sP4(X51,'F_0') | p(s(bool,X50))) ) | $spl4),
  inference(superposition,[],[f1193,f321])).
fof(f2195,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4(sK8,X49) | sP4(X49,'F_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f2181])).
fof(f2181,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP4(sK8,X49) | sP4(X49,'F_0') | p(s(bool,X48))) ) | $spl4),
  inference(superposition,[],[f1193,f323])).
fof(f2218,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4(sK8,X47) | sP4(X47,'F_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f2180,f324])).
fof(f2180,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP4(sK8,X47) | sP4(X47,'F_0') | s(bool,'T_0') = s(bool,X46)) ) | $spl4),
  inference(superposition,[],[f1193,f358])).
fof(f2158,plain,(
  ( ! [X196,X197] : (sP4(sK8,X196) | sP4(X196,'F_0') | p(s(bool,X197))) ) | $spl4),
  inference(resolution,[],[f1193,f333])).
fof(f2157,plain,(
  ( ! [X194,X195] : (sP4(sK8,X194) | sP4(X194,'F_0') | p(s(bool,X195))) ) | $spl4),
  inference(resolution,[],[f1193,f316])).
fof(f2156,plain,(
  ( ! [X192,X193] : (sP4(sK8,X192) | sP4(X192,'F_0') | p(s(bool,X193))) ) | $spl4),
  inference(resolution,[],[f1193,f309])).
fof(f2212,plain,(
  ( ! [X191,X190] : (sP4(sK8,X190) | sP4(X190,'F_0') | p(s(bool,X191))) ) | $spl4),
  inference(subsumption_resolution,[],[f2155,f361])).
fof(f2155,plain,(
  ( ! [X191,X190] : (sP4(sK8,X190) | sP4(X190,'F_0') | p(s(bool,X191)) | ~p(s(bool,'T_0'))) ) | $spl4),
  inference(resolution,[],[f1193,f295])).
fof(f2154,plain,(
  ( ! [X189,X188] : (sP4(sK8,X188) | sP4(X188,'F_0') | ~p(s(bool,X189))) ) | $spl4),
  inference(resolution,[],[f1193,f291])).
fof(f2153,plain,(
  ( ! [X187,X185,X186] : (sP4(sK8,X185) | sP4(X185,'F_0') | sP5(X186,X187)) ) | $spl4),
  inference(resolution,[],[f1193,f279])).
fof(f2152,plain,(
  ( ! [X182,X184,X183] : (sP4(sK8,X182) | sP4(X182,'F_0') | sP5(X183,X184)) ) | $spl4),
  inference(resolution,[],[f1193,f274])).
fof(f2207,plain,(
  ( ! [X0,X3,X1] : (sP4(sK8,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) | $spl68),
  inference(cnf_transformation,[],[f2207_D])).
fof(f2207_D,plain,(
  ( ! [X0,X3,X1] : (sP4(sK8,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP4(X0,X1)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2148,plain,(
  ( ! [X177,X175,X176,X174] : (sP4(sK8,X174) | sP4(X174,X175) | p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X175) != s(bool,X177)) ) | $spl4),
  inference(resolution,[],[f1193,f360])).
fof(f2106,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP4(sK8,X35) | sP4(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f1193,f225])).
fof(f2208,plain,(
  ( ! [X6,X5,X9] : (sP4(sK8,X5) | s(bool,X6) != s(bool,X9) | sP4(X5,X6)) ) | $spl70),
  inference(cnf_transformation,[],[f2208_D])).
fof(f2208_D,plain,(
  ( ! [X6,X5,X9] : (sP4(sK8,X5) | s(bool,X6) != s(bool,X9) | sP4(X5,X6)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2105,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP4(sK8,X30) | sP4(X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) ) | $spl4),
  inference(resolution,[],[f1193,f225])).
fof(f2103,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP4(sK8,X20) | sP4(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f1193,f224])).
fof(f2102,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP4(sK8,X15) | sP4(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f1193,f223])).
fof(f2101,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP4(sK8,X10) | sP4(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f1193,f221])).
fof(f2100,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP4(sK8,X5) | sP4(X5,X6) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f1193,f221])).
fof(f2099,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(sK8,X0) | sP4(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl4),
  inference(resolution,[],[f1193,f219])).
fof(f2060,plain,(
  ( ! [X0,X1] : (sP0('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1)) ) | $spl66),
  inference(cnf_transformation,[],[f2060_D])).
fof(f2060_D,plain,(
  ( ! [X0,X1] : (sP0('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2050,plain,(
  ( ! [X0,X1] : (sP0('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0,X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f984,f357])).
fof(f984,plain,(
  ( ! [X47,X46] : (sP0('const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X46,X47)) ) | $spl2),
  inference(resolution,[],[f366,f244])).
fof(f2045,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1)) ) | $spl64),
  inference(cnf_transformation,[],[f2045_D])).
fof(f2045_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2035,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X1) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f983,f357])).
fof(f983,plain,(
  ( ! [X45,X44] : (sP0(X44,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)),X45)) ) | $spl2),
  inference(resolution,[],[f366,f243])).
fof(f2016,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) | $spl62),
  inference(cnf_transformation,[],[f2016_D])).
fof(f2016_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2006,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f982,f357])).
fof(f982,plain,(
  ( ! [X43,X42] : (sP0(X42,X43,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f242])).
fof(f1972,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) | $spl60),
  inference(cnf_transformation,[],[f1972_D])).
fof(f1972_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1963,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f981,f357])).
fof(f981,plain,(
  ( ! [X41,X40] : (sP1(X40,X41,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f238])).
fof(f1914,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) | $spl58),
  inference(cnf_transformation,[],[f1914_D])).
fof(f1914_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1905,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f978,f357])).
fof(f978,plain,(
  ( ! [X35,X34] : (sP2(X34,X35,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f234])).
fof(f1837,plain,(
  ( ! [X0] : (sP5('const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')),X0) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f996,f357])).
fof(f1834,plain,(
  ( ! [X0] : (sP5(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) | $spl54),
  inference(cnf_transformation,[],[f1834_D])).
fof(f1834_D,plain,(
  ( ! [X0] : (sP5(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0')))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1823,plain,(
  ( ! [X0] : (sP5(X0,'const.option.IS_NONE_1'(s('type.option.option'(sK6),'const.option.NONE_0'))) | s('type.option.option'(sK6),sK7) = s('type.option.option'(sK6),'const.option.SOME_1'(s(sK6,sK9(sK7,sK6))))) ) | $spl2),
  inference(superposition,[],[f995,f357])).
fof(f1691,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97)) ) | $spl4),
  inference(forward_demodulation,[],[f1690,f900])).
fof(f1690,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,sK8) != s(bool,X98)) ) | $spl4),
  inference(subsumption_resolution,[],[f1496,f360])).
fof(f1496,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,sK8) != s(bool,X98)) ) | $spl4),
  inference(superposition,[],[f225,f900])).
fof(f1687,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl4),
  inference(subsumption_resolution,[],[f1492,f223])).
fof(f1492,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl4),
  inference(superposition,[],[f225,f900])).
fof(f1686,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl4),
  inference(subsumption_resolution,[],[f1491,f360])).
fof(f1491,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,sK8)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl4),
  inference(superposition,[],[f225,f900])).
fof(f1675,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X38) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl4),
  inference(forward_demodulation,[],[f1674,f900])).
fof(f1674,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,sK8) != s(bool,X38)) ) | $spl4),
  inference(subsumption_resolution,[],[f1476,f360])).
fof(f1476,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,sK8) != s(bool,X38)) ) | $spl4),
  inference(superposition,[],[f221,f900])).
fof(f1671,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl4),
  inference(subsumption_resolution,[],[f1472,f360])).
fof(f1472,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,sK8)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl4),
  inference(superposition,[],[f221,f900])).
fof(f1669,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(subsumption_resolution,[],[f1471,f223])).
fof(f1471,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(superposition,[],[f221,f900])).
fof(f1391,plain,(
  sP4('F_0',sK8) | $spl48),
  inference(cnf_transformation,[],[f1391_D])).
fof(f1391_D,plain,(
  sP4('F_0',sK8) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1382,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('F_0',sK8)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1371])).
fof(f1371,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4('F_0',sK8) | p(s(bool,X31))) ) | $spl4),
  inference(superposition,[],[f882,f321])).
fof(f1383,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('F_0',sK8)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1370])).
fof(f1370,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4('F_0',sK8) | p(s(bool,X30))) ) | $spl4),
  inference(superposition,[],[f882,f323])).
fof(f1398,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('F_0',sK8)) ) | $spl4),
  inference(subsumption_resolution,[],[f1369,f324])).
fof(f1369,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4('F_0',sK8) | s(bool,'T_0') = s(bool,X29)) ) | $spl4),
  inference(superposition,[],[f882,f358])).
fof(f1348,plain,(
  ( ! [X137] : (sP4('F_0',sK8) | p(s(bool,X137))) ) | $spl4),
  inference(resolution,[],[f882,f333])).
fof(f1347,plain,(
  ( ! [X136] : (sP4('F_0',sK8) | p(s(bool,X136))) ) | $spl4),
  inference(resolution,[],[f882,f316])).
fof(f1346,plain,(
  ( ! [X135] : (sP4('F_0',sK8) | p(s(bool,X135))) ) | $spl4),
  inference(resolution,[],[f882,f309])).
fof(f1392,plain,(
  ( ! [X134] : (sP4('F_0',sK8) | p(s(bool,X134))) ) | $spl4),
  inference(subsumption_resolution,[],[f1345,f361])).
fof(f1345,plain,(
  ( ! [X134] : (sP4('F_0',sK8) | p(s(bool,X134)) | ~p(s(bool,'T_0'))) ) | $spl4),
  inference(resolution,[],[f882,f295])).
fof(f1344,plain,(
  ( ! [X133] : (sP4('F_0',sK8) | ~p(s(bool,X133))) ) | $spl4),
  inference(resolution,[],[f882,f291])).
fof(f1343,plain,(
  ( ! [X132,X131] : (sP4('F_0',sK8) | sP5(X131,X132)) ) | $spl4),
  inference(resolution,[],[f882,f279])).
fof(f1342,plain,(
  ( ! [X130,X129] : (sP4('F_0',sK8) | sP5(X129,X130)) ) | $spl4),
  inference(resolution,[],[f882,f274])).
fof(f1386,plain,(
  ( ! [X2,X0] : (sP4(X0,sK8) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl44),
  inference(cnf_transformation,[],[f1386_D])).
fof(f1386_D,plain,(
  ( ! [X2,X0] : (sP4(X0,sK8) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1338,plain,(
  ( ! [X127,X125,X126] : (sP4(X125,sK8) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl4),
  inference(resolution,[],[f882,f360])).
fof(f1296,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(X28,sK8) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f882,f225])).
fof(f1387,plain,(
  ( ! [X4,X7] : (sP4(X4,sK8) | s(bool,X4) != s(bool,X7)) ) | $spl46),
  inference(cnf_transformation,[],[f1387_D])).
fof(f1387_D,plain,(
  ( ! [X4,X7] : (sP4(X4,sK8) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1295,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(X24,sK8) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl4),
  inference(resolution,[],[f882,f225])).
fof(f1293,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(X16,sK8) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f882,f224])).
fof(f1292,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(X12,sK8) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f882,f223])).
fof(f1291,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,sK8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl4),
  inference(resolution,[],[f882,f221])).
fof(f1290,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,sK8) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl4),
  inference(resolution,[],[f882,f221])).
fof(f1289,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,sK8) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f882,f219])).
fof(f1272,plain,(
  sP4(sK8,'F_0') | $spl42),
  inference(cnf_transformation,[],[f1272_D])).
fof(f1272_D,plain,(
  sP4(sK8,'F_0') <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1262,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4(sK8,'F_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1251])).
fof(f1251,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP4(sK8,'F_0') | p(s(bool,X31))) ) | $spl4),
  inference(superposition,[],[f881,f321])).
fof(f1263,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4(sK8,'F_0')) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1250])).
fof(f1250,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP4(sK8,'F_0') | p(s(bool,X30))) ) | $spl4),
  inference(superposition,[],[f881,f323])).
fof(f1279,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4(sK8,'F_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f1249,f324])).
fof(f1249,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP4(sK8,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl4),
  inference(superposition,[],[f881,f358])).
fof(f1228,plain,(
  ( ! [X137] : (sP4(sK8,'F_0') | p(s(bool,X137))) ) | $spl4),
  inference(resolution,[],[f881,f333])).
fof(f1227,plain,(
  ( ! [X136] : (sP4(sK8,'F_0') | p(s(bool,X136))) ) | $spl4),
  inference(resolution,[],[f881,f316])).
fof(f1226,plain,(
  ( ! [X135] : (sP4(sK8,'F_0') | p(s(bool,X135))) ) | $spl4),
  inference(resolution,[],[f881,f309])).
fof(f1273,plain,(
  ( ! [X134] : (sP4(sK8,'F_0') | p(s(bool,X134))) ) | $spl4),
  inference(subsumption_resolution,[],[f1225,f361])).
fof(f1225,plain,(
  ( ! [X134] : (sP4(sK8,'F_0') | p(s(bool,X134)) | ~p(s(bool,'T_0'))) ) | $spl4),
  inference(resolution,[],[f881,f295])).
fof(f1224,plain,(
  ( ! [X133] : (sP4(sK8,'F_0') | ~p(s(bool,X133))) ) | $spl4),
  inference(resolution,[],[f881,f291])).
fof(f1223,plain,(
  ( ! [X132,X131] : (sP4(sK8,'F_0') | sP5(X131,X132)) ) | $spl4),
  inference(resolution,[],[f881,f279])).
fof(f1222,plain,(
  ( ! [X130,X129] : (sP4(sK8,'F_0') | sP5(X129,X130)) ) | $spl4),
  inference(resolution,[],[f881,f274])).
fof(f1266,plain,(
  ( ! [X2,X0] : (sP4(sK8,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl36),
  inference(cnf_transformation,[],[f1266_D])).
fof(f1266_D,plain,(
  ( ! [X2,X0] : (sP4(sK8,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1218,plain,(
  ( ! [X127,X125,X126] : (sP4(sK8,X125) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X125) != s(bool,X127)) ) | $spl4),
  inference(resolution,[],[f881,f360])).
fof(f1176,plain,(
  ( ! [X30,X28,X31,X29] : (sP4(sK8,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f881,f225])).
fof(f1267,plain,(
  ( ! [X4,X7] : (sP4(sK8,X4) | s(bool,X4) != s(bool,X7)) ) | $spl38),
  inference(cnf_transformation,[],[f1267_D])).
fof(f1267_D,plain,(
  ( ! [X4,X7] : (sP4(sK8,X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1175,plain,(
  ( ! [X26,X24,X27,X25] : (sP4(sK8,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl4),
  inference(resolution,[],[f881,f225])).
fof(f1173,plain,(
  ( ! [X19,X17,X18,X16] : (sP4(sK8,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f881,f224])).
fof(f1172,plain,(
  ( ! [X14,X12,X15,X13] : (sP4(sK8,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f881,f223])).
fof(f1171,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(sK8,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl4),
  inference(resolution,[],[f881,f221])).
fof(f1170,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(sK8,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl4),
  inference(resolution,[],[f881,f221])).
fof(f1169,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK8,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f881,f219])).
fof(f1150,plain,(
  $false | $spl10),
  inference(equality_resolution,[],[f447])).
fof(f1102,plain,(
  s(bool,'T_0') = s(bool,sK8) | $spl32),
  inference(cnf_transformation,[],[f1102_D])).
fof(f1102_D,plain,(
  s(bool,'T_0') = s(bool,sK8) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1104,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl35),
  inference(cnf_transformation,[],[f1104_D])).
fof(f1104_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f1063,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'F_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f365,f358])).
fof(f1062,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | s(bool,'F_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f365,f358])).
fof(f1044,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl26),
  inference(cnf_transformation,[],[f1044_D])).
fof(f1044_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1014,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f360])).
fof(f972,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X22) != s(bool,X23)) ) | $spl2),
  inference(resolution,[],[f366,f225])).
fof(f971,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f225])).
fof(f969,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f224])).
fof(f968,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f223])).
fof(f967,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X7) != s(bool,X8)) ) | $spl2),
  inference(resolution,[],[f366,f221])).
fof(f966,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f221])).
fof(f965,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl2),
  inference(resolution,[],[f366,f219])).
fof(f532,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK8) != s(bool,X1)) ) | $spl20),
  inference(cnf_transformation,[],[f532_D])).
fof(f532_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK8) != s(bool,X1)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f906,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | s(bool,sK8) != s(bool,X77)) ) | $spl4),
  inference(resolution,[],[f369,f360])).
fof(f864,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f369,f225])).
fof(f447,plain,(
  ( ! [X5] : (s(bool,sK8) != s(bool,X5)) ) | $spl10),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  ( ! [X5] : (s(bool,sK8) != s(bool,X5)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f863,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,sK8) != s(bool,X20)) ) | $spl4),
  inference(resolution,[],[f369,f225])).
fof(f861,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK8) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f369,f224])).
fof(f860,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK8) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f369,f223])).
fof(f859,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X6) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f369,f221])).
fof(f858,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,sK8) != s(bool,X5)) ) | $spl4),
  inference(resolution,[],[f369,f221])).
fof(f857,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl4),
  inference(resolution,[],[f369,f219])).
fof(f848,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f803,f452])).
fof(f803,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(resolution,[],[f452,f338])).
fof(f847,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f802,f452])).
fof(f802,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(resolution,[],[f452,f336])).
fof(f846,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f796,f452])).
fof(f796,plain,(
  ( ! [X127] : (~p(s(bool,X127))) ) | $spl18),
  inference(resolution,[],[f452,f291])).
fof(f801,plain,(
  $false | $spl18),
  inference(resolution,[],[f452,f334])).
fof(f804,plain,(
  $false | ($spl5 | $spl18)),
  inference(resolution,[],[f452,f368])).
fof(f699,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X79) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f360])).
fof(f654,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f367,f225])).
fof(f653,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f224])).
fof(f651,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f223])).
fof(f650,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X9) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f367,f223])).
fof(f649,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f367,f221])).
fof(f648,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f367,f219])).
fof(f647,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) ) | $spl3),
  inference(resolution,[],[f367,f219])).
fof(f589,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,sK8) != s(bool,X79)) ) | $spl5),
  inference(resolution,[],[f368,f360])).
fof(f544,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f368,f225])).
fof(f543,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,sK8) != s(bool,X19)) ) | $spl5),
  inference(resolution,[],[f368,f224])).
fof(f541,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,sK8) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f540,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,sK8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f539,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f538,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK8) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f509,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | s(bool,sK8) != s(bool,X77)) ) | $spl4),
  inference(resolution,[],[f369,f360])).
fof(f467,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl4),
  inference(resolution,[],[f369,f225])).
fof(f466,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,sK8) != s(bool,X20)) ) | $spl4),
  inference(resolution,[],[f369,f225])).
fof(f464,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK8) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f369,f224])).
fof(f463,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK8) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f369,f223])).
fof(f462,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X6) | s(bool,X7) != s(bool,X8)) ) | $spl4),
  inference(resolution,[],[f369,f221])).
fof(f461,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,sK8) != s(bool,X5)) ) | $spl4),
  inference(resolution,[],[f369,f221])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl4),
  inference(resolution,[],[f369,f219])).
fof(f422,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,sK8) != s(bool,X79)) ) | $spl5),
  inference(resolution,[],[f368,f360])).
fof(f377,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK8) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f368,f225])).
fof(f376,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,sK8) != s(bool,X19)) ) | $spl5),
  inference(resolution,[],[f368,f224])).
fof(f374,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,sK8) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f373,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,sK8) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | $spl5),
  inference(resolution,[],[f368,f223])).
fof(f372,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK8) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f368,f221])).
fof(f371,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK8) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK8) != s(bool,X1)) ) | $spl5),
  inference(resolution,[],[f368,f219])).
fof(f364,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl0),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f368,plain,(
  ~p(s(bool,sK8)) | $spl5),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  ~p(s(bool,sK8)) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f366,plain,(
  p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) | $spl2),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f213,plain,(
  p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) | ~p(s(bool,sK8)) | s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0')))),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  (s('type.option.option'(sK6),'const.option.NONE_0') = s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | ~p(s(bool,sK8)) | p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))))) & (s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | (p(s(bool,sK8)) & ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f147])).
fof(f147,plain,(
  ? [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),X1),s('type.option.option'(X0),'const.option.NONE_0'))) | ~p(s(bool,X2)) | p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1))))) & (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),X1),s('type.option.option'(X0),'const.option.NONE_0'))) | (p(s(bool,X2)) & ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1)))))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ? [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),X1),s('type.option.option'(X0),'const.option.NONE_0'))) | (~p(s(bool,X2)) | p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1)))))) & (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),X1),s('type.option.option'(X0),'const.option.NONE_0'))) | (p(s(bool,X2)) & ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1)))))))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ? [X0,X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),X1),s('type.option.option'(X0),'const.option.NONE_0'))) <~> (~p(s(bool,X2)) | p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1))))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),X1),s('type.option.option'(X0),'const.option.NONE_0'))) <=> (p(s(bool,X2)) => p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1))))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X2),s('type.option.option'(X0),X1),s('type.option.option'(X0),'const.option.NONE_0'))) <=> (p(s(bool,X2)) => p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(X6),X7),s('type.option.option'(X6),'const.option.NONE_0'))) <=> (p(s(bool,X8)) => p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X6),X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.bool.COND_3'(s(bool,X8),s('type.option.option'(X6),X7),s('type.option.option'(X6),'const.option.NONE_0'))) <=> (p(s(bool,X8)) => p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X6),X7))))))),
  file('thm.option.IF_NONE_EQUALS_OPTION|split|0','thm.option.IF_NONE_EQUALS_OPTION|split|0')).
fof(f365,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) | $spl1),
  inference(cnf_transformation,[],[f365_D])).
fof(f365_D,plain,(
  s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f369,plain,(
  p(s(bool,sK8)) | $spl4),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  p(s(bool,sK8)) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f214,plain,(
  p(s(bool,sK8)) | s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0')))),
  inference(cnf_transformation,[],[f148])).
fof(f367,plain,(
  ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) | $spl3),
  inference(cnf_transformation,[],[f367_D])).
fof(f367_D,plain,(
  ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f215,plain,(
  ~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK6),sK7)))) | s('type.option.option'(sK6),'const.option.NONE_0') != s('type.option.option'(sK6),'const.bool.COND_3'(s(bool,sK8),s('type.option.option'(sK6),sK7),s('type.option.option'(sK6),'const.option.NONE_0')))),
  inference(cnf_transformation,[],[f148])).
% SZS output end Proof for thm.option
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 583
% Passive clauses: 4922
% Generated clauses: 22174
% Final active clauses: 240
% Final passive clauses: 1154
% Input formulas: 48
% Initial clauses: 147
% 
% Selected by SInE selection: 46
% SInE iterations: 1
% 
% Duplicate literals: 577
% Trivial inequalities: 84
% Fw subsumption resolutions: 2636
% Fw demodulations: 1501
% Bw demodulations: 10
% 
% Simple tautologies: 500
% Equational tautologies: 5
% Forward subsumptions: 11070
% 
% Binary resolution: 12212
% Factoring: 158
% Forward superposition: 3392
% Backward superposition: 1115
% Self superposition: 1
% Equality resolution: 9
% 
% Split clauses: 945
% Split components: 1946
% SAT solver clauses: 1020
% SAT solver unit clauses: 32
% SAT solver binary clauses: 34
% 
% Sat splits: 945
% Sat splitting refutations: 74
% 
% Memory used [KB]: 6012
% Time elapsed: 0.354 s
% ------------------------------
% ------------------------------
% Success in time 12.041 s
