Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.seq
% (104408)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 759
% Passive clauses: 10960
% Generated clauses: 20361
% Final active clauses: 751
% Final passive clauses: 9765
% Input formulas: 69
% Initial clauses: 231
% 
% Duplicate literals: 303
% Trivial inequalities: 36
% Fw subsumption resolutions: 1790
% Fw demodulations: 454
% Bw demodulations: 7
% 
% Simple tautologies: 248
% Equational tautologies: 56
% Forward subsumptions: 6268
% 
% Binary resolution: 11902
% Factoring: 1072
% Forward superposition: 2551
% Backward superposition: 1797
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 17
% 
% Split clauses: 343
% Split components: 690
% SAT solver clauses: 667
% SAT solver unit clauses: 21
% SAT solver binary clauses: 152
% 
% Sat splits: 343
% 
% TWLsolver clauses: 667
% TWLsolver calls for satisfiability: 406
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.seq
% (104424)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 471
% Passive clauses: 3555
% Generated clauses: 12900
% Final active clauses: 458
% Final passive clauses: 3034
% Input formulas: 69
% Initial clauses: 231
% 
% Duplicate literals: 182
% Trivial inequalities: 16
% Fw subsumption resolutions: 973
% Fw demodulations: 392
% Bw demodulations: 59
% Condensations: 221
% Global subsumptions: 18
% 
% Simple tautologies: 243
% Equational tautologies: 17
% Forward subsumptions: 3281
% 
% Binary resolution: 4552
% Factoring: 196
% Forward superposition: 1295
% Backward superposition: 1122
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 57
% Split components: 114
% SAT solver clauses: 4188
% SAT solver unit clauses: 241
% SAT solver binary clauses: 102
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 57
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 4116
% TWLsolver calls for satisfiability: 2155
% 
% Memory used [KB]: 14711
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.seq
% (104452)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 4324
% Passive clauses: 18645
% Generated clauses: 137156
% Final active clauses: 1624
% Final passive clauses: 3358
% Input formulas: 69
% Initial clauses: 231
% 
% Duplicate literals: 1026
% Trivial inequalities: 148
% Fw subsumption resolutions: 3138
% Fw demodulations: 29620
% Bw demodulations: 7436
% Condensations: 740
% Global subsumptions: 432
% 
% Simple tautologies: 222
% Equational tautologies: 3042
% Forward subsumptions: 77886
% Bw demodulations to eq. taut.: 794
% 
% Unit resulting resolution: 61613
% Factoring: 1
% Forward superposition: 13111
% Backward superposition: 13168
% Self superposition: 33
% Equality factoring: 32
% Equality resolution: 17
% 
% Split clauses: 145
% Split components: 298
% SAT solver clauses: 21003
% SAT solver unit clauses: 1353
% SAT solver binary clauses: 4177
% 
% Sat splits: 145
% Sat splitting refutations: 1633
% 
% Memory used [KB]: 85584
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.seq
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.seq
% SZS output start Proof for thm.seq
fof(f10329,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f334,f564,f333,f565,f561,f570,f571,f572,f573,f574,f575,f577,f670,f711,f712,f713,f714,f715,f717,f718,f810,f845,f846,f847,f848,f849,f695,f850,f852,f945,f693,f1012,f1013,f1014,f1020,f1021,f1022,f1023,f1024,f1025,f1027,f1120,f1238,f1239,f1240,f1241,f1242,f1244,f1172,f1245,f1337,f1377,f1399,f1400,f1401,f1402,f1403,f1405,f1541,f1406,f1498,f1540,f1504,f1546,f1544,f1686,f1685,f1670,f1666,f1757,f1758,f1759,f1760,f1761,f1762,f1763,f1764,f1765,f1769,f1770,f1771,f1772,f1773,f1775,f1893,f1776,f1868,f1962,f1961,f1966,f1968,f2061,f2060,f2064,f2065,f2212,f2208,f2207,f2186,f2185,f2292,f2293,f2294,f2295,f2296,f2297,f2302,f2303,f2304,f2305,f2306,f2307,f2348,f2397,f2457,f2498,f2499,f2500,f2501,f2502,f2504,f2628,f2505,f2597,f2627,f2989,f2994,f3002,f3011,f1965,f3624,f3627,f3634,f3639,f3647,f3654,f3832,f3833,f3849,f3987,f3985,f3855,f3856,f3857,f3858,f3859,f3989,f3860,f3862,f3955,f3988,f4020,f4021,f4022,f4023,f4024,f4150,f4025,f4027,f4120,f4149,f4168,f4169,f4170,f4171,f4172,f4174,f696,f4298,f4175,f4267,f4297,f6252,f6253,f8062,f8064,f8065,f8066,f8069,f8161,f8198,f9019,f6329,f6330,f6331,f9032,f9033,f1967,f9034,f9035,f9031,f1011,f562,f9045,f9080,f9082,f9083,f9084,f9087,f9179,f834,f9242,f1162,f1160,f9243,f1166,f1164,f9244,f1170,f1168,f9250,f9252,f9253,f9255,f9257,f9350,f1171,f9452,f9454,f9455,f694,f9456,f698,f9459,f697,f9551,f699,f1892,f10175,f10313,f10311,f10315,f10317])).
fof(f10317,plain,(
  $false | $spl29),
  inference(subsumption_resolution,[],[f10316,f346])).
fof(f346,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X7] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X7))))),
  file('thm.seq.SUBSEQ_SUC','thm.prim_rec.LESS_REFL')).
fof(f10316,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | $spl29),
  inference(forward_demodulation,[],[f10177,f344])).
fof(f344,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X8,X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X7)))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.LESS_EQ')).
fof(f10177,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK14)))) | $spl29),
  inference(resolution,[],[f1166,f337])).
fof(f337,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.NOT_LESS_EQUAL')).
fof(f10315,plain,(
  $false | $spl29),
  inference(subsumption_resolution,[],[f10314,f346])).
fof(f10314,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | $spl29),
  inference(forward_demodulation,[],[f10176,f344])).
fof(f10176,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK14)))) | $spl29),
  inference(resolution,[],[f1166,f339])).
fof(f339,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X8,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.NOT_LESS')).
fof(f10311,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))) | $spl85),
  inference(cnf_transformation,[],[f10311_D])).
fof(f10311_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f10313,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) | $spl86),
  inference(cnf_transformation,[],[f10313_D])).
fof(f10313_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f10175,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))) | $spl29),
  inference(resolution,[],[f1166,f341])).
fof(f341,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X8,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) | s('type.num.num',X7) = s('type.num.num',X8)))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.LESS_OR_EQ')).
fof(f1892,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl48),
  inference(cnf_transformation,[],[f1892_D])).
fof(f1892_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f699,plain,(
  ( ! [X21] : (p(s(bool,X21))) ) | $spl18),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ( ! [X21] : (p(s(bool,X21))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f9551,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl44),
  inference(resolution,[],[f1544,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X12,X18,X13,X19] : ((s(bool,X12) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X13) = s(bool,X19))) => ((p(s(bool,X12)) => p(s(bool,X13))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.IMP_CONG')).
fof(f697,plain,(
  ( ! [X14,X12] : (s(bool,X12) != s(bool,X14) | ~p(s(bool,X12))) ) | $spl14),
  inference(cnf_transformation,[],[f697_D])).
fof(f697_D,plain,(
  ( ! [X14,X12] : (s(bool,X12) != s(bool,X14) | ~p(s(bool,X12))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f9459,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X27) != s(bool,X28)) ) | $spl44),
  inference(resolution,[],[f1544,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f698,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) | $spl16),
  inference(cnf_transformation,[],[f698_D])).
fof(f698_D,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f9456,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl44),
  inference(resolution,[],[f1544,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f694,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | p(s(bool,X6))) ) | $spl8),
  inference(cnf_transformation,[],[f694_D])).
fof(f694_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | p(s(bool,X6))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f9455,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl44),
  inference(resolution,[],[f1544,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f9454,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X12) != s(bool,X13)) ) | $spl44),
  inference(resolution,[],[f1544,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f9452,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl44),
  inference(resolution,[],[f1544,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f1171,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl34),
  inference(cnf_transformation,[],[f1171_D])).
fof(f1171_D,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f9350,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl3),
  inference(resolution,[],[f564,f554])).
fof(f9257,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))))) | s(bool,X28) != s(bool,X30)) ) | $spl3),
  inference(resolution,[],[f564,f369])).
fof(f9255,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl3),
  inference(resolution,[],[f564,f368])).
fof(f9253,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl3),
  inference(resolution,[],[f564,f367])).
fof(f9252,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))))) | s(bool,X13) != s(bool,X15)) ) | $spl3),
  inference(resolution,[],[f564,f365])).
fof(f9250,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl3),
  inference(resolution,[],[f564,f363])).
fof(f1168,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) | $spl31),
  inference(cnf_transformation,[],[f1168_D])).
fof(f1168_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1170,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))) | $spl32),
  inference(cnf_transformation,[],[f1170_D])).
fof(f1170_D,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f9244,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) | $spl3),
  inference(resolution,[],[f564,f341])).
fof(f1164,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK16(sK13)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))))) | $spl27),
  inference(cnf_transformation,[],[f1164_D])).
fof(f1164_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK16(sK13)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1166,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))) | $spl29),
  inference(cnf_transformation,[],[f1166_D])).
fof(f1166_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f9243,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK16(sK13)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))))) | $spl3),
  inference(resolution,[],[f564,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK16(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X0)))))))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK16(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X0))))))))) | ! [X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f229])).
fof(f229,plain,(
  ! [X0] : (? [X3] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) | ! [X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))))),
  inference(rectify,[],[f160])).
fof(f160,plain,(
  ! [X0] : (? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) | ! [X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))))))))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) => ! [X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) => ! [X8,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X8))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))))))))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.STRICTLY_INCREASING_TC')).
fof(f1160,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK17(sK13)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK17(sK13))))))) | $spl23),
  inference(cnf_transformation,[],[f1160_D])).
fof(f1160_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK17(sK13)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK17(sK13))))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1162,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK14)))) | $spl25),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK14)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f9242,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK17(sK13)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK17(sK13))))))) | $spl3),
  inference(resolution,[],[f564,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK17(X0)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK17(X0)))))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK17(X0)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK17(X0))))))) | ! [X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f231])).
fof(f231,plain,(
  ! [X0] : (? [X3] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3)))))) | ! [X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))))))))),
  inference(rectify,[],[f161])).
fof(f161,plain,(
  ! [X0] : (? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1)))))) | ! [X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1)))))) => ! [X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7)))))) => ! [X8,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X8))))))))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.arithmetic_357_prove_ltb')).
fof(f834,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl20),
  inference(cnf_transformation,[],[f834_D])).
fof(f834_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f9179,plain,(
  ( ! [X177,X176] : (p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X177) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f554])).
fof(f9087,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X24) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f561,f369])).
fof(f9084,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f368])).
fof(f9083,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f367])).
fof(f9082,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X9) != s(bool,X10)) ) | $spl0),
  inference(resolution,[],[f561,f365])).
fof(f9080,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f363])).
fof(f9045,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f564,f565])).
fof(f562,plain,(
  ~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) | $spl1),
  inference(cnf_transformation,[],[f562_D])).
fof(f562_D,plain,(
  ~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1011,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))) | ~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) | $spl3),
  inference(resolution,[],[f564,f547])).
fof(f547,plain,(
  ( ! [X4,X0,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X4)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | ~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0))))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0] : ((~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X3,X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X4)))))))) & ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(X0)),s('type.num.num',sK30(X0))))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK29(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK30(X0)))))))) | p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29,sK30])],[f330])).
fof(f330,plain,(
  ! [X0] : ((~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X3,X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X4)))))))) & (? [X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))) | p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0))))))),
  inference(rectify,[],[f329])).
fof(f329,plain,(
  ! [X0] : ((~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2)))))))) & (? [X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))) | p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0))))))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) <=> ! [X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) <=> ! [X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X6)))) <=> ! [X8,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X8))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))))))))),
  file('thm.seq.SUBSEQ_SUC','thm.seq.subseq')).
fof(f9031,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl81),
  inference(cnf_transformation,[],[f9031_D])).
fof(f9031_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f9035,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X22))) ) | $spl76),
  inference(subsumption_resolution,[],[f6313,f530])).
fof(f530,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.EQ_CLAUSES|split|0')).
fof(f6313,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X22) | ~p(s(bool,X22))) ) | $spl76),
  inference(superposition,[],[f6253,f530])).
fof(f9034,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X23))) ) | $spl76),
  inference(subsumption_resolution,[],[f6314,f530])).
fof(f6314,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | $spl76),
  inference(superposition,[],[f6253,f532])).
fof(f532,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.EQ_CLAUSES|split|1')).
fof(f1967,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37)) ) | $spl54),
  inference(cnf_transformation,[],[f1967_D])).
fof(f1967_D,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f9033,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X24)) ) | $spl76),
  inference(subsumption_resolution,[],[f6315,f6253])).
fof(f6315,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X24) | s(bool,'F_0') = s(bool,X24)) ) | $spl76),
  inference(superposition,[],[f6253,f553])).
fof(f553,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.seq.SUBSEQ_SUC',aHL_BOOL_CASES)).
fof(f9032,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) | $spl82),
  inference(cnf_transformation,[],[f9032_D])).
fof(f9032_D,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f6331,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl76),
  inference(trivial_inequality_removal,[],[f6317])).
fof(f6317,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,X29) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl76),
  inference(superposition,[],[f6253,f553])).
fof(f6330,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl76),
  inference(trivial_inequality_removal,[],[f6318])).
fof(f6318,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,X30) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl76),
  inference(superposition,[],[f6253,f528])).
fof(f528,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.EQ_CLAUSES|split|3')).
fof(f6329,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl76),
  inference(trivial_inequality_removal,[],[f6319])).
fof(f6319,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,X31) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl76),
  inference(superposition,[],[f6253,f526])).
fof(f526,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.EQ_CLAUSES|split|2')).
fof(f9019,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f9018,f546])).
fof(f546,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.seq.SUBSEQ_SUC',aHL_TRUTH)).
fof(f9018,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f9017,f2595])).
fof(f2595,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13)))) | $spl1),
  inference(resolution,[],[f841,f530])).
fof(f841,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) | $spl1),
  inference(resolution,[],[f562,f548])).
fof(f548,plain,(
  ( ! [X0] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(X0)),s('type.num.num',sK30(X0)))))) )),
  inference(cnf_transformation,[],[f331])).
fof(f9017,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f8867,f565])).
fof(f8867,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK16(sK13)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))))) | $spl1),
  inference(resolution,[],[f842,f347])).
fof(f842,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK29(sK13)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK30(sK13))))))) | $spl1),
  inference(resolution,[],[f562,f549])).
fof(f549,plain,(
  ( ! [X0] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK29(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK30(X0)))))))) )),
  inference(cnf_transformation,[],[f331])).
fof(f8198,plain,(
  ( ! [X22,X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X22,sK13))))))) ) | $spl78),
  inference(cnf_transformation,[],[f8198_D])).
fof(f8198_D,plain,(
  ( ! [X22,X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X22,sK13))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8161,plain,(
  ( ! [X288,X287,X286] : (p(s(bool,X286)) | p(s(bool,X287)) | s(bool,X287) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X288),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X288,sK13))))))) ) | $spl4),
  inference(resolution,[],[f1386,f554])).
fof(f1386,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X2,sK13)))))))) ) | $spl4),
  inference(resolution,[],[f565,f349])).
fof(f349,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK19(X1,X0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK18(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(X0)))))))))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK18(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK18(X0))))))))) | ! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK19(X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f233])).
fof(f233,plain,(
  ! [X0] : (? [X3] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) | ! [X1] : ? [X2] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2)))))))),
  inference(rectify,[],[f162])).
fof(f162,plain,(
  ! [X0] : (? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) | ! [X2] : ? [X3] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3)))))))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) => ! [X2] : ? [X3] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3)))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) => ! [X9] : ? [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7)))))))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.STRICTLY_INCREASING_UNBOUNDED')).
fof(f8069,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X48),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X48,sK13)))))) | s(bool,X47) != s(bool,X49)) ) | $spl4),
  inference(resolution,[],[f1386,f369])).
fof(f8066,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X37) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X38,sK13))))))) ) | $spl4),
  inference(resolution,[],[f1386,f368])).
fof(f8065,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X34,sK13))))))) ) | $spl4),
  inference(resolution,[],[f1386,f367])).
fof(f8064,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X29,sK13)))))) | s(bool,X28) != s(bool,X30)) ) | $spl4),
  inference(resolution,[],[f1386,f365])).
fof(f8062,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK19(X22,sK13))))))) ) | $spl4),
  inference(resolution,[],[f1386,f363])).
fof(f6253,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl76),
  inference(cnf_transformation,[],[f6253_D])).
fof(f6253_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f6252,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl1),
  inference(duplicate_literal_removal,[],[f6246])).
fof(f6246,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl1),
  inference(resolution,[],[f3172,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f104,f213,f212,f211,f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f211,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f212,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f213,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> s(bool,X14) = s(bool,X15)) <=> ((p(s(bool,X16)) | p(s(bool,X15)) | p(s(bool,X14))) & (p(s(bool,X16)) | ~p(s(bool,X14)) | ~p(s(bool,X15))) & (p(s(bool,X15)) | ~p(s(bool,X14)) | ~p(s(bool,X16))) & (p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X16)))))),
  file('thm.seq.SUBSEQ_SUC','thm.sat.sat_2|49|_prove_ltb')).
fof(f3172,plain,(
  ( ! [X4,X3] : (sP9(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) ) | $spl1),
  inference(superposition,[],[f3138,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f284])).
fof(f3138,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(superposition,[],[f942,f940])).
fof(f940,plain,(
  s(bool,'F_0') = s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | $spl1),
  inference(resolution,[],[f562,f526])).
fof(f942,plain,(
  s(bool,'T_0') != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | $spl1),
  inference(resolution,[],[f562,f529])).
fof(f529,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f325])).
fof(f4297,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl72),
  inference(cnf_transformation,[],[f4297_D])).
fof(f4297_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4267,plain,(
  ( ! [X175,X174] : (p(s(bool,X174)) | p(s(bool,X175)) | s(bool,X175) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f2493,f554])).
fof(f2493,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) | $spl1),
  inference(resolution,[],[f841,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f4175,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f2493,f369])).
fof(f4298,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl74),
  inference(cnf_transformation,[],[f4298_D])).
fof(f4298_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f696,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9)) | p(s(bool,X10))) ) | $spl12),
  inference(cnf_transformation,[],[f696_D])).
fof(f696_D,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9)) | p(s(bool,X10))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4174,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f2493,f369])).
fof(f4172,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f2493,f368])).
fof(f4171,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f2493,f367])).
fof(f4170,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f2493,f365])).
fof(f4169,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f2493,f365])).
fof(f4168,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f2493,f363])).
fof(f4149,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl68),
  inference(cnf_transformation,[],[f4149_D])).
fof(f4149_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4120,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl1),
  inference(resolution,[],[f2491,f554])).
fof(f2491,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) | $spl1),
  inference(resolution,[],[f841,f338])).
fof(f338,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f4027,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13)))) | s(bool,X27) != s(bool,X29)) ) | $spl1),
  inference(resolution,[],[f2491,f369])).
fof(f4025,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl1),
  inference(resolution,[],[f2491,f368])).
fof(f4150,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl70),
  inference(cnf_transformation,[],[f4150_D])).
fof(f4150_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4024,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13)))) | s(bool,X18) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f2491,f367])).
fof(f4023,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl1),
  inference(resolution,[],[f2491,f367])).
fof(f4022,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13)))) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f2491,f365])).
fof(f4021,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13)))) | s(bool,X9) != s(bool,X10)) ) | $spl1),
  inference(resolution,[],[f2491,f363])).
fof(f4020,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl1),
  inference(resolution,[],[f2491,f363])).
fof(f3988,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl64),
  inference(cnf_transformation,[],[f3988_D])).
fof(f3988_D,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3955,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl1),
  inference(resolution,[],[f2489,f554])).
fof(f2489,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) | $spl1),
  inference(resolution,[],[f841,f336])).
fof(f336,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X8,X7] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.LESS_ANTISYM')).
fof(f3862,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13)))) | s(bool,X28) != s(bool,X30)) ) | $spl1),
  inference(resolution,[],[f2489,f369])).
fof(f3860,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl1),
  inference(resolution,[],[f2489,f368])).
fof(f3989,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl66),
  inference(cnf_transformation,[],[f3989_D])).
fof(f3989_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3859,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13)))) | s(bool,X19) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f2489,f367])).
fof(f3858,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl1),
  inference(resolution,[],[f2489,f367])).
fof(f3857,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13)))) | s(bool,X13) != s(bool,X15)) ) | $spl1),
  inference(resolution,[],[f2489,f365])).
fof(f3856,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13)))) | s(bool,X10) != s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f2489,f363])).
fof(f3855,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) ) | $spl1),
  inference(resolution,[],[f2489,f363])).
fof(f3985,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) | $spl61),
  inference(cnf_transformation,[],[f3985_D])).
fof(f3985_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f3987,plain,(
  s('type.num.num',sK29(sK13)) = s('type.num.num',sK30(sK13)) | $spl62),
  inference(cnf_transformation,[],[f3987_D])).
fof(f3987_D,plain,(
  s('type.num.num',sK29(sK13)) = s('type.num.num',sK30(sK13)) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3849,plain,(
  s('type.num.num',sK29(sK13)) = s('type.num.num',sK30(sK13)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK30(sK13)),s('type.num.num',sK29(sK13))))) | $spl1),
  inference(resolution,[],[f2489,f341])).
fof(f3833,plain,(
  $false | ($spl4 | $spl50)),
  inference(trivial_inequality_removal,[],[f3808])).
fof(f3808,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl50)),
  inference(superposition,[],[f1893,f1866])).
fof(f1866,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13))) | $spl4),
  inference(resolution,[],[f1387,f530])).
fof(f1387,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) | $spl4),
  inference(resolution,[],[f565,f351])).
fof(f351,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK21(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK21(X0)))))))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK21(X0)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK21(X0))))))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21])],[f164])).
fof(f164,plain,(
  ! [X0] : (? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) => p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) => p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X6)))))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.STRICTLY_INCREASING_ONE_ONE')).
fof(f3832,plain,(
  $false | $spl50),
  inference(equality_resolution,[],[f1893])).
fof(f3654,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') != s(bool,X109) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,X108)) ) | $spl4),
  inference(forward_demodulation,[],[f3653,f1866])).
fof(f3653,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,X108) | s(bool,X109) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(subsumption_resolution,[],[f3383,f499])).
fof(f499,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.IMP_CLAUSES|split|3')).
fof(f3383,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,'T_0')) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,X108) | s(bool,X109) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(superposition,[],[f369,f1866])).
fof(f3647,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X93) != s(bool,X94)) ) | $spl4),
  inference(subsumption_resolution,[],[f3378,f554])).
fof(f3378,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) | ~p(s(bool,X94)) | s(bool,X93) != s(bool,X94)) ) | $spl4),
  inference(superposition,[],[f369,f1866])).
fof(f3639,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') != s(bool,X62) | ~p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X63) != s(bool,X64)) ) | $spl4),
  inference(subsumption_resolution,[],[f3368,f369])).
fof(f3368,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') != s(bool,X62) | p(s(bool,X62)) | ~p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X63) != s(bool,X64)) ) | $spl4),
  inference(superposition,[],[f367,f1866])).
fof(f3634,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,X49) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,X48)) ) | $spl4),
  inference(forward_demodulation,[],[f3633,f1866])).
fof(f3633,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,X48) | s(bool,X49) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(subsumption_resolution,[],[f3363,f499])).
fof(f3363,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,X48) | s(bool,X49) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(superposition,[],[f365,f1866])).
fof(f3627,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl4),
  inference(subsumption_resolution,[],[f3359,f554])).
fof(f3359,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl4),
  inference(superposition,[],[f365,f1866])).
fof(f3624,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X19)) ) | $spl4),
  inference(subsumption_resolution,[],[f3353,f369])).
fof(f3353,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X17) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X19)) ) | $spl4),
  inference(superposition,[],[f363,f1866])).
fof(f1965,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl52),
  inference(cnf_transformation,[],[f1965_D])).
fof(f1965_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3011,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl1),
  inference(subsumption_resolution,[],[f2743,f554])).
fof(f2743,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl1),
  inference(superposition,[],[f369,f940])).
fof(f3002,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X60) != s(bool,X61)) ) | $spl1),
  inference(subsumption_resolution,[],[f2733,f369])).
fof(f2733,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') != s(bool,X59) | p(s(bool,X59)) | ~p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X60) != s(bool,X61)) ) | $spl1),
  inference(superposition,[],[f367,f940])).
fof(f2994,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') != s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X34)) ) | $spl1),
  inference(subsumption_resolution,[],[f2724,f554])).
fof(f2724,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') != s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X34)) ) | $spl1),
  inference(superposition,[],[f365,f940])).
fof(f2989,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,X14) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X16)) ) | $spl1),
  inference(subsumption_resolution,[],[f2718,f369])).
fof(f2718,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,X14) | p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X16)) ) | $spl1),
  inference(superposition,[],[f363,f940])).
fof(f2627,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl56),
  inference(cnf_transformation,[],[f2627_D])).
fof(f2627_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2597,plain,(
  ( ! [X182,X183] : (p(s(bool,X182)) | p(s(bool,X183)) | s(bool,X183) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f841,f554])).
fof(f2505,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13)))) | s(bool,X30) != s(bool,X31)) ) | $spl1),
  inference(resolution,[],[f841,f369])).
fof(f2628,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl58),
  inference(cnf_transformation,[],[f2628_D])).
fof(f2628_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2504,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f841,f369])).
fof(f2502,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f841,f368])).
fof(f2501,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f841,f367])).
fof(f2500,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13)))) | s(bool,X15) != s(bool,X16)) ) | $spl1),
  inference(resolution,[],[f841,f365])).
fof(f2499,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f841,f365])).
fof(f2498,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK13)),s('type.num.num',sK30(sK13))))) ) | $spl1),
  inference(resolution,[],[f841,f363])).
fof(f2457,plain,(
  $false | $spl10),
  inference(equality_resolution,[],[f695])).
fof(f2397,plain,(
  $false | $spl54),
  inference(equality_resolution,[],[f1967])).
fof(f2348,plain,(
  $false | $spl52),
  inference(equality_resolution,[],[f1965])).
fof(f2307,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2218,f2059])).
fof(f2059,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(equality_resolution,[],[f697])).
fof(f2218,plain,(
  ( ! [X341,X342] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X341),s('type.num.num',X342))))) ) | $spl14),
  inference(resolution,[],[f2059,f337])).
fof(f2306,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2217,f2059])).
fof(f2217,plain,(
  ( ! [X339,X340] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X339),s('type.num.num',X340))))) ) | $spl14),
  inference(resolution,[],[f2059,f339])).
fof(f2305,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2209,f2059])).
fof(f2209,plain,(
  ( ! [X326] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X326))))) ) | $spl14),
  inference(resolution,[],[f2059,f548])).
fof(f2304,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2200,f2059])).
fof(f2200,plain,(
  ( ! [X310,X311] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X310),s('type.num.num',X311))))) ) | $spl14),
  inference(resolution,[],[f2059,f337])).
fof(f2303,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2199,f2059])).
fof(f2199,plain,(
  ( ! [X308,X309] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X308),s('type.num.num',X309))))) ) | $spl14),
  inference(resolution,[],[f2059,f339])).
fof(f2302,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2193,f2059])).
fof(f2193,plain,(
  ( ! [X293] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(X293)),s('type.num.num',sK30(X293)))))) ) | $spl14),
  inference(resolution,[],[f2059,f548])).
fof(f2297,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2184,f2059])).
fof(f2184,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f2059,f543])).
fof(f543,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f33])).
fof(f33,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.NOT_CLAUSES|split|1')).
fof(f2296,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2183,f2059])).
fof(f2183,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f2059,f541])).
fof(f541,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f38])).
fof(f38,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.NOT_CLAUSES|split|2')).
fof(f2295,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2180,f2059])).
fof(f2180,plain,(
  ( ! [X284] : (p(s(bool,X284))) ) | $spl14),
  inference(resolution,[],[f2059,f507])).
fof(f507,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.IMP_CLAUSES|split|1')).
fof(f2294,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2178,f2059])).
fof(f2178,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(resolution,[],[f2059,f504])).
fof(f504,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.IMP_CLAUSES|split|2')).
fof(f2293,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2177,f2059])).
fof(f2177,plain,(
  ( ! [X282] : (p(s(bool,X282))) ) | $spl14),
  inference(resolution,[],[f2059,f501])).
fof(f501,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X10] : ((p(s(bool,'T_0')) => p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.IMP_CLAUSES|split|0')).
fof(f2292,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2175,f2059])).
fof(f2175,plain,(
  ( ! [X280] : (p(s(bool,X280))) ) | $spl14),
  inference(resolution,[],[f2059,f498])).
fof(f498,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f2185,plain,(
  $false | $spl14),
  inference(resolution,[],[f2059,f545])).
fof(f545,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f22])).
fof(f22,axiom,(
  p(s(bool,'T_0'))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.TRUTH')).
fof(f2186,plain,(
  $false | $spl14),
  inference(resolution,[],[f2059,f546])).
fof(f2207,plain,(
  $false | ($spl2 | $spl14)),
  inference(resolution,[],[f2059,f563])).
fof(f563,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) | $spl2),
  inference(cnf_transformation,[],[f563_D])).
fof(f563_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2208,plain,(
  $false | ($spl4 | $spl14)),
  inference(resolution,[],[f2059,f565])).
fof(f2212,plain,(
  $false | ($spl4 | $spl14)),
  inference(resolution,[],[f2059,f1387])).
fof(f2065,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl14),
  inference(subsumption_resolution,[],[f2024,f697])).
fof(f2024,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl14),
  inference(superposition,[],[f697,f553])).
fof(f2064,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl14),
  inference(subsumption_resolution,[],[f2023,f697])).
fof(f2023,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl14),
  inference(superposition,[],[f697,f553])).
fof(f2060,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2022])).
fof(f2022,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl14),
  inference(superposition,[],[f697,f532])).
fof(f2061,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2021])).
fof(f2021,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl14),
  inference(superposition,[],[f697,f530])).
fof(f1968,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38))) ) | $spl8),
  inference(subsumption_resolution,[],[f1925,f694])).
fof(f1925,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl8),
  inference(superposition,[],[f694,f553])).
fof(f1966,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36))) ) | $spl8),
  inference(subsumption_resolution,[],[f1924,f694])).
fof(f1924,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl8),
  inference(superposition,[],[f694,f553])).
fof(f1961,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f1921])).
fof(f1921,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl8),
  inference(superposition,[],[f694,f528])).
fof(f1962,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f1920])).
fof(f1920,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f694,f526])).
fof(f1868,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(resolution,[],[f1387,f554])).
fof(f1776,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X27) != s(bool,X28)) ) | $spl4),
  inference(resolution,[],[f1387,f369])).
fof(f1893,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl50),
  inference(cnf_transformation,[],[f1893_D])).
fof(f1893_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1775,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X24) | s(bool,X25) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(resolution,[],[f1387,f369])).
fof(f1773,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(resolution,[],[f1387,f368])).
fof(f1772,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(resolution,[],[f1387,f367])).
fof(f1771,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X12) != s(bool,X13)) ) | $spl4),
  inference(resolution,[],[f1387,f365])).
fof(f1770,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(resolution,[],[f1387,f365])).
fof(f1769,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl4),
  inference(resolution,[],[f1387,f363])).
fof(f1765,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1699,f699])).
fof(f1699,plain,(
  ( ! [X366,X365] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X365),s('type.num.num',X366))))) ) | $spl18),
  inference(resolution,[],[f699,f338])).
fof(f1764,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1698,f699])).
fof(f1698,plain,(
  ( ! [X364,X363] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X363),s('type.num.num',X364))))) ) | $spl18),
  inference(resolution,[],[f699,f340])).
fof(f340,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f1763,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1679,f699])).
fof(f1679,plain,(
  ( ! [X318,X317] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X317),s('type.num.num',X318))))) ) | $spl18),
  inference(resolution,[],[f699,f340])).
fof(f1762,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1678,f699])).
fof(f1678,plain,(
  ( ! [X315,X316] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X315),s('type.num.num',X316))))) ) | $spl18),
  inference(resolution,[],[f699,f338])).
fof(f1761,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1677,f699])).
fof(f1677,plain,(
  ( ! [X313,X314] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X313),s('type.num.num',X314))))) ) | $spl18),
  inference(resolution,[],[f699,f336])).
fof(f1760,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1676,f699])).
fof(f1676,plain,(
  ( ! [X312,X311] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X311),s('type.num.num',X312))))) ) | $spl18),
  inference(resolution,[],[f699,f336])).
fof(f1759,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1668,f699])).
fof(f1668,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(resolution,[],[f699,f544])).
fof(f544,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f328])).
fof(f1758,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1667,f699])).
fof(f1667,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(resolution,[],[f699,f542])).
fof(f542,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f327])).
fof(f1757,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1661,f699])).
fof(f1661,plain,(
  ( ! [X284] : (~p(s(bool,X284))) ) | $spl18),
  inference(resolution,[],[f699,f496])).
fof(f496,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X10)))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.IMP_CLAUSES|split|4')).
fof(f1666,plain,(
  $false | $spl18),
  inference(resolution,[],[f699,f540])).
fof(f540,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.seq.SUBSEQ_SUC',aHL_FALSITY)).
fof(f1670,plain,(
  $false | ($spl1 | $spl18)),
  inference(resolution,[],[f699,f562])).
fof(f1685,plain,(
  $false | $spl18),
  inference(resolution,[],[f699,f346])).
fof(f1686,plain,(
  $false | $spl18),
  inference(resolution,[],[f699,f345])).
fof(f345,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK15(X0)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK15(X0)))))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK15(X0)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK15(X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f159])).
fof(f159,plain,(
  ! [X0] : ? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ~! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ~! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))))))),
  file('thm.seq.SUBSEQ_SUC','thm.arithmetic.NOT_STRICTLY_DECREASING')).
fof(f1544,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) | $spl44),
  inference(cnf_transformation,[],[f1544_D])).
fof(f1544_D,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1546,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK24(sK13,'type.num.num','type.num.num')))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK25(sK13,'type.num.num','type.num.num'))))))))) | $spl46),
  inference(cnf_transformation,[],[f1546_D])).
fof(f1546_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK24(sK13,'type.num.num','type.num.num')))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK25(sK13,'type.num.num','type.num.num'))))))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1504,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK24(sK13,'type.num.num','type.num.num')))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK25(sK13,'type.num.num','type.num.num'))))))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK13)))) | $spl4),
  inference(superposition,[],[f565,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,sK24(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK25(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & ((s(X0,happ(s(fun(X1,X0),X2),s(X1,sK24(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK25(X2,X1,X0)))) & s(X1,sK24(X2,X1,X0)) != s(X1,sK25(X2,X1,X0))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24,sK25])],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4)))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X17,X11] : ! [X6] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X11,X17),X6)))) <=> ! [X28,X29] : (s(X17,happ(s(fun(X11,X17),X6),s(X11,X28))) = s(X17,happ(s(fun(X11,X17),X6),s(X11,X29))) => s(X11,X28) = s(X11,X29)))),
  file('thm.seq.SUBSEQ_SUC','thm.bool.ONE_ONE_THM')).
fof(f1540,plain,(
  ( ! [X23,X25] : (p(s(bool,X23)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X25))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))))) ) | $spl40),
  inference(cnf_transformation,[],[f1540_D])).
fof(f1540_D,plain,(
  ( ! [X23,X25] : (p(s(bool,X23)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X25))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1498,plain,(
  ( ! [X290,X289,X291] : (p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X290) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X291))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X291)))))))) ) | $spl4),
  inference(resolution,[],[f565,f554])).
fof(f1406,plain,(
  ( ! [X52,X50,X51,X49] : (p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X51))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))))))) | s(bool,X50) != s(bool,X52)) ) | $spl4),
  inference(resolution,[],[f565,f369])).
fof(f1541,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X29))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))))) ) | $spl42),
  inference(cnf_transformation,[],[f1541_D])).
fof(f1541_D,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X29))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1405,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X46) | s(bool,X48) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X47))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X47)))))))) ) | $spl4),
  inference(resolution,[],[f565,f369])).
fof(f1403,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X41))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X41)))))))) ) | $spl4),
  inference(resolution,[],[f565,f368])).
fof(f1402,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X37))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37)))))))) ) | $spl4),
  inference(resolution,[],[f565,f367])).
fof(f1401,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X32))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))))))) | s(bool,X31) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f565,f365])).
fof(f1400,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X29))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))))) ) | $spl4),
  inference(resolution,[],[f565,f365])).
fof(f1399,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X25))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))))) ) | $spl4),
  inference(resolution,[],[f565,f363])).
fof(f1377,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl38),
  inference(cnf_transformation,[],[f1377_D])).
fof(f1377_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1337,plain,(
  ( ! [X182,X183] : (p(s(bool,X182)) | p(s(bool,X183)) | s(bool,X183) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl2),
  inference(resolution,[],[f563,f554])).
fof(f1245,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))))) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f563,f369])).
fof(f1172,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl36),
  inference(cnf_transformation,[],[f1172_D])).
fof(f1172_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1244,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl2),
  inference(resolution,[],[f563,f369])).
fof(f1242,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl2),
  inference(resolution,[],[f563,f368])).
fof(f1241,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl2),
  inference(resolution,[],[f563,f367])).
fof(f1240,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))))) | s(bool,X15) != s(bool,X16)) ) | $spl2),
  inference(resolution,[],[f563,f365])).
fof(f1239,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl2),
  inference(resolution,[],[f563,f365])).
fof(f1238,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl2),
  inference(resolution,[],[f563,f363])).
fof(f1120,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl3),
  inference(resolution,[],[f564,f554])).
fof(f1027,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))))) | s(bool,X28) != s(bool,X30)) ) | $spl3),
  inference(resolution,[],[f564,f369])).
fof(f1025,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl3),
  inference(resolution,[],[f564,f368])).
fof(f1024,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))))) | s(bool,X19) != s(bool,X20)) ) | $spl3),
  inference(resolution,[],[f564,f367])).
fof(f1023,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl3),
  inference(resolution,[],[f564,f367])).
fof(f1022,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))))) | s(bool,X13) != s(bool,X15)) ) | $spl3),
  inference(resolution,[],[f564,f365])).
fof(f1021,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))))) | s(bool,X10) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f564,f363])).
fof(f1020,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) ) | $spl3),
  inference(resolution,[],[f564,f363])).
fof(f1014,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) | $spl3),
  inference(resolution,[],[f564,f341])).
fof(f1013,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK16(sK13)))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(sK13))))))))) | $spl3),
  inference(resolution,[],[f564,f347])).
fof(f1012,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK17(sK13)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK17(sK13))))))) | $spl3),
  inference(resolution,[],[f564,f348])).
fof(f693,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl6),
  inference(cnf_transformation,[],[f693_D])).
fof(f693_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f945,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl1),
  inference(resolution,[],[f562,f554])).
fof(f852,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X27) != s(bool,X29)) ) | $spl1),
  inference(resolution,[],[f562,f369])).
fof(f850,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl1),
  inference(resolution,[],[f562,f368])).
fof(f695,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl10),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f849,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X20) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X18) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f562,f367])).
fof(f848,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl1),
  inference(resolution,[],[f562,f367])).
fof(f847,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f562,f365])).
fof(f846,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X9) != s(bool,X10)) ) | $spl1),
  inference(resolution,[],[f562,f363])).
fof(f845,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl1),
  inference(resolution,[],[f562,f363])).
fof(f810,plain,(
  ( ! [X177,X176] : (p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X177) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f554])).
fof(f718,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X24) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f561,f369])).
fof(f717,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f369])).
fof(f715,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f368])).
fof(f714,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f367])).
fof(f713,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X9) != s(bool,X10)) ) | $spl0),
  inference(resolution,[],[f561,f365])).
fof(f712,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X7) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f365])).
fof(f711,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl0),
  inference(resolution,[],[f561,f363])).
fof(f670,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl1),
  inference(resolution,[],[f562,f554])).
fof(f577,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X27) != s(bool,X29)) ) | $spl1),
  inference(resolution,[],[f562,f369])).
fof(f575,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl1),
  inference(resolution,[],[f562,f368])).
fof(f574,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X20) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X18) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f562,f367])).
fof(f573,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl1),
  inference(resolution,[],[f562,f367])).
fof(f572,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f562,f365])).
fof(f571,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))) | s(bool,X9) != s(bool,X10)) ) | $spl1),
  inference(resolution,[],[f562,f363])).
fof(f570,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) ) | $spl1),
  inference(resolution,[],[f562,f363])).
fof(f561,plain,(
  p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) | $spl0),
  inference(cnf_transformation,[],[f561_D])).
fof(f561_D,plain,(
  p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f565,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) ) | $spl4),
  inference(cnf_transformation,[],[f565_D])).
fof(f565_D,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f333,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))) | p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) | ! [X2] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) & (~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f222])).
fof(f222,plain,(
  ? [X0] : ((p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X2] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) & (~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | ? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ? [X0] : ((p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) & (~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) | ? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ? [X0] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) <~> ! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X0)))) <=> ! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X6)))) <=> ! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),X6)))) <=> ! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))))),
  file('thm.seq.SUBSEQ_SUC','thm.seq.SUBSEQ_SUC')).
fof(f564,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) | $spl3),
  inference(cnf_transformation,[],[f564_D])).
fof(f564_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f334,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num',sK14))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK13),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14)))))))) | ~p(s(bool,'const.seq.subseq_1'(s(fun('type.num.num','type.num.num'),sK13))))),
  inference(cnf_transformation,[],[f223])).
% SZS output end Proof for thm.seq
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 409
% Passive clauses: 3279
% Generated clauses: 9821
% Final active clauses: 200
% Final passive clauses: 498
% Input formulas: 69
% Initial clauses: 221
% 
% Selected by SInE selection: 62
% SInE iterations: 1
% 
% Duplicate literals: 272
% Trivial inequalities: 61
% Fw subsumption resolutions: 1617
% Fw demodulations: 560
% Bw demodulations: 11
% 
% Simple tautologies: 501
% Equational tautologies: 1
% Forward subsumptions: 3214
% 
% Binary resolution: 3469
% Forward superposition: 1878
% Backward superposition: 1843
% Equality resolution: 8
% 
% Split clauses: 153
% Split components: 321
% SAT solver clauses: 193
% SAT solver unit clauses: 32
% SAT solver binary clauses: 9
% 
% Sat splits: 153
% Sat splitting refutations: 39
% 
% Memory used [KB]: 4221
% Time elapsed: 0.168 s
% ------------------------------
% ------------------------------
% Success in time 11.9 s
