Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (83806)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 114
% Passive clauses: 3403
% Generated clauses: 12542
% Final active clauses: 107
% Final passive clauses: 3026
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 484
% Trivial inequalities: 199
% Fw subsumption resolutions: 1386
% Fw demodulations: 373
% Bw demodulations: 5
% 
% Simple tautologies: 284
% Equational tautologies: 89
% Forward subsumptions: 5415
% 
% Binary resolution: 282
% Factoring: 2
% Forward superposition: 3808
% Backward superposition: 5304
% Self superposition: 270
% Equality factoring: 190
% Equality resolution: 33
% 
% Split clauses: 465
% Split components: 974
% SAT solver clauses: 1031
% SAT solver unit clauses: 35
% SAT solver binary clauses: 368
% 
% Sat splits: 465
% 
% TWLsolver clauses: 1001
% TWLsolver calls for satisfiability: 126
% 
% Memory used [KB]: 8571
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (83823)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 157
% Passive clauses: 1349
% Generated clauses: 15249
% Final active clauses: 142
% Final passive clauses: 1129
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 290
% Trivial inequalities: 113
% Fw subsumption resolutions: 515
% Fw demodulations: 2192
% Bw demodulations: 71
% Condensations: 664
% Global subsumptions: 12
% 
% Simple tautologies: 145
% Equational tautologies: 46
% Forward subsumptions: 7719
% 
% Binary resolution: 68
% Factoring: 2
% Forward superposition: 4295
% Backward superposition: 5310
% Self superposition: 140
% Equality factoring: 68
% Equality resolution: 27
% 
% Split clauses: 28
% Split components: 56
% SAT solver clauses: 2341
% SAT solver unit clauses: 158
% SAT solver binary clauses: 69
% SAT solver learnt clauses: 24
% 
% Sat splits: 28
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2293
% TWLsolver calls for satisfiability: 426
% 
% Memory used [KB]: 25202
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f106120,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f106119])).
fof(f106119,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3)))),
  inference(superposition,[],[f105618,f245])).
fof(f245,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f185])).
fof(f185,plain,(
  ? [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|8','thm.ringNorm.polynom_distinct|split|8')).
fof(f105618,plain,(
  ( ! [X930,X931] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),X930),s('type.ringNorm.polynom'(sK0),X931)))) )),
  inference(superposition,[],[f337,f104369])).
fof(f104369,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(backward_demodulation,[],[f104368,f103036])).
fof(f103036,plain,(
  ( ! [X65] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.ringNorm.polynom'(sK0),'type.ringNorm.polynom'(sK0)),X65),s('type.ringNorm.polynom'(sK0),sK3)))) )),
  inference(superposition,[],[f50560,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6] : ! [X11,X24] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X11),s('type.ringNorm.polynom'(X6),X24))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X6),X11))),s('type.ringNorm.polynom'(X6),X24)))),
  file('thm.ringNorm.polynom_distinct|split|8','ARITY_0_2_const.ringNorm.Pmult')).
fof(f50560,plain,(
  ( ! [X6,X11,X9] : (s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),X11),s('type.ringNorm.polynom'(sK0),sK3))) = s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X6)),X9),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(superposition,[],[f11096,f350])).
fof(f350,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X15,X6] : ! [X9,X16,X36,X37,X38,X39] : s(X15,happ(s(fun('type.ringNorm.polynom'(X6),X15),X39),s('type.ringNorm.polynom'(X6),X9))) = s(X15,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))),s(fun('type.quote.index',X15),X16),s(fun(X6,X15),X36),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X37),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X38),s(fun('type.ringNorm.polynom'(X6),X15),X39)))),
  file('thm.ringNorm.polynom_distinct|split|8','thm.ringNorm.polynom_case_def|split|4')).
fof(f11096,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),X0),X5)))) )),
  inference(superposition,[],[f351,f245])).
fof(f351,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X15,X6] : ! [X8,X7,X16,X36,X37,X38,X39] : s(X15,happ(s(fun('type.ringNorm.polynom'(X6),X15),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X37),s('type.ringNorm.polynom'(X6),X8))),s('type.ringNorm.polynom'(X6),X7))) = s(X15,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7))),s(fun('type.quote.index',X15),X16),s(fun(X6,X15),X36),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X37),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X38),s(fun('type.ringNorm.polynom'(X6),X15),X39)))),
  file('thm.ringNorm.polynom_distinct|split|8','thm.ringNorm.polynom_case_def|split|2')).
fof(f104368,plain,(
  ( ! [X66] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.ringNorm.polynom'(sK0),'type.ringNorm.polynom'(sK0)),X66),s('type.ringNorm.polynom'(sK0),sK3)))) )),
  inference(forward_demodulation,[],[f103037,f245])).
fof(f103037,plain,(
  ( ! [X66] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.ringNorm.polynom'(sK0),'type.ringNorm.polynom'(sK0)),X66),s('type.ringNorm.polynom'(sK0),sK3)))) )),
  inference(superposition,[],[f50560,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X11,X24] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X11),s('type.ringNorm.polynom'(X6),X24))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X6),X11))),s('type.ringNorm.polynom'(X6),X24)))),
  file('thm.ringNorm.polynom_distinct|split|8','ARITY_0_2_const.ringNorm.Pplus')).
fof(f337,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X4),s('type.ringNorm.polynom'(X0),X2))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X3),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X4),s('type.ringNorm.polynom'(X0),X2))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X3),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X4),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X3),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ! [X41,X7,X40,X8] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X40),s('type.ringNorm.polynom'(X6),X41)))),
  file('thm.ringNorm.polynom_distinct|split|8','thm.ringNorm.polynom_distinct|split|7')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 538
% Passive clauses: 14435
% Generated clauses: 117898
% Final active clauses: 318
% Final passive clauses: 1374
% Input formulas: 69
% Initial clauses: 125
% Discarded non-redundant clauses: 4901
% 
% Duplicate literals: 2527
% Trivial inequalities: 705
% Fw subsumption resolutions: 4415
% Fw demodulations: 13635
% Bw demodulations: 8840
% Condensations: 5520
% Global subsumptions: 249
% 
% Simple tautologies: 628
% Equational tautologies: 584
% Forward subsumptions: 70961
% Fw demodulations to eq. taut.: 2986
% Bw demodulations to eq. taut.: 1258
% 
% Unit resulting resolution: 19961
% Factoring: 8
% Forward superposition: 29193
% Backward superposition: 18906
% Self superposition: 312
% Equality factoring: 138
% Equality resolution: 62
% 
% Split clauses: 236
% Split components: 493
% SAT solver clauses: 15632
% SAT solver unit clauses: 632
% SAT solver binary clauses: 2811
% 
% Sat splits: 236
% Sat splitting refutations: 474
% 
% Memory used [KB]: 70105
% Time elapsed: 8.910 s
% ------------------------------
% ------------------------------
% Success in time 10.231 s
