Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.combin
% (74668)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 307
% Passive clauses: 4106
% Generated clauses: 28015
% Final active clauses: 292
% Final passive clauses: 3640
% Input formulas: 69
% Initial clauses: 136
% 
% Duplicate literals: 94
% Trivial inequalities: 568
% Fw subsumption resolutions: 1973
% Fw demodulations: 5283
% Bw demodulations: 12
% 
% Simple tautologies: 162
% Equational tautologies: 23
% Forward subsumptions: 15270
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 660
% Factoring: 2
% Forward superposition: 7551
% Backward superposition: 11545
% Self superposition: 111
% Equality factoring: 16
% Equality resolution: 19
% 
% Split clauses: 87
% Split components: 183
% SAT solver clauses: 258
% SAT solver unit clauses: 6
% SAT solver binary clauses: 37
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 5
% 
% Sat splits: 87
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 255
% TWLsolver calls for satisfiability: 47
% 
% Memory used [KB]: 23283
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.combin
% (74683)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 175
% Passive clauses: 2192
% Generated clauses: 11743
% Final active clauses: 115
% Final passive clauses: 1667
% Input formulas: 69
% Initial clauses: 136
% 
% Duplicate literals: 3
% Trivial inequalities: 6
% Fw subsumption resolutions: 423
% Fw demodulations: 2157
% Bw demodulations: 406
% Condensations: 62
% Global subsumptions: 14
% 
% Simple tautologies: 46
% Equational tautologies: 10
% Forward subsumptions: 4300
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 658
% Factoring: 1
% Forward superposition: 2752
% Backward superposition: 2769
% Self superposition: 93
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 33
% Split components: 74
% SAT solver clauses: 2684
% SAT solver unit clauses: 166
% SAT solver binary clauses: 89
% SAT solver learnt clauses: 30
% SAT solver learnt literals: 6
% 
% Sat splits: 33
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2627
% TWLsolver calls for satisfiability: 754
% 
% Memory used [KB]: 24434
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.combin
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.combin
% SZS output start Proof for thm.combin
fof(f42720,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f274,f418,f416,f273,f420,f412,f272,f271,f417,f270,f269,f415,f268,f413,f419,f411,f278,f424,f422,f277,f425,f423,f432,f434,f1243,f1245,f1263,f1264,f1244,f1396,f1399,f1408,f1421,f1422,f1423,f1424,f1425,f1426,f1427,f1428,f1429,f1430,f1431,f1432,f1433,f1434,f1435,f1436,f1437,f3286,f3289,f3287,f20283,f20284,f20288,f20289,f20290,f20285,f20291,f29975,f29980,f29891,f29890,f29888,f29887,f30006,f414,f36716,f36823,f36821,f42427,f42713,f42717,f42719])).
fof(f42719,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f42718,f416])).
fof(f42718,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f417,f411])).
fof(f42717,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f42716,f416])).
fof(f42716,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f42715,f42714])).
fof(f42714,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl6 | $spl8)),
  inference(forward_demodulation,[],[f417,f419])).
fof(f42715,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f411,f419])).
fof(f42713,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f42712])).
fof(f42712,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl1 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f42711,f414])).
fof(f42711,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f412,f419])).
fof(f42427,plain,(
  $false | ($spl0 | $spl2 | $spl9)),
  inference(subsumption_resolution,[],[f42426,f420])).
fof(f42426,plain,(
  s(sK1,sK5) = s(sK1,sK6) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f42141,f6374])).
fof(f6374,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,happ(s(fun(X2,X0),'const.combin.UPDATE_3'(s(X2,X3),s(X0,X1),s(fun(X2,X0),X4))),s(X2,X3)))) )),
  inference(unit_resulting_resolution,[],[f389,f387])).
fof(f387,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,X1),X2) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))) | s(X1,X4) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : ((s(X1,X4) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X3))) | s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2)))) & (s(fun(X0,X1),X2) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))) | s(X1,X4) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X3)))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4] : (s(X1,X4) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X3))) <=> s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (s(X1,X4) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X3))) <=> s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X8,X9,X10] : (s(X7,X10) = s(X7,happ(s(fun(X6,X7),X8),s(X6,X9))) <=> s(fun(X6,X7),X8) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X10),s(fun(X6,X7),X8))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.combin.UPDATE_APPLY_ID')).
fof(f389,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2)))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2)))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),X8))) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X10),s(fun(X6,X7),X8)))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.combin.UPDATE_EQ')).
fof(f42141,plain,(
  s(sK1,sK6) = s(sK1,happ(s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))),s(sK0,sK3))) | ($spl0 | $spl2)),
  inference(superposition,[],[f6374,f30094])).
fof(f30094,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f414,f411])).
fof(f36821,plain,(
  ( ! [X220] : (p(s(bool,sK9('const.combin.I_0','const.combin.UPDATE_3'(s(bool,X220),s(bool,X220),s(fun(bool,bool),'const.combin.I_0')),bool)))) ) | $spl36),
  inference(cnf_transformation,[],[f36821_D])).
fof(f36821_D,plain,(
  ( ! [X220] : (p(s(bool,sK9('const.combin.I_0','const.combin.UPDATE_3'(s(bool,X220),s(bool,X220),s(fun(bool,bool),'const.combin.I_0')),bool)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f36823,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0')))) | $spl39),
  inference(cnf_transformation,[],[f36823_D])).
fof(f36823_D,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f36716,plain,(
  ( ! [X220] : (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0')))) | p(s(bool,sK9('const.combin.I_0','const.combin.UPDATE_3'(s(bool,X220),s(bool,X220),s(fun(bool,bool),'const.combin.I_0')),bool)))) )),
  inference(superposition,[],[f3587,f4318])).
fof(f4318,plain,(
  ( ! [X0,X1] : (s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'const.combin.UPDATE_3'(s(X0,X1),s(X0,X1),s(fun(X0,X0),'const.combin.I_0')))) )),
  inference(superposition,[],[f378,f427])).
fof(f427,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f331,f330])).
fof(f330,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6] : ! [X17] : s(X6,X17) = s(X6,'const.combin.I_1'(s(X6,X17)))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.combin.I_THM')).
fof(f331,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6] : ! [X13] : s(X6,'const.combin.I_1'(s(X6,X13))) = s(X6,happ(s(fun(X6,X6),'const.combin.I_0'),s(X6,X13)))),
  file('thm.combin.UPD|49||49|_SAME_BASE','ARITY_0_1_const.combin.I')).
fof(f378,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,happ(s(fun(X0,X1),X2),s(X0,X3))),s(fun(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,happ(s(fun(X0,X1),X2),s(X0,X3))),s(fun(X0,X1),X2)))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ! [X2,X3] : s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,happ(s(fun(X0,X1),X2),s(X0,X3))),s(fun(X0,X1),X2)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ! [X8,X9] : s(fun(X6,X7),X8) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,happ(s(fun(X6,X7),X8),s(X6,X9))),s(fun(X6,X7),X8)))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.combin.APPLY_UPDATE_ID')).
fof(f3587,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(bool,bool),X3),s(fun(bool,bool),'const.combin.I_0')))) | p(s(bool,sK9('const.combin.I_0',X3,bool)))) )),
  inference(superposition,[],[f345,f427])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK9(X2,X1,X0))))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK9(X2,X1,X0)),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK9(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X5,X35] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X5),s(fun(X6,bool),X35)))) <=> ? [X17] : (p(s(bool,'const.bool.IN_2'(s(X6,X17),s(fun(X6,bool),X5)))) & p(s(bool,happ(s(fun(X6,bool),X35),s(X6,X17))))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.RES_EXISTS_DEF')).
fof(f414,plain,(
  s(sK0,sK3) = s(sK0,sK4) | $spl2),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  s(sK0,sK3) = s(sK0,sK4) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f30006,plain,(
  ( ! [X1] : (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X1),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X1),s(fun(sK0,sK1),sK2)))))) ) | $spl34),
  inference(cnf_transformation,[],[f30006_D])).
fof(f30006_D,plain,(
  ( ! [X1] : (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X1),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X1),s(fun(sK0,sK1),sK2)))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f29887,plain,(
  ( ! [X1] : (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X1),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X1),s(fun(sK0,sK1),sK2))))) | s(sK0,sK3) = s(sK0,sK4)) ) | $spl0),
  inference(superposition,[],[f6504,f393])).
fof(f393,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))))) | s(X0,X3) = s(X0,X4)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(X0,X3) = s(X0,X4) | s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))))))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(X0,X3) != s(X0,X4) => s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : (~s(X0,X3) = s(X0,X4) => s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : ! [X8,X9,X10,X11,X12] : (~s(X6,X9) = s(X6,X10) => s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X10),s(X7,X12),s(fun(X6,X7),X8))))) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X10),s(X7,X12),s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),X8))))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.combin.UPDATE_COMMUTES')).
fof(f6504,plain,(
  ( ! [X0] : (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X0),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X0),s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2)))))) ) | $spl0),
  inference(superposition,[],[f389,f411])).
fof(f29888,plain,(
  ( ! [X2] : (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X2),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X2),s(fun(sK0,sK1),sK2))))) | s(sK0,sK3) = s(sK0,sK4)) ) | $spl0),
  inference(superposition,[],[f6504,f393])).
fof(f29890,plain,(
  ( ! [X0] : (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X0),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X0),s(fun(sK0,sK1),sK2))))) | s(sK0,sK3) = s(sK0,sK4)) ) | $spl0),
  inference(superposition,[],[f393,f6504])).
fof(f29891,plain,(
  ( ! [X1] : (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X1),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,X1),s(fun(sK0,sK1),sK2))))) | s(sK0,sK3) = s(sK0,sK4)) ) | $spl0),
  inference(superposition,[],[f393,f6504])).
fof(f29980,plain,(
  $false | ($spl0 | $spl3 | $spl7)),
  inference(subsumption_resolution,[],[f29979,f1053])).
fof(f1053,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl7)),
  inference(superposition,[],[f418,f411])).
fof(f29979,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f29978,f378])).
fof(f29978,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK4))),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f29977,f6890])).
fof(f6890,plain,(
  ( ! [X2,X0,X1] : (s(X2,happ(s(fun(sK0,X2),X0),s(sK0,sK4))) = s(X2,'const.combin.UPDATE_4'(s(sK0,sK3),s(X2,X1),s(fun(sK0,X2),X0),s(sK0,sK4)))) ) | $spl3),
  inference(forward_demodulation,[],[f6697,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,X19) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X18),s(X6,X19)))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.COND_CLAUSES|split|1')).
fof(f6697,plain,(
  ( ! [X2,X0,X1] : (s(X2,'const.combin.UPDATE_4'(s(sK0,sK3),s(X2,X1),s(fun(sK0,X2),X0),s(sK0,sK4))) = s(X2,'const.bool.COND_3'(s(bool,'F_0'),s(X2,X1),s(X2,happ(s(fun(sK0,X2),X0),s(sK0,sK4)))))) ) | $spl3),
  inference(superposition,[],[f392,f1129])).
fof(f1129,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,sK19(sK4,X0,X1,sK3,sK0,X2))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1108,f295])).
fof(f295,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.EQ_CLAUSES|split|2')).
fof(f1108,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19(sK4,X0,X1,sK3,sK0,X2)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f413,f390])).
fof(f390,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19(X5,X4,X3,X2,X1,X0))) | s(X1,X2) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,sK19(X5,X4,X3,X2,X1,X0))) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,sK19(X5,X4,X3,X2,X1,X0)))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,sK19(X5,X4,X3,X2,X1,X0)),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((~p(s(bool,X6)) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,X6))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : (((~p(s(bool,X6)) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,X6)))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X2) = s(X1,X5)) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X2) = s(X1,X5)) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7,X6] : ! [X9,X10,X5,X22] : ? [X23] : ((p(s(bool,X23)) <=> s(X6,X9) = s(X6,X22)) & s(X7,'const.combin.UPDATE_4'(s(X6,X9),s(X7,X10),s(fun(X6,X7),X5),s(X6,X22))) = s(X7,'const.bool.COND_3'(s(bool,X23),s(X7,X10),s(X7,happ(s(fun(X6,X7),X5),s(X6,X22))))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.combin.UPDATE_def')).
fof(f392,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,sK19(X5,X4,X3,X2,X1,X0)),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f29977,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2),s(sK0,sK4))),s(fun(sK0,sK1),sK2))) | $spl0),
  inference(forward_demodulation,[],[f29889,f388])).
fof(f388,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X6] : ! [X13,X14,X15,X16] : s(X7,'const.combin.UPDATE_4'(s(X6,X13),s(X7,X14),s(fun(X6,X7),X15),s(X6,X16))) = s(X7,happ(s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X13),s(X7,X14),s(fun(X6,X7),X15))),s(X6,X16)))),
  file('thm.combin.UPD|49||49|_SAME_BASE','ARITY_3_4_const.combin.UPDATE')).
fof(f29889,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,happ(s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))),s(sK0,sK4))),s(fun(sK0,sK1),sK2))) | $spl0),
  inference(superposition,[],[f378,f6504])).
fof(f29975,plain,(
  $false | ($spl0 | $spl3 | $spl7)),
  inference(subsumption_resolution,[],[f29974,f1053])).
fof(f29974,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f29973,f378])).
fof(f29973,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK4))),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f29972,f6890])).
fof(f29972,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2),s(sK0,sK4))),s(fun(sK0,sK1),sK2))) | $spl0),
  inference(forward_demodulation,[],[f29886,f388])).
fof(f29886,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,happ(s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))),s(sK0,sK4))),s(fun(sK0,sK1),sK2))) | $spl0),
  inference(superposition,[],[f6504,f378])).
fof(f20291,plain,(
  ( ! [X82] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X82))),s(bool,'T_0'))))) ) | $spl32),
  inference(cnf_transformation,[],[f20291_D])).
fof(f20291_D,plain,(
  ( ! [X82] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X82))),s(bool,'T_0'))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f20285,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,sK10(X48,X49,X50,X51))) | s(bool,'T_0') = s(bool,X50)) ) | $spl28),
  inference(cnf_transformation,[],[f20285_D])).
fof(f20285_D,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,sK10(X48,X49,X50,X51))) | s(bool,'T_0') = s(bool,X50)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f20290,plain,(
  ( ! [X80,X78,X81,X79,X82] : (p(s(bool,sK10(X78,X79,X80,X81))) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X82))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X80)) ) | $spl3),
  inference(forward_demodulation,[],[f20180,f330])).
fof(f20180,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X82))),s(bool,'T_0')))) | p(s(bool,'const.combin.I_1'(s(bool,sK10(X78,X79,X80,X81))))) | s(bool,'T_0') = s(bool,X80)) ) | $spl3),
  inference(superposition,[],[f371,f19682])).
fof(f19682,plain,(
  ( ! [X6,X4,X0,X5] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK10(X4,X5,X0,X6)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(forward_demodulation,[],[f19681,f330])).
fof(f19681,plain,(
  ( ! [X6,X4,X0,X5] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK10(X4,X5,X0,X6)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl3),
  inference(subsumption_resolution,[],[f19573,f1129])).
fof(f19573,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') != s(bool,sK19(sK4,X1,X2,sK3,sK0,X3)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK10(X4,X5,X0,X6)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl3),
  inference(superposition,[],[f18212,f16205])).
fof(f16205,plain,(
  ( ! [X389,X387,X388,X390] : (s(bool,X387) = s(bool,sK19(sK4,X388,X389,sK3,sK0,X390)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X387)))) ) | $spl3),
  inference(superposition,[],[f16113,f330])).
fof(f16113,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,X125) = s(bool,sK19(sK4,X122,X123,sK3,sK0,X124)) | s(bool,'T_0') = s(bool,X125)) ) | $spl3),
  inference(subsumption_resolution,[],[f16074,f426])).
fof(f426,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f280,f275])).
fof(f275,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f12])).
fof(f12,axiom,(
  p(s(bool,'T_0'))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.TRUTH')).
fof(f280,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f35])).
fof(f35,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.NOT_CLAUSES|split|2')).
fof(f16074,plain,(
  ( ! [X125,X123,X124,X122] : (p(s(bool,'F_0')) | s(bool,X125) = s(bool,sK19(sK4,X122,X123,sK3,sK0,X124)) | s(bool,'T_0') = s(bool,X125)) ) | $spl3),
  inference(superposition,[],[f674,f1129])).
fof(f674,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f662,f426])).
fof(f662,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f356,f287])).
fof(f287,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.combin.UPD|49||49|_SAME_BASE',aHL_BOOL_CASES)).
fof(f356,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X18,X19] : ((p(s(bool,X18)) => p(s(bool,X19))) => ((p(s(bool,X19)) => p(s(bool,X18))) => s(bool,X18) = s(bool,X19)))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.IMP_ANTISYM_AX')).
fof(f18212,plain,(
  ( ! [X420,X418,X421,X419] : (s(bool,'F_0') != s(bool,X420) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK10(X418,X419,X420,X421))))) ) | $spl3),
  inference(subsumption_resolution,[],[f18211,f426])).
fof(f18211,plain,(
  ( ! [X420,X418,X421,X419] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X420) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK10(X418,X419,X420,X421))))) ) | $spl3),
  inference(forward_demodulation,[],[f17698,f1129])).
fof(f17698,plain,(
  ( ! [X420,X422,X418,X424,X421,X423,X419] : (p(s(bool,sK19(sK4,X422,X423,sK3,sK0,X424))) | s(bool,'F_0') != s(bool,X420) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK10(X418,X419,X420,X421))))) ) | $spl3),
  inference(superposition,[],[f348,f16205])).
fof(f348,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK10(X3,X2,X1,X0))) | s(bool,'F_0') != s(bool,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,sK10(X3,X2,X1,X0))) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,sK10(X3,X2,X1,X0)))) & ((~p(s(bool,sK11(X3,X2,X1,X0))) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,sK11(X3,X2,X1,X0)))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,sK11(X3,X2,X1,X0)))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,sK10(X3,X2,X1,X0)))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,X4))) & ? [X5] : ((~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,X5))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (((~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) & (s(bool,'F_0') != s(bool,X1) | p(s(bool,X4)))) & ? [X5] : (((~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X1)) & (s(bool,'T_0') != s(bool,X1) | p(s(bool,X5)))) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s(bool,'F_0') = s(bool,X1)) & ? [X5] : ((p(s(bool,X5)) <=> s(bool,'T_0') = s(bool,X1)) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s(bool,'F_0') = s(bool,X1)) & ? [X5] : ((p(s(bool,X5)) <=> s(bool,'T_0') = s(bool,X1)) & s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.min.@_1'(s(fun(X0,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X5))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X4))),s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3)))))))))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X5,X22,X35] : ? [X23] : ((p(s(bool,X23)) <=> s(bool,'F_0') = s(bool,X5)) & ? [X44] : ((p(s(bool,X44)) <=> s(bool,'T_0') = s(bool,X5)) & s(X6,'const.bool.COND_3'(s(bool,X5),s(X6,X22),s(X6,X35))) = s(X6,'const.min.@_1'(s(fun(X6,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X44))),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(X6,X22))))),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'==>_1'(s(bool,X23))),s(fun(X6,bool),'_dst_x1x0_2'(s(fun(X6,fun(X6,bool)),'=_0'),s(X6,X35)))))))))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.COND_DEF')).
fof(f371,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & ((~p(s(bool,X0)) | ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) & p(s(bool,X1))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X46,X47] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X46))),s(bool,X47)))) <=> (p(s(bool,X46)) & p(s(bool,X47))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','|slash||bslash|_HL_0')).
fof(f20289,plain,(
  ( ! [X67] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,X67))))) ) | $spl30),
  inference(cnf_transformation,[],[f20289_D])).
fof(f20289_D,plain,(
  ( ! [X67] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,X67))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f20288,plain,(
  ( ! [X66,X64,X67,X65,X63] : (p(s(bool,sK10(X63,X64,X65,X66))) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,X67)))) | s(bool,'T_0') = s(bool,X65)) ) | $spl3),
  inference(forward_demodulation,[],[f20177,f330])).
fof(f20177,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,X67)))) | p(s(bool,'const.combin.I_1'(s(bool,sK10(X63,X64,X65,X66))))) | s(bool,'T_0') = s(bool,X65)) ) | $spl3),
  inference(superposition,[],[f370,f19682])).
fof(f370,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f255])).
fof(f20284,plain,(
  ( ! [X52] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X52))))) ) | $spl26),
  inference(cnf_transformation,[],[f20284_D])).
fof(f20284_D,plain,(
  ( ! [X52] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X52))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f20283,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,sK10(X48,X49,X50,X51))) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X52)))) | s(bool,'T_0') = s(bool,X50)) ) | $spl3),
  inference(forward_demodulation,[],[f20174,f330])).
fof(f20174,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,'T_0'))),s(bool,X52)))) | p(s(bool,'const.combin.I_1'(s(bool,sK10(X48,X49,X50,X51))))) | s(bool,'T_0') = s(bool,X50)) ) | $spl3),
  inference(superposition,[],[f368,f19682])).
fof(f368,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X46,X47] : (p(s(bool,happ(s(fun(bool,bool),'==>_1'(s(bool,X46))),s(bool,X47)))) <=> (p(s(bool,X46)) => p(s(bool,X47))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','==>_HL_1')).
fof(f3287,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl22),
  inference(cnf_transformation,[],[f3287_D])).
fof(f3287_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3289,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl24),
  inference(cnf_transformation,[],[f3289_D])).
fof(f3289_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3286,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f3270,f427])).
fof(f3270,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f338,f427])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X39,X17] : (p(s(bool,happ(s(fun(X6,bool),X39),s(X6,X17)))) => p(s(bool,happ(s(fun(X6,bool),X39),s(X6,'const.min.@_1'(s(fun(X6,bool),X39)))))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.SELECT_AX')).
fof(f1437,plain,(
  $false | ($spl3 | $spl16)),
  inference(subsumption_resolution,[],[f1390,f275])).
fof(f1390,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1204])).
fof(f1204,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19(sK19(sK4,X0,X1,sK3,sK0,X2),X3,X4,'T_0',bool,X5)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1130,f390])).
fof(f1130,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK19(sK4,X0,X1,sK3,sK0,X2))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1108,f290])).
fof(f290,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.EQ_CLAUSES|split|0')).
fof(f1436,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1386,f275])).
fof(f1386,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f1244,f1112])).
fof(f1112,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19(sK11(X0,X1,'F_0',X2),X3,X4,'T_0',bool,X5)))) )),
  inference(unit_resulting_resolution,[],[f629,f390])).
fof(f629,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK11(X0,X1,'F_0',X2))) )),
  inference(unit_resulting_resolution,[],[f625,f290])).
fof(f625,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK11(X0,X1,'F_0',X2)))) )),
  inference(unit_resulting_resolution,[],[f284,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK11(X3,X2,X1,X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f284,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(flattening,[],[f69])).
fof(f69,axiom,(
  ~s(bool,'T_0') = s(bool,'F_0')),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.BOOL_EQ_DISTINCT|split|1')).
fof(f1435,plain,(
  $false | ($spl11 | $spl16)),
  inference(subsumption_resolution,[],[f1385,f275])).
fof(f1385,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1113])).
fof(f1113,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19(sK11(X0,X1,sK7,X2),X3,X4,'T_0',bool,X5)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f633,f390])).
fof(f633,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK11(X0,X1,sK7,X2))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f626,f290])).
fof(f626,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK11(X0,X1,sK7,X2)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f436,f349])).
fof(f436,plain,(
  s(bool,'T_0') != s(bool,sK7) | $spl11),
  inference(unit_resulting_resolution,[],[f422,f290])).
fof(f1434,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1382,f275])).
fof(f1382,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f1244,f1111])).
fof(f1111,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19(sK10(X0,X1,'T_0',X2),X3,X4,'T_0',bool,X5)))) )),
  inference(unit_resulting_resolution,[],[f577,f390])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK10(X0,X1,'T_0',X2))) )),
  inference(unit_resulting_resolution,[],[f574,f290])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK10(X0,X1,'T_0',X2)))) )),
  inference(unit_resulting_resolution,[],[f284,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK10(X3,X2,X1,X0))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f1433,plain,(
  $false | ($spl11 | $spl16)),
  inference(subsumption_resolution,[],[f1380,f275])).
fof(f1380,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1110])).
fof(f1110,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19(sK7,X0,X1,'T_0',bool,X2)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f436,f390])).
fof(f1432,plain,(
  $false | ($spl7 | $spl16)),
  inference(subsumption_resolution,[],[f1379,f275])).
fof(f1379,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1122])).
fof(f1122,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19(sK2,X0,X1,'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)),fun(sK0,sK1),X2)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f418,f390])).
fof(f1431,plain,(
  $false | ($spl0 | $spl7 | $spl16)),
  inference(subsumption_resolution,[],[f1378,f275])).
fof(f1378,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl7 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1123])).
fof(f1123,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19(sK2,X0,X1,'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2)),fun(sK0,sK1),X2)))) ) | ($spl0 | $spl7)),
  inference(unit_resulting_resolution,[],[f1053,f390])).
fof(f1430,plain,(
  $false | ($spl7 | $spl16)),
  inference(subsumption_resolution,[],[f1370,f275])).
fof(f1370,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1114])).
fof(f1114,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19('const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)),X0,X1,sK2,fun(sK0,sK1),X2)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f418,f390])).
fof(f1429,plain,(
  $false | ($spl0 | $spl7 | $spl16)),
  inference(subsumption_resolution,[],[f1369,f275])).
fof(f1369,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl7 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1115])).
fof(f1115,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19('const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2)),X0,X1,sK2,fun(sK0,sK1),X2)))) ) | ($spl0 | $spl7)),
  inference(unit_resulting_resolution,[],[f1053,f390])).
fof(f1428,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1363,f275])).
fof(f1363,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f1244,f1109])).
fof(f1109,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19('F_0',X0,X1,'T_0',bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f284,f390])).
fof(f1427,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1361,f275])).
fof(f1361,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f1244,f1117])).
fof(f1117,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19('T_0',X0,X1,'F_0',bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f284,f390])).
fof(f1426,plain,(
  $false | ($spl11 | $spl16)),
  inference(subsumption_resolution,[],[f1360,f275])).
fof(f1360,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1118])).
fof(f1118,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK19('T_0',X0,X1,sK7,bool,X2)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f436,f390])).
fof(f1425,plain,(
  $false | ($spl3 | $spl16)),
  inference(subsumption_resolution,[],[f1359,f275])).
fof(f1359,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1206])).
fof(f1206,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19('T_0',X0,X1,sK19(sK4,X2,X3,sK3,sK0,X4),bool,X5)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1130,f390])).
fof(f1424,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1357,f275])).
fof(f1357,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f1244,f1119])).
fof(f1119,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19('T_0',X0,X1,sK10(X2,X3,'T_0',X4),bool,X5)))) )),
  inference(unit_resulting_resolution,[],[f577,f390])).
fof(f1423,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1356,f275])).
fof(f1356,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f1244,f1120])).
fof(f1120,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19('T_0',X0,X1,sK11(X2,X3,'F_0',X4),bool,X5)))) )),
  inference(unit_resulting_resolution,[],[f629,f390])).
fof(f1422,plain,(
  $false | ($spl11 | $spl16)),
  inference(subsumption_resolution,[],[f1355,f275])).
fof(f1355,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1121])).
fof(f1121,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK19('T_0',X0,X1,sK11(X2,X3,sK7,X4),bool,X5)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f633,f390])).
fof(f1421,plain,(
  $false | ($spl3 | $spl16)),
  inference(subsumption_resolution,[],[f1347,f275])).
fof(f1347,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl16)),
  inference(backward_demodulation,[],[f1244,f1200])).
fof(f1200,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK11(X0,X1,sK19(sK4,X2,X3,sK3,sK0,X4),X5)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1130,f349])).
fof(f1408,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1407,f275])).
fof(f1407,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(forward_demodulation,[],[f1311,f536])).
fof(f536,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f517,f291])).
fof(f291,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f517,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f275,f428])).
fof(f428,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))) )),
  inference(subsumption_resolution,[],[f365,f366])).
fof(f366,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK14(X1,X0))) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | ! [X3] : ((p(s(bool,X0)) & ~p(s(bool,X3))) | (p(s(bool,X1)) & ~p(s(bool,X3))) | p(s(bool,X3)))) & (((~p(s(bool,X0)) | p(s(bool,sK14(X1,X0)))) & (~p(s(bool,X1)) | p(s(bool,sK14(X1,X0)))) & ~p(s(bool,sK14(X1,X0)))) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | ! [X3] : ((p(s(bool,X0)) & ~p(s(bool,X3))) | (p(s(bool,X1)) & ~p(s(bool,X3))) | p(s(bool,X3)))) & (? [X2] : ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))) & ~p(s(bool,X2))) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))))),
  inference(rectify,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | ! [X2] : ((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2))) | p(s(bool,X2)))) & (? [X2] : ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))) & ~p(s(bool,X2))) | p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) <=> ! [X2] : ((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2))) | p(s(bool,X2))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) <=> ! [X2] : ((p(s(bool,X0)) & ~p(s(bool,X2))) | ((p(s(bool,X1)) & ~p(s(bool,X2))) | p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) <=> ! [X2] : ((p(s(bool,X0)) => p(s(bool,X2))) => ((p(s(bool,X1)) => p(s(bool,X2))) => p(s(bool,X2)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X5,X35] : (p(s(bool,'|dollar|or_2'(s(bool,X5),s(bool,X35)))) <=> ! [X21] : ((p(s(bool,X5)) => p(s(bool,X21))) => ((p(s(bool,X35)) => p(s(bool,X21))) => p(s(bool,X21)))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.OR_DEF')).
fof(f365,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1)))) | p(s(bool,sK14(X1,X0))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f251])).
fof(f1311,plain,(
  ~p(s(bool,'|dollar|or_2'(s(bool,'T_0'),s(bool,'T_0')))) | $spl16),
  inference(backward_demodulation,[],[f1244,f1212])).
fof(f1212,plain,(
  ~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f426,f426,f403])).
fof(f403,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,X0)))) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(condensation,[],[f360])).
fof(f360,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,'|dollar|or_2'(s(bool,X0),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f1399,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1287,f275])).
fof(f1287,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f1244,f426])).
fof(f1396,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1395,f1244])).
fof(f1395,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl16),
  inference(subsumption_resolution,[],[f1282,f290])).
fof(f1282,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(backward_demodulation,[],[f1244,f294])).
fof(f294,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f1244,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f1244_D])).
fof(f1244_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1264,plain,(
  ( ! [X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl20),
  inference(cnf_transformation,[],[f1264_D])).
fof(f1264_D,plain,(
  ( ! [X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'F_0')))) | p(s(bool,X1))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1263,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'F_0')))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1227,f291])).
fof(f1227,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,X1),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f403,f287])).
fof(f1245,plain,(
  ( ! [X1] : (~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,X1)))) | p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f1245_D])).
fof(f1245_D,plain,(
  ( ! [X1] : (~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,X1)))) | p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1243,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,X1)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1213,f291])).
fof(f1213,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|or_2'(s(bool,'F_0'),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f403,f287])).
fof(f434,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f433,f418])).
fof(f433,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f411,f415])).
fof(f432,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f431,f415])).
fof(f431,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f412,f417])).
fof(f423,plain,(
  ~p(s(bool,'F_0')) | $spl13),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f425,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f277,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (~p(s(bool,sK7)) | p(s(bool,'F_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f205])).
fof(f205,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(rectify,[],[f204])).
fof(f204,plain,(
  (~p(s(bool,'F_0')) | ! [X0] : p(s(bool,X0))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  p(s(bool,'F_0')) <=> ! [X0] : p(s(bool,X0))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  p(s(bool,'F_0')) <=> ! [X21] : p(s(bool,X21))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.bool.F_DEF')).
fof(f422,plain,(
  ~p(s(bool,sK7)) | $spl11),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  ~p(s(bool,sK7)) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f424,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(cnf_transformation,[],[f424_D])).
fof(f424_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f278,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK7))),
  inference(cnf_transformation,[],[f206])).
fof(f411,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | $spl0),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f419,plain,(
  s(sK1,sK5) = s(sK1,sK6) | $spl8),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  s(sK1,sK5) = s(sK1,sK6) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f413,plain,(
  s(sK0,sK3) != s(sK0,sK4) | $spl3),
  inference(cnf_transformation,[],[f413_D])).
fof(f413_D,plain,(
  s(sK0,sK3) != s(sK0,sK4) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f268,plain,(
  s(sK0,sK3) != s(sK0,sK4) | s(sK1,sK5) = s(sK1,sK6) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | (s(sK0,sK3) = s(sK0,sK4) & s(sK1,sK5) = s(sK1,sK6)) | (s(sK0,sK3) != s(sK0,sK4) & s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) & s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))))) & (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | ((s(sK0,sK3) != s(sK0,sK4) | s(sK1,sK5) != s(sK1,sK6)) & (s(sK0,sK3) = s(sK0,sK4) | s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f202])).
fof(f202,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))) | (s(X0,X3) = s(X0,X4) & s(X1,X5) = s(X1,X6)) | (s(X0,X3) != s(X0,X4) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))))) & (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))) | ((s(X0,X3) != s(X0,X4) | s(X1,X5) != s(X1,X6)) & (s(X0,X3) = s(X0,X4) | s(fun(X0,X1),X2) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) | s(fun(X0,X1),X2) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2)))))))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))) | ((s(X0,X3) = s(X0,X4) & s(X1,X5) = s(X1,X6)) | (s(X0,X3) != s(X0,X4) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2)))))) & (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))) | ((s(X0,X3) != s(X0,X4) | s(X1,X5) != s(X1,X6)) & (s(X0,X3) = s(X0,X4) | s(fun(X0,X1),X2) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) | s(fun(X0,X1),X2) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2)))))))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))) <~> ((s(X0,X3) = s(X0,X4) & s(X1,X5) = s(X1,X6)) | (s(X0,X3) != s(X0,X4) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))) <=> ((s(X0,X3) = s(X0,X4) & s(X1,X5) = s(X1,X6)) | (s(X0,X3) != s(X0,X4) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5,X6] : (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))) <=> ((s(X0,X3) = s(X0,X4) & s(X1,X5) = s(X1,X6)) | (~s(X0,X3) = s(X0,X4) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) & s(fun(X0,X1),X2) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X4),s(X1,X6),s(fun(X0,X1),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11,X12] : (s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),X8))) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X10),s(X7,X12),s(fun(X6,X7),X8))) <=> ((s(X6,X9) = s(X6,X10) & s(X7,X11) = s(X7,X12)) | (~s(X6,X9) = s(X6,X10) & s(fun(X6,X7),X8) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),X8))) & s(fun(X6,X7),X8) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X10),s(X7,X12),s(fun(X6,X7),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11,X12] : (s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),X8))) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X10),s(X7,X12),s(fun(X6,X7),X8))) <=> ((s(X6,X9) = s(X6,X10) & s(X7,X11) = s(X7,X12)) | (~s(X6,X9) = s(X6,X10) & s(fun(X6,X7),X8) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),X8))) & s(fun(X6,X7),X8) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X10),s(X7,X12),s(fun(X6,X7),X8))))))),
  file('thm.combin.UPD|49||49|_SAME_BASE','thm.combin.UPD|49||49|_SAME_BASE')).
fof(f415,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | $spl4),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f269,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | s(sK0,sK3) = s(sK0,sK4) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f203])).
fof(f270,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | s(sK1,sK5) = s(sK1,sK6) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f203])).
fof(f417,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | $spl6),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f271,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | s(sK0,sK3) = s(sK0,sK4) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f203])).
fof(f272,plain,(
  s(fun(sK0,sK1),sK2) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | s(sK1,sK5) = s(sK1,sK6) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f203])).
fof(f412,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | $spl1),
  inference(cnf_transformation,[],[f412_D])).
fof(f412_D,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f420,plain,(
  s(sK1,sK5) != s(sK1,sK6) | $spl9),
  inference(cnf_transformation,[],[f420_D])).
fof(f420_D,plain,(
  s(sK1,sK5) != s(sK1,sK6) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f273,plain,(
  s(sK1,sK5) != s(sK1,sK6) | s(sK0,sK3) != s(sK0,sK4) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f203])).
fof(f416,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | $spl5),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f418,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | $spl7),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f274,plain,(
  s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2))) | s(fun(sK0,sK1),sK2) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | s(sK0,sK3) = s(sK0,sK4) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK4),s(sK1,sK6),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f203])).
% SZS output end Proof for thm.combin
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 431
% Passive clauses: 5324
% Generated clauses: 42811
% Final active clauses: 108
% Final passive clauses: 1979
% Input formulas: 69
% Initial clauses: 136
% 
% Duplicate literals: 415
% Trivial inequalities: 298
% Fw subsumption resolutions: 1411
% Fw demodulations: 11434
% Bw demodulations: 1603
% Condensations: 761
% 
% Simple tautologies: 148
% Equational tautologies: 154
% Forward subsumptions: 22774
% Fw demodulations to eq. taut.: 44
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 6519
% Factoring: 1
% Forward superposition: 9697
% Backward superposition: 9855
% Self superposition: 167
% Equality factoring: 34
% Equality resolution: 14
% 
% Split clauses: 20
% Split components: 48
% SAT solver clauses: 5394
% SAT solver unit clauses: 1036
% SAT solver binary clauses: 2678
% 
% Sat splits: 20
% Sat splitting refutations: 28
% 
% Memory used [KB]: 46694
% Time elapsed: 1.806 s
% ------------------------------
% ------------------------------
% Success in time 3.13 s
