Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f20476,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f20403])).
fof(f20403,plain,(
  s(fun(sK12,fun(sK12,bool)),'const.relation.EQC_1'(s(fun(sK12,fun(sK12,bool)),sK13))) != s(fun(sK12,fun(sK12,bool)),'const.relation.EQC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))),
  inference(superposition,[],[f292,f19840])).
fof(f19840,plain,(
  ( ! [X2,X3] : (s(fun(X2,fun(X2,bool)),'const.relation.EQC_1'(s(fun(X2,fun(X2,bool)),X3))) = s(fun(X2,fun(X2,bool)),'const.relation.EQC_1'(s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(forward_demodulation,[],[f5751,f5729])).
fof(f5729,plain,(
  ( ! [X4,X5] : (s(fun(X4,fun(X4,bool)),'const.relation.EQC_1'(s(fun(X4,fun(X4,bool)),X5))) = s(fun(X4,fun(X4,bool)),'const.relation.RC_1'(s(fun(X4,fun(X4,bool)),'const.relation.EQC_1'(s(fun(X4,fun(X4,bool)),X5)))))) )),
  inference(superposition,[],[f339,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X7)))))))),
  file('thm.relation.EQC_MOVES_IN|split|0','thm.relation.EQC_DEF')).
fof(f339,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.EQC_MOVES_IN|split|0','thm.relation.RC_IDEM')).
fof(f5751,plain,(
  ( ! [X2,X3] : (s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),'const.relation.EQC_1'(s(fun(X2,fun(X2,bool)),X3))))) = s(fun(X2,fun(X2,bool)),'const.relation.EQC_1'(s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),X3)))))) )),
  inference(forward_demodulation,[],[f5750,f342])).
fof(f5750,plain,(
  ( ! [X2,X3] : (s(fun(X2,fun(X2,bool)),'const.relation.EQC_1'(s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),X3))))) = s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),'const.relation.TC_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))))))) )),
  inference(forward_demodulation,[],[f5726,f340])).
fof(f340,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X7))))) = s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.EQC_MOVES_IN|split|0','thm.relation.RC_MOVES_OUT|split|2')).
fof(f5726,plain,(
  ( ! [X2,X3] : (s(fun(X2,fun(X2,bool)),'const.relation.EQC_1'(s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),X3))))) = s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),'const.relation.TC_1'(s(fun(X2,fun(X2,bool)),'const.relation.RC_1'(s(fun(X2,fun(X2,bool)),'const.relation.SC_1'(s(fun(X2,fun(X2,bool)),X3)))))))))) )),
  inference(superposition,[],[f342,f341])).
fof(f341,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),X1))))) = s(fun(X0,fun(X0,bool)),'const.relation.SC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),X7))))) = s(fun(X6,fun(X6,bool)),'const.relation.SC_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.EQC_MOVES_IN|split|0','thm.relation.RC_MOVES_OUT|split|0')).
fof(f292,plain,(
  s(fun(sK12,fun(sK12,bool)),'const.relation.EQC_1'(s(fun(sK12,fun(sK12,bool)),sK13))) != s(fun(sK12,fun(sK12,bool)),'const.relation.EQC_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.RC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  s(fun(sK12,fun(sK12,bool)),'const.relation.EQC_1'(s(fun(sK12,fun(sK12,bool)),sK13))) != s(fun(sK12,fun(sK12,bool)),'const.relation.EQC_1'(s(fun(sK12,fun(sK12,bool)),'const.relation.RC_1'(s(fun(sK12,fun(sK12,bool)),sK13)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f148])).
fof(f148,plain,(
  ? [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) != s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ~! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),X1))) = s(fun(X0,fun(X0,bool)),'const.relation.EQC_1'(s(fun(X0,fun(X0,bool)),'const.relation.RC_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),X7))) = s(fun(X6,fun(X6,bool)),'const.relation.EQC_1'(s(fun(X6,fun(X6,bool)),'const.relation.RC_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.relation.EQC_MOVES_IN|split|0','thm.relation.EQC_MOVES_IN|split|0')).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 752
% Passive clauses: 6124
% Generated clauses: 19856
% Final active clauses: 747
% Final passive clauses: 5128
% Input formulas: 56
% Initial clauses: 218
% 
% Duplicate literals: 590
% Trivial inequalities: 46
% Fw subsumption resolutions: 1734
% Fw demodulations: 1786
% Bw demodulations: 2
% 
% Simple tautologies: 274
% Equational tautologies: 57
% Forward subsumptions: 8681
% 
% Binary resolution: 7264
% Factoring: 1232
% Forward superposition: 3024
% Backward superposition: 3683
% Self superposition: 42
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 480
% Split components: 975
% SAT solver clauses: 742
% SAT solver unit clauses: 12
% SAT solver binary clauses: 119
% 
% Sat splits: 480
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 742
% TWLsolver calls for satisfiability: 379
% 
% Memory used [KB]: 10490
% Time elapsed: 0.550 s
% ------------------------------
% ------------------------------
% Success in time 0.553 s
