Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (45652)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 72
% Passive clauses: 2771
% Generated clauses: 15686
% Final active clauses: 70
% Final passive clauses: 2625
% Input formulas: 23
% Initial clauses: 24
% 
% Duplicate literals: 898
% Trivial inequalities: 146
% Fw subsumption resolutions: 1571
% Fw demodulations: 598
% 
% Simple tautologies: 292
% Equational tautologies: 114
% Forward subsumptions: 7312
% 
% Binary resolution: 25
% Forward superposition: 6056
% Backward superposition: 5497
% Self superposition: 440
% Equality factoring: 340
% Equality resolution: 10
% 
% Split clauses: 667
% Split components: 1362
% SAT solver clauses: 1844
% SAT solver unit clauses: 451
% SAT solver binary clauses: 584
% 
% Sat splits: 667
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1834
% TWLsolver calls for satisfiability: 47
% 
% Memory used [KB]: 11769
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (45849)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 39
% Passive clauses: 271
% Generated clauses: 6466
% Final active clauses: 36
% Final passive clauses: 230
% Input formulas: 23
% Initial clauses: 24
% 
% Duplicate literals: 919
% Trivial inequalities: 57
% Fw subsumption resolutions: 154
% Fw demodulations: 56
% Bw demodulations: 1
% Condensations: 1787
% 
% Simple tautologies: 119
% Equational tautologies: 77
% Forward subsumptions: 2264
% 
% Binary resolution: 16
% Forward superposition: 1855
% Backward superposition: 1665
% Self superposition: 272
% Equality factoring: 222
% Equality resolution: 5
% 
% Split clauses: 29
% Split components: 58
% SAT solver clauses: 703
% SAT solver unit clauses: 42
% SAT solver binary clauses: 66
% 
% Sat splits: 29
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 703
% TWLsolver calls for satisfiability: 57
% 
% Memory used [KB]: 6268
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f67725,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f67724])).
fof(f67724,plain,(
  s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),sK4))) != s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),sK4)))),
  inference(superposition,[],[f71,f12844])).
fof(f12844,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(X7,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X7),X8),s('type.quote.varmap'(X7),X9),s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11))) = s(X7,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X7),X8),s('type.quote.varmap'(X7),X9),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11)))))) )),
  inference(forward_demodulation,[],[f12685,f11566])).
fof(f11566,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(forward_demodulation,[],[f11382,f92])).
fof(f92,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X17,X18,X19,X20] : s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X17),s('type.quote.varmap'(X6),X18),s('type.quote.index',X19),s('type.list.list'('type.quote.index'),X20))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6)),happ(s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X6),X17))),s('type.quote.varmap'(X6),X18))),s('type.quote.index',X19))),s('type.list.list'('type.quote.index'),X20)))),
  file('thm.ringNorm.interp_vl_def|split|1','ARITY_1_4_const.ringNorm.r_ivl_aux')).
fof(f11382,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(superposition,[],[f93,f76])).
fof(f76,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_vl_def|split|1','thm.ringNorm.r_ivl_aux_def')).
fof(f93,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X17,X18,X19,X20] : s(X6,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X6),X17),s('type.quote.varmap'(X6),X18),s('type.quote.index',X19),s('type.list.list'('type.quote.index'),X20))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6)),happ(s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X6),X17))),s('type.quote.varmap'(X6),X18))),s('type.quote.index',X19))),s('type.list.list'('type.quote.index'),X20)))),
  file('thm.ringNorm.interp_vl_def|split|1','ARITY_1_4_const.canonical.ivl_aux')).
fof(f12685,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(X7,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X7),X8),s('type.quote.varmap'(X7),X9),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11))))) = s(X7,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X7),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X7),X8))),s('type.quote.varmap'(X7),X9),s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11)))) )),
  inference(superposition,[],[f8636,f91])).
fof(f91,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X12,X8,X9,X10] : s(X6,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X6),X12),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),X10))) = s(X6,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X6),X12),s('type.quote.varmap'(X6),X8),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X9),s('type.list.list'('type.quote.index'),X10)))))),
  file('thm.ringNorm.interp_vl_def|split|1','thm.canonical.interp_vl_def|split|1')).
fof(f8636,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(forward_demodulation,[],[f8508,f86])).
fof(f86,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X17,X18,X19] : s(X6,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X6),X17),s('type.quote.varmap'(X6),X18),s('type.list.list'('type.quote.index'),X19))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.list.list'('type.quote.index'),X6)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X6),X17))),s('type.quote.varmap'(X6),X18))),s('type.list.list'('type.quote.index'),X19)))),
  file('thm.ringNorm.interp_vl_def|split|1','ARITY_1_3_const.ringNorm.r_interp_vl')).
fof(f8508,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(superposition,[],[f85,f75])).
fof(f75,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.list.list'('type.quote.index'),X6)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.list.list'('type.quote.index'),X6)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_vl_def|split|1','thm.ringNorm.r_interp_vl_def')).
fof(f85,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X17,X18,X19] : s(X6,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X6),X17),s('type.quote.varmap'(X6),X18),s('type.list.list'('type.quote.index'),X19))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.list.list'('type.quote.index'),X6)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X6),X17))),s('type.quote.varmap'(X6),X18))),s('type.list.list'('type.quote.index'),X19)))),
  file('thm.ringNorm.interp_vl_def|split|1','ARITY_1_3_const.canonical.interp_vl')).
fof(f71,plain,(
  s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),sK4))) != s(sK0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),sK4)))))),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),sK4))) != s(sK0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) != s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ~! [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),X10))) = s(X6,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X9),s('type.list.list'('type.quote.index'),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),X10))) = s(X6,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X9),s('type.list.list'('type.quote.index'),X10)))))),
  file('thm.ringNorm.interp_vl_def|split|1','thm.ringNorm.interp_vl_def|split|1')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 230
% Passive clauses: 7481
% Generated clauses: 73935
% Final active clauses: 100
% Final passive clauses: 774
% Input formulas: 23
% Initial clauses: 24
% Discarded non-redundant clauses: 2765
% 
% Duplicate literals: 5482
% Trivial inequalities: 469
% Fw subsumption resolutions: 3151
% Fw demodulations: 4527
% Bw demodulations: 5002
% Condensations: 10151
% Global subsumptions: 151
% 
% Simple tautologies: 833
% Equational tautologies: 1472
% Forward subsumptions: 40029
% Bw demodulations to eq. taut.: 482
% 
% Unit resulting resolution: 3440
% Forward superposition: 19323
% Backward superposition: 13863
% Self superposition: 473
% Equality factoring: 297
% Equality resolution: 14
% 
% Split clauses: 162
% Split components: 325
% SAT solver clauses: 8127
% SAT solver unit clauses: 41
% SAT solver binary clauses: 559
% 
% Sat splits: 162
% Sat splitting refutations: 169
% 
% Memory used [KB]: 33005
% Time elapsed: 7.262 s
% ------------------------------
% ------------------------------
% Success in time 8.577 s
