Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quotient_option
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient_option
% SZS output start Proof for thm.quotient_option
fof(f14053,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14032,f262,f265,f2972,f2969,f178,f268])).
fof(f268,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f268_D])).
fof(f268_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f178,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient_option.OPTION_REL_def|split|2',aHL_FALSITY)).
fof(f2969,plain,(
  s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK0,bool)),sK2),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))))) | $spl32),
  inference(cnf_transformation,[],[f2969_D])).
fof(f2969_D,plain,(
  s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK0,bool)),sK2),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2972,plain,(
  s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK0,bool)),sK2),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))))) | $spl3),
  inference(subsumption_resolution,[],[f2957,f179])).
fof(f179,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X0),'const.option.NONE_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X6),'const.option.NONE_0')))),
  file('thm.quotient_option.OPTION_REL_def|split|2','thm.option.IS_SOME_DEF|split|1')).
fof(f2957,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X4),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK0,bool)),sK2),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1)))))) ) | $spl3),
  inference(superposition,[],[f171,f505])).
fof(f505,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X4),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(subsumption_resolution,[],[f499,f268])).
fof(f499,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.option.IS_SOME_1'(s('type.option.option'(X4),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f350,f179])).
fof(f350,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(resolution,[],[f243,f184])).
fof(f184,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.quotient_option.OPTION_REL_def|split|2','thm.bool.EQ_CLAUSES|split|0')).
fof(f243,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.quotient_option.OPTION_REL_def|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f171,plain,(
  s(bool,'F_0') != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK0,bool)),sK2),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1)))))),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  s(bool,'F_0') != s(bool,'const.option.OPTREL_3'(s(fun(sK0,fun(sK0,bool)),sK2),s('type.option.option'(sK0),'const.option.NONE_0'),s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f123])).
fof(f123,plain,(
  ? [X0,X1,X2] : s(bool,'F_0') != s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X0,bool)),X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ~! [X0,X1,X2] : s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X0,bool)),X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ~! [X0] : ! [X1,X2] : s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X0,bool)),X2),s('type.option.option'(X0),'const.option.NONE_0'),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X6,fun(X6,bool)),X8),s('type.option.option'(X6),'const.option.NONE_0'),s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(bool,'F_0') = s(bool,'const.option.OPTREL_3'(s(fun(X6,fun(X6,bool)),X8),s('type.option.option'(X6),'const.option.NONE_0'),s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7)))))),
  file('thm.quotient_option.OPTION_REL_def|split|2','thm.quotient_option.OPTION_REL_def|split|2')).
fof(f265,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f262,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f230])).
fof(f230,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X12] : (s(X6,X12) = s(X6,X12) <=> p(s(bool,'T_0')))),
  file('thm.quotient_option.OPTION_REL_def|split|2','thm.bool.REFL_CLAUSE')).
fof(f14032,plain,(
  $false | ($spl0 | $spl32)),
  inference(subsumption_resolution,[],[f14031,f238])).
fof(f238,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1] : ~s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X12] : ~s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12)))),
  file('thm.quotient_option.OPTION_REL_def|split|2','thm.option.NOT_NONE_SOME')).
fof(f14031,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK4('const.option.SOME_1'(s(sK0,sK1)),'const.option.NONE_0',sK2,sK0,sK0)))) | ($spl0 | $spl32)),
  inference(subsumption_resolution,[],[f14030,f238])).
fof(f14030,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK4('const.option.SOME_1'(s(sK0,sK1)),'const.option.NONE_0',sK2,sK0,sK0)))) | ($spl0 | $spl32)),
  inference(subsumption_resolution,[],[f13918,f265])).
fof(f13918,plain,(
  ~p(s(bool,'T_0')) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK1))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK4('const.option.SOME_1'(s(sK0,sK1)),'const.option.NONE_0',sK2,sK0,sK0)))) | $spl32),
  inference(superposition,[],[f255,f2969])).
fof(f255,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) | s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK4(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK4(X4,X3,X2,X1,X0)))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,sK5(X4,X3,X2,X1,X0)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK4(X4,X3,X2,X1,X0)))),s(X1,sK5(X4,X3,X2,X1,X0))))))) & (((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4)) & ! [X5,X6] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X7,X8] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X7))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X7))),s(X1,X8)))))) & (((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4)) & ! [X5,X6] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))))),
  inference(rectify,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X5,X6] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) & (((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4)) & ! [X5,X6] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) | ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X5,X6] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6))))))) & (((s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4)) & ! [X5,X6] : (s('type.option.option'(X0),X3) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))) | p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) <=> ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X5,X6] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.option.OPTREL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.option.option'(X0),X3),s('type.option.option'(X1),X4)))) <=> ((s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X3) & s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4)) | ? [X5,X6] : (s('type.option.option'(X0),X3) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) & s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6,X13] : ! [X8,X12,X7] : (p(s(bool,'const.option.OPTREL_3'(s(fun(X6,fun(X13,bool)),X8),s('type.option.option'(X6),X12),s('type.option.option'(X13),X7)))) <=> ((s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X12) & s('type.option.option'(X13),'const.option.NONE_0') = s('type.option.option'(X13),X7)) | ? [X19,X20] : (s('type.option.option'(X6),X12) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X19))) & s('type.option.option'(X13),X7) = s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X20))) & p(s(bool,happ(s(fun(X13,bool),happ(s(fun(X6,fun(X13,bool)),X8),s(X6,X19))),s(X13,X20)))))))),
  file('thm.quotient_option.OPTION_REL_def|split|2','thm.option.OPTREL_def')).
% SZS output end Proof for thm.quotient_option
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 140
% Passive clauses: 1684
% Generated clauses: 13619
% Final active clauses: 124
% Final passive clauses: 1335
% Input formulas: 49
% Initial clauses: 91
% 
% Duplicate literals: 603
% Trivial inequalities: 312
% Fw subsumption resolutions: 1471
% Fw demodulations: 1379
% Bw demodulations: 14
% 
% Simple tautologies: 376
% Equational tautologies: 46
% Forward subsumptions: 7399
% Fw demodulations to eq. taut.: 43
% 
% Binary resolution: 209
% Factoring: 1
% Forward superposition: 5049
% Backward superposition: 4208
% Self superposition: 115
% Equality factoring: 70
% Equality resolution: 31
% 
% Split clauses: 169
% Split components: 340
% SAT solver clauses: 464
% SAT solver unit clauses: 32
% SAT solver binary clauses: 183
% SAT solver learnt clauses: 1
% 
% Sat splits: 169
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 462
% TWLsolver calls for satisfiability: 119
% 
% Memory used [KB]: 7675
% Time elapsed: 0.308 s
% ------------------------------
% ------------------------------
% Success in time 0.315 s
