Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quote
% (84895)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 201
% Passive clauses: 4602
% Generated clauses: 21086
% Final active clauses: 192
% Final passive clauses: 4191
% Input formulas: 69
% Initial clauses: 122
% 
% Duplicate literals: 243
% Trivial inequalities: 265
% Fw subsumption resolutions: 1706
% Fw demodulations: 3177
% Bw demodulations: 8
% 
% Simple tautologies: 340
% Equational tautologies: 30
% Forward subsumptions: 10431
% 
% Binary resolution: 847
% Factoring: 2
% Forward superposition: 7135
% Backward superposition: 7280
% Self superposition: 157
% Equality factoring: 40
% Equality resolution: 24
% 
% Split clauses: 195
% Split components: 391
% SAT solver clauses: 451
% SAT solver unit clauses: 7
% SAT solver binary clauses: 101
% 
% Sat splits: 195
% 
% TWLsolver clauses: 451
% TWLsolver calls for satisfiability: 115
% 
% Memory used [KB]: 21620
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quote
% (84921)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 144
% Passive clauses: 1875
% Generated clauses: 11955
% Final active clauses: 135
% Final passive clauses: 1682
% Input formulas: 69
% Initial clauses: 122
% 
% Duplicate literals: 30
% Trivial inequalities: 58
% Fw subsumption resolutions: 287
% Fw demodulations: 2072
% Bw demodulations: 56
% Condensations: 41
% Global subsumptions: 5
% 
% Simple tautologies: 63
% Equational tautologies: 12
% Forward subsumptions: 5221
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 465
% Factoring: 1
% Forward superposition: 3225
% Backward superposition: 3583
% Self superposition: 110
% Equality factoring: 4
% Equality resolution: 11
% 
% Split clauses: 39
% Split components: 78
% SAT solver clauses: 2396
% SAT solver unit clauses: 143
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 10
% 
% Sat splits: 39
% 
% TWLsolver clauses: 2376
% TWLsolver calls for satisfiability: 462
% 
% Memory used [KB]: 24050
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quote
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quote
% SZS output start Proof for thm.quote
fof(f16170,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1270,f1280,f1279,f1768,f1789,f2030,f2039,f2051,f2065,f2063,f2388,f2402,f2400,f2657,f2669,f2667,f4727,f4756,f4754,f5247,f5248,f2040,f5251,f5253,f5255,f5259,f5261,f5263,f5265,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5324,f5340,f5342,f5344,f5346,f5348,f5350,f5352,f5358,f5431,f5436,f5438,f5440,f5442,f5445,f5450,f5488,f5489,f5490,f5491,f5492,f5493,f5494,f5495,f5496,f5497,f5498,f5499,f5500,f5501,f5502,f5503,f5504,f5505,f5506,f5507,f5508,f5509,f5510,f5511,f5512,f5513,f5514,f5515,f5516,f5517,f5518,f5519,f5520,f5521,f5522,f5523,f5524,f5525,f8833,f8894,f8897,f8896,f8898,f8899,f8901,f8900,f8909,f8913,f8911,f9953,f9980,f10050,f7226,f10053,f7190,f10056,f2058,f10060,f1901,f10061,f1900,f10062,f1899,f10063,f1898,f1791,f10064,f10098,f10100,f10102,f10105,f10108,f10110,f10113,f10117,f10120,f10122,f10124,f10126,f10129,f10419,f10423,f10427,f10431,f10436,f10442,f10447,f10451,f10456,f10461,f10465,f10469,f10472,f10478,f12144,f12170,f12239,f12277,f12315,f12313,f13695,f13697,f13698,f13700,f13701,f13702,f13704,f13705,f13706,f13708,f13707,f13709,f13710,f13711,f13712,f13715,f13716,f13717,f13718,f13719,f13720,f13722,f13723,f13724,f13725,f13727,f13728,f13729,f13730,f13731,f13713,f13732,f13734,f13735,f13736,f13737,f13738,f13739,f13696,f13742,f13703,f13743,f13745,f13746,f13699,f13748,f13750,f13749,f13765,f13767,f13766,f13768,f13769,f13775,f14798,f14800,f14801,f14802,f14803,f14804,f14805,f14806,f14808,f14810,f14812,f14814,f14816,f14818,f14823,f14824,f14825,f14826,f14827,f14828,f14829,f14830,f14831,f14832,f14833,f14834,f14835,f14836,f14837,f14838,f14839,f14840,f14841,f14842,f14843,f14844,f14845,f14846,f14847,f14848,f14849,f14850,f14851,f14852,f14853,f14854,f14889,f14892,f14893,f14894,f14895,f14896,f14897,f14898,f14899,f14900,f14901,f14902,f14903,f14904,f14905,f14906,f14907,f14908,f14909,f14910,f14911,f14912,f14913,f14914,f14915,f14916,f14917,f14918,f14919,f14927,f14939,f14944,f14946,f14948,f14950,f14952,f14954,f14955,f14956,f14957,f14958,f14959,f14960,f14962,f14963,f14966,f15016,f15033,f16107,f16109,f16110,f16111,f16112,f16113,f16114,f16115,f16117,f16119,f16121,f16123,f16125,f16127,f16132,f16133,f16134,f16135,f16136,f16137,f16138,f16139,f16140,f16141,f16142,f16143,f16144,f16145,f16146,f16147,f16148,f16149,f16150,f16151,f16152,f16153,f16154,f16155,f16156,f16157,f16158,f16159,f16160,f16161,f16162,f16163,f16164,f16165,f16166,f16167,f16168,f16169])).
fof(f16169,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f15007,f16131])).
fof(f16131,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl62),
  inference(subsumption_resolution,[],[f15417,f243])).
fof(f243,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quote.index_distinct|split|0',aHL_TRUTH)).
fof(f15417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2)) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X44,X41,X42] : (p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X44,fun(X44,bool)),'=_0'),s(X44,X41))),s(X44,X42)))) <=> s(X44,X41) = s(X44,X42))),
  file('thm.quote.index_distinct|split|0',arity_0_HL_EQ)).
fof(f15015,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl62),
  inference(unit_resulting_resolution,[],[f328,f13696,f398])).
fof(f398,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f257,f252])).
fof(f252,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quote.index_distinct|split|0',aHL_BOOL_CASES)).
fof(f257,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.quote.index_distinct|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f328,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X25,X23] : ! [X26,X28] : s(X23,X26) = s(X23,happ(s(fun(X25,X23),'_dst_x0_1'(s(X23,X26))),s(X25,X28)))),
  file('thm.quote.index_distinct|split|0','_dst_x0_ax')).
fof(f15007,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK0)))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f380,f13696,f398])).
fof(f380,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f246,f253])).
fof(f253,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.quote.index_distinct|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f246,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quote.index_distinct|split|0',aHL_FALSITY)).
fof(f16168,plain,(
  $false | ($spl2 | $spl62)),
  inference(subsumption_resolution,[],[f15008,f16131])).
fof(f15008,plain,(
  ( ! [X0,X1] : (s(bool,sK2(X0,'_i_0_0',bool)) != s(bool,happ(s(fun('type.quote.index',bool),X1),s('type.quote.index',sK0)))) ) | ($spl2 | $spl62)),
  inference(unit_resulting_resolution,[],[f1289,f13696,f398])).
fof(f1289,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK2(X0,'_i_0_0',bool))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1280,f253])).
fof(f16167,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f15009,f16131])).
fof(f15009,plain,(
  ( ! [X0] : (s(bool,sK5('F_0','_i_0_0',bool)) != s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK0)))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f1075,f13696,f398])).
fof(f1075,plain,(
  s(bool,'T_0') != s(bool,sK5('F_0','_i_0_0',bool))),
  inference(subsumption_resolution,[],[f1050,f246])).
fof(f1050,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK5('F_0','_i_0_0',bool))),
  inference(superposition,[],[f253,f1033])).
fof(f1033,plain,(
  s(bool,'F_0') = s(bool,sK5('F_0','_i_0_0',bool))),
  inference(unit_resulting_resolution,[],[f246,f1018])).
fof(f1018,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,sK5(X3,'_i_0_0',bool))) )),
  inference(superposition,[],[f293,f282])).
fof(f282,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X25] : ! [X27] : s(X25,X27) = s(X25,happ(s(fun(X25,X25),'_i_0_0'),s(X25,X27)))),
  file('thm.quote.index_distinct|split|0','_i0_ax')).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : (((s(X0,X2) = s(X0,sK5(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((? [X4] : (s(X0,X2) = s(X0,X4) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X9] : ! [X10,X12] : (! [X19] : (s(X9,X12) = s(X9,X19) => p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X19))))) <=> p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X12)))))),
  file('thm.quote.index_distinct|split|0','thm.bool.UNWIND_FORALL_THM2')).
fof(f16166,plain,(
  $false | ($spl2 | $spl62)),
  inference(subsumption_resolution,[],[f15010,f16131])).
fof(f15010,plain,(
  ( ! [X0,X1] : (s(bool,sK5(sK2(X0,'_i_0_0',bool),'_i_0_0',bool)) != s(bool,happ(s(fun('type.quote.index',bool),X1),s('type.quote.index',sK0)))) ) | ($spl2 | $spl62)),
  inference(unit_resulting_resolution,[],[f1538,f13696,f398])).
fof(f1538,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK5(sK2(X1,'_i_0_0',bool),'_i_0_0',bool))) ) | $spl2),
  inference(subsumption_resolution,[],[f1508,f246])).
fof(f1508,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK5(sK2(X1,'_i_0_0',bool),'_i_0_0',bool))) ) | $spl2),
  inference(superposition,[],[f253,f1295])).
fof(f1295,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK5(sK2(X0,'_i_0_0',bool),'_i_0_0',bool))) ) | $spl2),
  inference(forward_demodulation,[],[f1281,f1288])).
fof(f1288,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK2(X0,'_i_0_0',bool))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1280,f258])).
fof(f258,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f198])).
fof(f1281,plain,(
  ( ! [X0] : (s(bool,sK2(X0,'_i_0_0',bool)) = s(bool,sK5(sK2(X0,'_i_0_0',bool),'_i_0_0',bool))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1280,f1018])).
fof(f16165,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f15011,f16131])).
fof(f15011,plain,(
  ( ! [X0] : (s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool)) != s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK0)))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f1250,f13696,f398])).
fof(f1250,plain,(
  s(bool,'T_0') != s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool))),
  inference(subsumption_resolution,[],[f1224,f246])).
fof(f1224,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool))),
  inference(superposition,[],[f253,f1085])).
fof(f1085,plain,(
  s(bool,'F_0') = s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool))),
  inference(forward_demodulation,[],[f1084,f1033])).
fof(f1084,plain,(
  s(bool,sK5('F_0','_i_0_0',bool)) = s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool))),
  inference(subsumption_resolution,[],[f1060,f246])).
fof(f1060,plain,(
  p(s(bool,'F_0')) | s(bool,sK5('F_0','_i_0_0',bool)) = s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool))),
  inference(superposition,[],[f1018,f1033])).
fof(f16164,plain,(
  $false | ($spl2 | $spl23 | $spl62)),
  inference(subsumption_resolution,[],[f15012,f16131])).
fof(f15012,plain,(
  ( ! [X0] : (s(bool,sK5(sK5(sK6('_i_0_0',bool),'_i_0_0',bool),'_i_0_0',bool)) != s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK0)))) ) | ($spl2 | $spl23 | $spl62)),
  inference(unit_resulting_resolution,[],[f6042,f13696,f398])).
fof(f6042,plain,(
  s(bool,'T_0') != s(bool,sK5(sK5(sK6('_i_0_0',bool),'_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(subsumption_resolution,[],[f5979,f246])).
fof(f5979,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK5(sK5(sK6('_i_0_0',bool),'_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(superposition,[],[f253,f2886])).
fof(f2886,plain,(
  s(bool,'F_0') = s(bool,sK5(sK5(sK6('_i_0_0',bool),'_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(forward_demodulation,[],[f2885,f2690])).
fof(f2690,plain,(
  s(bool,'F_0') = s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(backward_demodulation,[],[f2689,f2684])).
fof(f2684,plain,(
  s(bool,sK6('_i_0_0',bool)) = s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f2669,f1018])).
fof(f2689,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(forward_demodulation,[],[f2683,f1288])).
fof(f2683,plain,(
  ( ! [X0] : (s(bool,sK6('_i_0_0',bool)) = s(bool,sK2(X0,'_i_0_0',bool))) ) | ($spl2 | $spl23)),
  inference(unit_resulting_resolution,[],[f1289,f2669,f454])).
fof(f454,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f448,f246])).
fof(f448,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f302,f252])).
fof(f302,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.quote.index_distinct|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f2885,plain,(
  s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool)) = s(bool,sK5(sK5(sK6('_i_0_0',bool),'_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(subsumption_resolution,[],[f2847,f246])).
fof(f2847,plain,(
  p(s(bool,'F_0')) | s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool)) = s(bool,sK5(sK5(sK6('_i_0_0',bool),'_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(superposition,[],[f1018,f2690])).
fof(f16163,plain,(
  $false | ($spl2 | $spl23 | $spl62)),
  inference(subsumption_resolution,[],[f15013,f16131])).
fof(f15013,plain,(
  ( ! [X0] : (s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool)) != s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK0)))) ) | ($spl2 | $spl23 | $spl62)),
  inference(unit_resulting_resolution,[],[f2876,f13696,f398])).
fof(f2876,plain,(
  s(bool,'T_0') != s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(subsumption_resolution,[],[f2828,f246])).
fof(f2828,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl23)),
  inference(superposition,[],[f253,f2690])).
fof(f16162,plain,(
  $false | ($spl23 | $spl62)),
  inference(subsumption_resolution,[],[f15014,f16131])).
fof(f15014,plain,(
  ( ! [X0] : (s(bool,sK6('_i_0_0',bool)) != s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK0)))) ) | ($spl23 | $spl62)),
  inference(unit_resulting_resolution,[],[f2670,f13696,f398])).
fof(f2670,plain,(
  s(bool,'T_0') != s(bool,sK6('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f2669,f253])).
fof(f16161,plain,(
  $false | ($spl62 | $spl84)),
  inference(subsumption_resolution,[],[f15980,f243])).
fof(f15980,plain,(
  ~p(s(bool,'T_0')) | ($spl62 | $spl84)),
  inference(backward_demodulation,[],[f15015,f13716])).
fof(f16160,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15973,f243])).
fof(f15973,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5048])).
fof(f5048,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))) | $spl27),
  inference(forward_demodulation,[],[f5047,f338])).
fof(f338,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X23,X25,X24] : ! [X26,X27,X28] : s(X23,happ(s(fun(X24,X23),happ(s(fun(X25,fun(X24,X23)),X26),s(X25,X28))),s(X24,X27))) = s(X23,happ(s(fun(X25,X23),'_dst_x1x0_2'(s(fun(X25,fun(X24,X23)),X26),s(X24,X27))),s(X25,X28)))),
  file('thm.quote.index_distinct|split|0','_dst_x1x0_ax')).
fof(f5047,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | $spl27),
  inference(unit_resulting_resolution,[],[f4756,f334])).
fof(f16159,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15972,f243])).
fof(f15972,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5162])).
fof(f5162,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))) | $spl27),
  inference(forward_demodulation,[],[f5161,f338])).
fof(f5161,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5046,f334])).
fof(f5046,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))) | $spl27),
  inference(unit_resulting_resolution,[],[f4756,f306])).
fof(f306,plain,(
  ( ! [X0,X1] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) | s('type.quote.index',X0) = s('type.quote.index',X1)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) | s('type.quote.index',X0) = s('type.quote.index',X1)) & (s('type.quote.index',X0) != s('type.quote.index',X1) | s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1)))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) <=> s('type.quote.index',X0) = s('type.quote.index',X1))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X6] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X6))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X7))) <=> s('type.quote.index',X6) = s('type.quote.index',X7))),
  file('thm.quote.index_distinct|split|0','thm.quote.index_|49||49||split|0')).
fof(f16158,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15971,f243])).
fof(f15971,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5738])).
fof(f5738,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))))) | $spl27),
  inference(forward_demodulation,[],[f5737,f338])).
fof(f5737,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5160,f334])).
fof(f5160,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5046,f306])).
fof(f16157,plain,(
  $false | ($spl19 | $spl62)),
  inference(subsumption_resolution,[],[f15970,f243])).
fof(f15970,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2975])).
fof(f2975,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))))) | $spl19),
  inference(forward_demodulation,[],[f2974,f338])).
fof(f2974,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f2636,f334])).
fof(f2636,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f2402,f306])).
fof(f16156,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15969,f243])).
fof(f15969,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5615])).
fof(f5615,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))))) | $spl27),
  inference(forward_demodulation,[],[f5614,f338])).
fof(f5614,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5055,f334])).
fof(f5055,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5045,f306])).
fof(f5045,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))) | $spl27),
  inference(unit_resulting_resolution,[],[f4756,f308])).
fof(f308,plain,(
  ( ! [X0,X1] : (s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))) | s('type.quote.index',X0) = s('type.quote.index',X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))) | s('type.quote.index',X0) = s('type.quote.index',X1)) & (s('type.quote.index',X0) != s('type.quote.index',X1) | s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : (s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))) <=> s('type.quote.index',X0) = s('type.quote.index',X1))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X6] : (s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X6))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X7))) <=> s('type.quote.index',X6) = s('type.quote.index',X7))),
  file('thm.quote.index_distinct|split|0','thm.quote.index_|49||49||split|1')).
fof(f16155,plain,(
  $false | ($spl15 | $spl62)),
  inference(subsumption_resolution,[],[f15968,f243])).
fof(f15968,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2383])).
fof(f2383,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))))) | $spl15),
  inference(forward_demodulation,[],[f2382,f338])).
fof(f2382,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2219,f334])).
fof(f2219,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2065,f306])).
fof(f16154,plain,(
  $false | ($spl7 | $spl62)),
  inference(subsumption_resolution,[],[f15967,f243])).
fof(f15967,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl62)),
  inference(backward_demodulation,[],[f15015,f10562])).
fof(f10562,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl7),
  inference(forward_demodulation,[],[f10561,f338])).
fof(f10561,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1791,f334])).
fof(f16153,plain,(
  $false | ($spl19 | $spl62)),
  inference(subsumption_resolution,[],[f15966,f243])).
fof(f15966,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2638])).
fof(f2638,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | $spl19),
  inference(forward_demodulation,[],[f2637,f338])).
fof(f2637,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | $spl19),
  inference(unit_resulting_resolution,[],[f2402,f334])).
fof(f16152,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15965,f243])).
fof(f15965,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5624])).
fof(f5624,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))))) | $spl27),
  inference(forward_demodulation,[],[f5623,f338])).
fof(f5623,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5159,f334])).
fof(f5159,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5046,f308])).
fof(f16151,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15964,f243])).
fof(f15964,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5057])).
fof(f5057,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))) | $spl27),
  inference(forward_demodulation,[],[f5056,f338])).
fof(f5056,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5045,f334])).
fof(f16150,plain,(
  $false | ($spl15 | $spl62)),
  inference(subsumption_resolution,[],[f15963,f243])).
fof(f15963,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2312])).
fof(f2312,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | $spl15),
  inference(forward_demodulation,[],[f2246,f338])).
fof(f2246,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2065,f334])).
fof(f16149,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15962,f243])).
fof(f15962,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5171])).
fof(f5171,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))))) | $spl27),
  inference(forward_demodulation,[],[f5170,f338])).
fof(f5170,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5054,f334])).
fof(f5054,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5045,f308])).
fof(f16148,plain,(
  $false | ($spl19 | $spl62)),
  inference(subsumption_resolution,[],[f15961,f243])).
fof(f15961,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2647])).
fof(f2647,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))))) | $spl19),
  inference(forward_demodulation,[],[f2646,f338])).
fof(f2646,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f2635,f334])).
fof(f2635,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f2402,f308])).
fof(f16147,plain,(
  $false | ($spl15 | $spl62)),
  inference(subsumption_resolution,[],[f15960,f243])).
fof(f15960,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2310])).
fof(f2310,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))))) | $spl15),
  inference(forward_demodulation,[],[f2248,f338])).
fof(f2248,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2218,f334])).
fof(f2218,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2065,f308])).
fof(f16146,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15956,f243])).
fof(f15956,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5044])).
fof(f5044,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))) | $spl27),
  inference(unit_resulting_resolution,[],[f4756,f334])).
fof(f16145,plain,(
  $false | ($spl7 | $spl62)),
  inference(subsumption_resolution,[],[f15955,f243])).
fof(f15955,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl62)),
  inference(backward_demodulation,[],[f15015,f10558])).
fof(f10558,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1791,f334])).
fof(f16144,plain,(
  $false | ($spl19 | $spl62)),
  inference(subsumption_resolution,[],[f15954,f243])).
fof(f15954,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2634])).
fof(f2634,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f2402,f334])).
fof(f16143,plain,(
  $false | ($spl15 | $spl62)),
  inference(subsumption_resolution,[],[f15953,f243])).
fof(f15953,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2232])).
fof(f2232,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2065,f334])).
fof(f16142,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15952,f243])).
fof(f15952,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5166])).
fof(f5166,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5054,f334])).
fof(f16141,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15951,f243])).
fof(f15951,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5052])).
fof(f5052,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5045,f334])).
fof(f16140,plain,(
  $false | ($spl19 | $spl62)),
  inference(subsumption_resolution,[],[f15950,f243])).
fof(f15950,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2642])).
fof(f2642,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f2635,f334])).
fof(f16139,plain,(
  $false | ($spl15 | $spl62)),
  inference(subsumption_resolution,[],[f15949,f243])).
fof(f15949,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2234])).
fof(f2234,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2218,f334])).
fof(f16138,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15948,f243])).
fof(f15948,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5619])).
fof(f5619,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5159,f334])).
fof(f16137,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15944,f243])).
fof(f15944,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5157])).
fof(f5157,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5046,f334])).
fof(f16136,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15943,f243])).
fof(f15943,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5610])).
fof(f5610,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5055,f334])).
fof(f16135,plain,(
  $false | ($spl19 | $spl62)),
  inference(subsumption_resolution,[],[f15942,f243])).
fof(f15942,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2970])).
fof(f2970,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f2636,f334])).
fof(f16134,plain,(
  $false | ($spl15 | $spl62)),
  inference(subsumption_resolution,[],[f15941,f243])).
fof(f15941,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2378])).
fof(f2378,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2219,f334])).
fof(f16133,plain,(
  $false | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f15940,f243])).
fof(f15940,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62)),
  inference(backward_demodulation,[],[f15015,f5733])).
fof(f5733,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))))))) | $spl27),
  inference(unit_resulting_resolution,[],[f5160,f334])).
fof(f16132,plain,(
  $false | ($spl62 | $spl80)),
  inference(subsumption_resolution,[],[f15773,f243])).
fof(f15773,plain,(
  ~p(s(bool,'T_0')) | ($spl62 | $spl80)),
  inference(backward_demodulation,[],[f15015,f13712])).
fof(f16127,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f16126,f15032])).
fof(f15032,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f328,f13696,f3694])).
fof(f3694,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X6)) )),
  inference(forward_demodulation,[],[f3656,f282])).
fof(f3656,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,X6) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X6))))) )),
  inference(superposition,[],[f295,f282])).
fof(f295,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f16126,plain,(
  ( ! [X414,X420,X416,X418,X415,X417,X419] : (~p(s(bool,happ(s(fun(X414,bool),happ(s(fun(X415,fun(X414,bool)),X416),s(X415,X417))),s(X414,happ(s(fun(X418,X414),X419),s(X418,X420))))))) ) | $spl62),
  inference(subsumption_resolution,[],[f15214,f15015])).
fof(f15214,plain,(
  ( ! [X414,X420,X416,X418,X415,X417,X419] : (s(bool,'T_0') != s(bool,happ(s(fun(X418,bool),'_dst_x0x0x1_3'(s(fun(X415,fun(X414,bool)),X416),s(X415,X417),s(fun(X418,X414),X419))),s(X418,X420))) | ~p(s(bool,happ(s(fun(X414,bool),happ(s(fun(X415,fun(X414,bool)),X416),s(X415,X417))),s(X414,happ(s(fun(X418,X414),X419),s(X418,X420))))))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f7616])).
fof(f7616,plain,(
  ( ! [X414,X420,X416,X418,X415,X417,X419] : (s(bool,'F_0') != s(bool,happ(s(fun(X418,bool),'_dst_x0x0x1_3'(s(fun(X415,fun(X414,bool)),X416),s(X415,X417),s(fun(X418,X414),X419))),s(X418,X420))) | ~p(s(bool,happ(s(fun(X414,bool),happ(s(fun(X415,fun(X414,bool)),X416),s(X415,X417))),s(X414,happ(s(fun(X418,X414),X419),s(X418,X420))))))) )),
  inference(superposition,[],[f257,f356])).
fof(f356,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X23,X24,X38,X25] : ! [X26,X27,X40,X28] : s(X23,happ(s(fun(X25,X23),'_dst_x0x0x1_3'(s(fun(X24,fun(X38,X23)),X26),s(X24,X27),s(fun(X25,X38),X40))),s(X25,X28))) = s(X23,happ(s(fun(X38,X23),happ(s(fun(X24,fun(X38,X23)),X26),s(X24,X27))),s(X38,happ(s(fun(X25,X38),X40),s(X25,X28)))))),
  file('thm.quote.index_distinct|split|0','_dst_x0x0x1_ax')).
fof(f16125,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f16124,f15032])).
fof(f16124,plain,(
  ( ! [X412,X414,X410,X413,X415,X409,X411] : (~p(s(bool,happ(s(fun(X409,bool),'_dst_x0x0x1_3'(s(fun(X410,fun(X411,bool)),X412),s(X410,X413),s(fun(X409,X411),X414))),s(X409,X415))))) ) | $spl62),
  inference(subsumption_resolution,[],[f15211,f15015])).
fof(f15211,plain,(
  ( ! [X412,X414,X410,X413,X415,X409,X411] : (s(bool,'T_0') != s(bool,happ(s(fun(X411,bool),happ(s(fun(X410,fun(X411,bool)),X412),s(X410,X413))),s(X411,happ(s(fun(X409,X411),X414),s(X409,X415))))) | ~p(s(bool,happ(s(fun(X409,bool),'_dst_x0x0x1_3'(s(fun(X410,fun(X411,bool)),X412),s(X410,X413),s(fun(X409,X411),X414))),s(X409,X415))))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f7482])).
fof(f7482,plain,(
  ( ! [X412,X414,X410,X413,X415,X409,X411] : (s(bool,'F_0') != s(bool,happ(s(fun(X411,bool),happ(s(fun(X410,fun(X411,bool)),X412),s(X410,X413))),s(X411,happ(s(fun(X409,X411),X414),s(X409,X415))))) | ~p(s(bool,happ(s(fun(X409,bool),'_dst_x0x0x1_3'(s(fun(X410,fun(X411,bool)),X412),s(X410,X413),s(fun(X409,X411),X414))),s(X409,X415))))) )),
  inference(superposition,[],[f257,f356])).
fof(f16123,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f16122,f15032])).
fof(f16122,plain,(
  ( ! [X292,X294,X293,X295,X291] : (~p(s(bool,happ(s(fun(X291,bool),happ(s(fun(X292,fun(X291,bool)),X293),s(X292,X294))),s(X291,X295))))) ) | $spl62),
  inference(subsumption_resolution,[],[f15202,f15015])).
fof(f15202,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s(bool,'T_0') != s(bool,happ(s(fun(X292,bool),'_dst_x1x0_2'(s(fun(X292,fun(X291,bool)),X293),s(X291,X295))),s(X292,X294))) | ~p(s(bool,happ(s(fun(X291,bool),happ(s(fun(X292,fun(X291,bool)),X293),s(X292,X294))),s(X291,X295))))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f6382])).
fof(f6382,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s(bool,'F_0') != s(bool,happ(s(fun(X292,bool),'_dst_x1x0_2'(s(fun(X292,fun(X291,bool)),X293),s(X291,X295))),s(X292,X294))) | ~p(s(bool,happ(s(fun(X291,bool),happ(s(fun(X292,fun(X291,bool)),X293),s(X292,X294))),s(X291,X295))))) )),
  inference(superposition,[],[f257,f338])).
fof(f16121,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f16120,f15032])).
fof(f16120,plain,(
  ( ! [X280,X277,X279,X281,X278] : (~p(s(bool,happ(s(fun(X277,bool),'_dst_x1x0_2'(s(fun(X277,fun(X278,bool)),X279),s(X278,X280))),s(X277,X281))))) ) | $spl62),
  inference(subsumption_resolution,[],[f15199,f15015])).
fof(f15199,plain,(
  ( ! [X280,X277,X279,X281,X278] : (s(bool,'T_0') != s(bool,happ(s(fun(X278,bool),happ(s(fun(X277,fun(X278,bool)),X279),s(X277,X281))),s(X278,X280))) | ~p(s(bool,happ(s(fun(X277,bool),'_dst_x1x0_2'(s(fun(X277,fun(X278,bool)),X279),s(X278,X280))),s(X277,X281))))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f6263])).
fof(f6263,plain,(
  ( ! [X280,X277,X279,X281,X278] : (s(bool,'F_0') != s(bool,happ(s(fun(X278,bool),happ(s(fun(X277,fun(X278,bool)),X279),s(X277,X281))),s(X278,X280))) | ~p(s(bool,happ(s(fun(X277,bool),'_dst_x1x0_2'(s(fun(X277,fun(X278,bool)),X279),s(X278,X280))),s(X277,X281))))) )),
  inference(superposition,[],[f257,f338])).
fof(f16119,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f16118,f15032])).
fof(f16118,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,happ(s(fun(X210,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X210,X211),X213))),s(X210,X214))))) ) | $spl62),
  inference(subsumption_resolution,[],[f15189,f15015])).
fof(f15189,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,'T_0') != s(bool,happ(s(fun(X211,bool),X212),s(X211,happ(s(fun(X210,X211),X213),s(X210,X214))))) | ~p(s(bool,happ(s(fun(X210,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X210,X211),X213))),s(X210,X214))))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f4975])).
fof(f4975,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,'F_0') != s(bool,happ(s(fun(X211,bool),X212),s(X211,happ(s(fun(X210,X211),X213),s(X210,X214))))) | ~p(s(bool,happ(s(fun(X210,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X210,X211),X213))),s(X210,X214))))) )),
  inference(superposition,[],[f257,f337])).
fof(f337,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X23,X24,X25] : ! [X26,X27,X28] : s(X23,happ(s(fun(X24,X23),X26),s(X24,happ(s(fun(X25,X24),X27),s(X25,X28))))) = s(X23,happ(s(fun(X25,X23),'_dst_x0x1_2'(s(fun(X24,X23),X26),s(fun(X25,X24),X27))),s(X25,X28)))),
  file('thm.quote.index_distinct|split|0','_dst_x0x1_ax')).
fof(f16117,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f16116,f15032])).
fof(f16116,plain,(
  ( ! [X246,X244,X248,X247,X245] : (~p(s(bool,happ(s(fun(X244,bool),X245),s(X244,happ(s(fun(X246,X244),X247),s(X246,X248))))))) ) | $spl62),
  inference(subsumption_resolution,[],[f15186,f15015])).
fof(f15186,plain,(
  ( ! [X246,X244,X248,X247,X245] : (s(bool,'T_0') != s(bool,happ(s(fun(X246,bool),'_dst_x0x1_2'(s(fun(X244,bool),X245),s(fun(X246,X244),X247))),s(X246,X248))) | ~p(s(bool,happ(s(fun(X244,bool),X245),s(X244,happ(s(fun(X246,X244),X247),s(X246,X248))))))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f4878])).
fof(f4878,plain,(
  ( ! [X246,X244,X248,X247,X245] : (s(bool,'F_0') != s(bool,happ(s(fun(X246,bool),'_dst_x0x1_2'(s(fun(X244,bool),X245),s(fun(X246,X244),X247))),s(X246,X248))) | ~p(s(bool,happ(s(fun(X244,bool),X245),s(X244,happ(s(fun(X246,X244),X247),s(X246,X248))))))) )),
  inference(superposition,[],[f257,f337])).
fof(f16115,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f15176,f15032])).
fof(f15176,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl62),
  inference(backward_demodulation,[],[f15015,f3153])).
fof(f3153,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f243,f246,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X41))),s(bool,X42)))) <=> (p(s(bool,X41)) => p(s(bool,X42))))),
  file('thm.quote.index_distinct|split|0','==>_HL_0')).
fof(f16114,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f15173,f15032])).
fof(f15173,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl62),
  inference(backward_demodulation,[],[f15015,f2976])).
fof(f2976,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f246,f246,f317])).
fof(f317,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X41))),s(bool,X42)))) <=> (p(s(bool,X41)) | p(s(bool,X42))))),
  file('thm.quote.index_distinct|split|0','|bslash||slash|_HL_0')).
fof(f16113,plain,(
  $false | ($spl2 | $spl62)),
  inference(subsumption_resolution,[],[f15159,f15032])).
fof(f15159,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl2 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2326])).
fof(f2326,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))) | $spl2),
  inference(forward_demodulation,[],[f2230,f1288])).
fof(f2230,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,sK2(X0,'_i_0_0',bool)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1289,f334])).
fof(f16112,plain,(
  $false | ($spl2 | $spl62)),
  inference(subsumption_resolution,[],[f15158,f15032])).
fof(f15158,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl2 | $spl62)),
  inference(backward_demodulation,[],[f15015,f2316])).
fof(f2316,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl2),
  inference(forward_demodulation,[],[f2315,f338])).
fof(f2315,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))) | $spl2),
  inference(forward_demodulation,[],[f2314,f1288])).
fof(f2314,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,sK2(X0,'_i_0_0',bool)))))) ) | $spl2),
  inference(forward_demodulation,[],[f2244,f338])).
fof(f2244,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,sK2(X0,'_i_0_0',bool)))),s(bool,'T_0'))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1289,f334])).
fof(f16111,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f15141,f15032])).
fof(f15141,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,'T_0'))))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f1433])).
fof(f1433,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f246,f315])).
fof(f315,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) & p(s(bool,X1)))) & ((~p(s(bool,X0)) | ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) & p(s(bool,X1))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X41))),s(bool,X42)))) <=> (p(s(bool,X41)) & p(s(bool,X42))))),
  file('thm.quote.index_distinct|split|0','|slash||bslash|_HL_0')).
fof(f16110,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f15140,f15032])).
fof(f15140,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,X0))))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f1337])).
fof(f1337,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f246,f314])).
fof(f314,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f16109,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f16108,f15032])).
fof(f16108,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl62),
  inference(subsumption_resolution,[],[f15111,f15015])).
fof(f15111,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl62),
  inference(backward_demodulation,[],[f15015,f257])).
fof(f16107,plain,(
  $false | $spl62),
  inference(subsumption_resolution,[],[f15110,f243])).
fof(f15110,plain,(
  ~p(s(bool,'T_0')) | $spl62),
  inference(backward_demodulation,[],[f15015,f246])).
fof(f15033,plain,(
  $false | $spl62),
  inference(unit_resulting_resolution,[],[f246,f328,f13696,f3694])).
fof(f15016,plain,(
  $false | $spl62),
  inference(unit_resulting_resolution,[],[f380,f328,f13696,f398])).
fof(f14966,plain,(
  $false | ($spl68 | $spl72)),
  inference(subsumption_resolution,[],[f14965,f243])).
fof(f14965,plain,(
  ~p(s(bool,'T_0')) | ($spl68 | $spl72)),
  inference(forward_demodulation,[],[f13703,f13700])).
fof(f14963,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f246,f13776])).
fof(f13776,plain,(
  ( ! [X384] : (p(s(bool,X384))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f13746,f13707])).
fof(f14962,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f14961,f13776])).
fof(f14961,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f257,f14942])).
fof(f14942,plain,(
  ( ! [X225,X223,X224] : (s(X223,X224) = s(X223,X225)) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f6366,f13776])).
fof(f6366,plain,(
  ( ! [X225,X223,X224] : (~p(s(bool,happ(s(fun(X223,bool),'_dst_x1x0_2'(s(fun(X223,fun(X223,bool)),'=_0'),s(X223,X225))),s(X223,X224)))) | s(X223,X224) = s(X223,X225)) )),
  inference(superposition,[],[f334,f338])).
fof(f14960,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f1337,f13776])).
fof(f14959,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f1433,f13776])).
fof(f14958,plain,(
  $false | ($spl2 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2316,f13776])).
fof(f14957,plain,(
  $false | ($spl2 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2326,f13776])).
fof(f14956,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2976,f13776])).
fof(f14955,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f3153,f13776])).
fof(f14954,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f14953,f13776])).
fof(f14953,plain,(
  ( ! [X246,X244,X248,X247,X245] : (~p(s(bool,happ(s(fun(X244,bool),X245),s(X244,happ(s(fun(X246,X244),X247),s(X246,X248))))))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f4878,f14941])).
fof(f14941,plain,(
  ( ! [X202,X200,X199,X203,X201] : (s(X199,X200) = s(X199,happ(s(fun(X201,X199),X202),s(X201,X203)))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f7720,f13776])).
fof(f7720,plain,(
  ( ! [X202,X200,X199,X203,X201] : (~p(s(bool,happ(s(fun(X201,bool),'_dst_x0x0x1_3'(s(fun(X199,fun(X199,bool)),'=_0'),s(X199,X200),s(fun(X201,X199),X202))),s(X201,X203)))) | s(X199,X200) = s(X199,happ(s(fun(X201,X199),X202),s(X201,X203)))) )),
  inference(backward_demodulation,[],[f7435,f4872])).
fof(f4872,plain,(
  ( ! [X202,X200,X199,X203,X201] : (~p(s(bool,happ(s(fun(X201,bool),'_dst_x0x1_2'(s(fun(X199,bool),happ(s(fun(X199,fun(X199,bool)),'=_0'),s(X199,X200))),s(fun(X201,X199),X202))),s(X201,X203)))) | s(X199,X200) = s(X199,happ(s(fun(X201,X199),X202),s(X201,X203)))) )),
  inference(superposition,[],[f334,f337])).
fof(f7435,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(X16,happ(s(fun(X21,X16),'_dst_x0x0x1_3'(s(fun(X18,fun(X17,X16)),X19),s(X18,X20),s(fun(X21,X17),X22))),s(X21,X23))) = s(X16,happ(s(fun(X21,X16),'_dst_x0x1_2'(s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),X19),s(X18,X20))),s(fun(X21,X17),X22))),s(X21,X23)))) )),
  inference(superposition,[],[f356,f337])).
fof(f14952,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f14951,f13776])).
fof(f14951,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,happ(s(fun(X210,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X210,X211),X213))),s(X210,X214))))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f4975,f14941])).
fof(f14950,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f14949,f13776])).
fof(f14949,plain,(
  ( ! [X280,X277,X279,X281,X278] : (~p(s(bool,happ(s(fun(X277,bool),'_dst_x1x0_2'(s(fun(X277,fun(X278,bool)),X279),s(X278,X280))),s(X277,X281))))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f6263,f14941])).
fof(f14948,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f14947,f13776])).
fof(f14947,plain,(
  ( ! [X292,X294,X293,X295,X291] : (~p(s(bool,happ(s(fun(X291,bool),happ(s(fun(X292,fun(X291,bool)),X293),s(X292,X294))),s(X291,X295))))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f6382,f14941])).
fof(f14946,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f14945,f13776])).
fof(f14945,plain,(
  ( ! [X412,X414,X410,X413,X415,X409,X411] : (~p(s(bool,happ(s(fun(X409,bool),'_dst_x0x0x1_3'(s(fun(X410,fun(X411,bool)),X412),s(X410,X413),s(fun(X409,X411),X414))),s(X409,X415))))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f7482,f14941])).
fof(f14944,plain,(
  $false | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f14943,f13776])).
fof(f14943,plain,(
  ( ! [X414,X420,X416,X418,X415,X417,X419] : (~p(s(bool,happ(s(fun(X414,bool),happ(s(fun(X415,fun(X414,bool)),X416),s(X415,X417))),s(X414,happ(s(fun(X418,X414),X419),s(X418,X420))))))) ) | ($spl76 | $spl108)),
  inference(subsumption_resolution,[],[f7616,f14941])).
fof(f14939,plain,(
  $false | ($spl76 | $spl80 | $spl108)),
  inference(subsumption_resolution,[],[f13712,f13776])).
fof(f14927,plain,(
  $false | ($spl72 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f13703,f13776])).
fof(f14919,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5733,f13776])).
fof(f14918,plain,(
  $false | ($spl15 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2378,f13776])).
fof(f14917,plain,(
  $false | ($spl19 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2970,f13776])).
fof(f14916,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5610,f13776])).
fof(f14915,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5157,f13776])).
fof(f14914,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5619,f13776])).
fof(f14913,plain,(
  $false | ($spl15 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2234,f13776])).
fof(f14912,plain,(
  $false | ($spl19 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2642,f13776])).
fof(f14911,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5166,f13776])).
fof(f14910,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5052,f13776])).
fof(f14909,plain,(
  $false | ($spl15 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2232,f13776])).
fof(f14908,plain,(
  $false | ($spl19 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2634,f13776])).
fof(f14907,plain,(
  $false | ($spl7 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f10558,f13776])).
fof(f14906,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5044,f13776])).
fof(f14905,plain,(
  $false | ($spl15 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2310,f13776])).
fof(f14904,plain,(
  $false | ($spl19 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2647,f13776])).
fof(f14903,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5171,f13776])).
fof(f14902,plain,(
  $false | ($spl15 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2312,f13776])).
fof(f14901,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5057,f13776])).
fof(f14900,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5624,f13776])).
fof(f14899,plain,(
  $false | ($spl19 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2638,f13776])).
fof(f14898,plain,(
  $false | ($spl7 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f10562,f13776])).
fof(f14897,plain,(
  $false | ($spl15 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2383,f13776])).
fof(f14896,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5615,f13776])).
fof(f14895,plain,(
  $false | ($spl19 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f2975,f13776])).
fof(f14894,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5738,f13776])).
fof(f14893,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5162,f13776])).
fof(f14892,plain,(
  $false | ($spl27 | $spl76 | $spl108)),
  inference(subsumption_resolution,[],[f5048,f13776])).
fof(f14889,plain,(
  $false | ($spl76 | $spl84 | $spl108)),
  inference(subsumption_resolution,[],[f13716,f13776])).
fof(f14854,plain,(
  $false | ($spl64 | $spl68)),
  inference(subsumption_resolution,[],[f13700,f13697])).
fof(f14853,plain,(
  $false | ($spl84 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14667,f243])).
fof(f14667,plain,(
  ~p(s(bool,'T_0')) | ($spl84 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f13716])).
fof(f13777,plain,(
  ( ! [X361] : (s(bool,'T_0') = s(bool,X361)) ) | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13737,f13739])).
fof(f14852,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14654,f243])).
fof(f14654,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5048])).
fof(f14851,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14653,f243])).
fof(f14653,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5162])).
fof(f14850,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14652,f243])).
fof(f14652,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5738])).
fof(f14849,plain,(
  $false | ($spl19 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14651,f243])).
fof(f14651,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2975])).
fof(f14848,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14650,f243])).
fof(f14650,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5615])).
fof(f14847,plain,(
  $false | ($spl15 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14649,f243])).
fof(f14649,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2383])).
fof(f14846,plain,(
  $false | ($spl7 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14648,f243])).
fof(f14648,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f10562])).
fof(f14845,plain,(
  $false | ($spl19 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14647,f243])).
fof(f14647,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2638])).
fof(f14844,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14646,f243])).
fof(f14646,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5624])).
fof(f14843,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14645,f243])).
fof(f14645,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5057])).
fof(f14842,plain,(
  $false | ($spl15 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14644,f243])).
fof(f14644,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2312])).
fof(f14841,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14643,f243])).
fof(f14643,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5171])).
fof(f14840,plain,(
  $false | ($spl19 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14642,f243])).
fof(f14642,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2647])).
fof(f14839,plain,(
  $false | ($spl15 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14641,f243])).
fof(f14641,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2310])).
fof(f14838,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14637,f243])).
fof(f14637,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5044])).
fof(f14837,plain,(
  $false | ($spl7 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14636,f243])).
fof(f14636,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f10558])).
fof(f14836,plain,(
  $false | ($spl19 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14635,f243])).
fof(f14635,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2634])).
fof(f14835,plain,(
  $false | ($spl15 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14634,f243])).
fof(f14634,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2232])).
fof(f14834,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14633,f243])).
fof(f14633,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5052])).
fof(f14833,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14632,f243])).
fof(f14632,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5166])).
fof(f14832,plain,(
  $false | ($spl19 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14631,f243])).
fof(f14631,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2642])).
fof(f14831,plain,(
  $false | ($spl15 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14630,f243])).
fof(f14630,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2234])).
fof(f14830,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14629,f243])).
fof(f14629,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5619])).
fof(f14829,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14623,f243])).
fof(f14623,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5157])).
fof(f14828,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14622,f243])).
fof(f14622,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5610])).
fof(f14827,plain,(
  $false | ($spl19 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14621,f243])).
fof(f14621,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2970])).
fof(f14826,plain,(
  $false | ($spl15 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14620,f243])).
fof(f14620,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2378])).
fof(f14825,plain,(
  $false | ($spl27 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14619,f243])).
fof(f14619,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f5733])).
fof(f14824,plain,(
  $false | ($spl72 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14598,f243])).
fof(f14598,plain,(
  ~p(s(bool,'T_0')) | ($spl72 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f13703])).
fof(f14823,plain,(
  $false | ($spl80 | $spl102 | $spl104)),
  inference(subsumption_resolution,[],[f14450,f243])).
fof(f14450,plain,(
  ~p(s(bool,'T_0')) | ($spl80 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f13712])).
fof(f14818,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f14817,f13776])).
fof(f14817,plain,(
  ( ! [X414,X420,X416,X418,X415,X417,X419] : (~p(s(bool,happ(s(fun(X414,bool),happ(s(fun(X415,fun(X414,bool)),X416),s(X415,X417))),s(X414,happ(s(fun(X418,X414),X419),s(X418,X420))))))) ) | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13882,f13777])).
fof(f13882,plain,(
  ( ! [X414,X420,X416,X418,X415,X417,X419] : (s(bool,'T_0') != s(bool,happ(s(fun(X418,bool),'_dst_x0x0x1_3'(s(fun(X415,fun(X414,bool)),X416),s(X415,X417),s(fun(X418,X414),X419))),s(X418,X420))) | ~p(s(bool,happ(s(fun(X414,bool),happ(s(fun(X415,fun(X414,bool)),X416),s(X415,X417))),s(X414,happ(s(fun(X418,X414),X419),s(X418,X420))))))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f7616])).
fof(f14816,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f14815,f13776])).
fof(f14815,plain,(
  ( ! [X412,X414,X410,X413,X415,X409,X411] : (~p(s(bool,happ(s(fun(X409,bool),'_dst_x0x0x1_3'(s(fun(X410,fun(X411,bool)),X412),s(X410,X413),s(fun(X409,X411),X414))),s(X409,X415))))) ) | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13879,f13777])).
fof(f13879,plain,(
  ( ! [X412,X414,X410,X413,X415,X409,X411] : (s(bool,'T_0') != s(bool,happ(s(fun(X411,bool),happ(s(fun(X410,fun(X411,bool)),X412),s(X410,X413))),s(X411,happ(s(fun(X409,X411),X414),s(X409,X415))))) | ~p(s(bool,happ(s(fun(X409,bool),'_dst_x0x0x1_3'(s(fun(X410,fun(X411,bool)),X412),s(X410,X413),s(fun(X409,X411),X414))),s(X409,X415))))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f7482])).
fof(f14814,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f14813,f13776])).
fof(f14813,plain,(
  ( ! [X292,X294,X293,X295,X291] : (~p(s(bool,happ(s(fun(X291,bool),happ(s(fun(X292,fun(X291,bool)),X293),s(X292,X294))),s(X291,X295))))) ) | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13870,f13777])).
fof(f13870,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s(bool,'T_0') != s(bool,happ(s(fun(X292,bool),'_dst_x1x0_2'(s(fun(X292,fun(X291,bool)),X293),s(X291,X295))),s(X292,X294))) | ~p(s(bool,happ(s(fun(X291,bool),happ(s(fun(X292,fun(X291,bool)),X293),s(X292,X294))),s(X291,X295))))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f6382])).
fof(f14812,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f14811,f13776])).
fof(f14811,plain,(
  ( ! [X280,X277,X279,X281,X278] : (~p(s(bool,happ(s(fun(X277,bool),'_dst_x1x0_2'(s(fun(X277,fun(X278,bool)),X279),s(X278,X280))),s(X277,X281))))) ) | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13867,f13777])).
fof(f13867,plain,(
  ( ! [X280,X277,X279,X281,X278] : (s(bool,'T_0') != s(bool,happ(s(fun(X278,bool),happ(s(fun(X277,fun(X278,bool)),X279),s(X277,X281))),s(X278,X280))) | ~p(s(bool,happ(s(fun(X277,bool),'_dst_x1x0_2'(s(fun(X277,fun(X278,bool)),X279),s(X278,X280))),s(X277,X281))))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f6263])).
fof(f14810,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f14809,f13776])).
fof(f14809,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,happ(s(fun(X210,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X210,X211),X213))),s(X210,X214))))) ) | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13857,f13777])).
fof(f13857,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,'T_0') != s(bool,happ(s(fun(X211,bool),X212),s(X211,happ(s(fun(X210,X211),X213),s(X210,X214))))) | ~p(s(bool,happ(s(fun(X210,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X210,X211),X213))),s(X210,X214))))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f4975])).
fof(f14808,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f14807,f13776])).
fof(f14807,plain,(
  ( ! [X246,X244,X248,X247,X245] : (~p(s(bool,happ(s(fun(X244,bool),X245),s(X244,happ(s(fun(X246,X244),X247),s(X246,X248))))))) ) | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13854,f13777])).
fof(f13854,plain,(
  ( ! [X246,X244,X248,X247,X245] : (s(bool,'T_0') != s(bool,happ(s(fun(X246,bool),'_dst_x0x1_2'(s(fun(X244,bool),X245),s(fun(X246,X244),X247))),s(X246,X248))) | ~p(s(bool,happ(s(fun(X244,bool),X245),s(X244,happ(s(fun(X246,X244),X247),s(X246,X248))))))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f4878])).
fof(f14806,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f13844,f13776])).
fof(f13844,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f3153])).
fof(f14805,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f13841,f13776])).
fof(f13841,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2976])).
fof(f14804,plain,(
  $false | ($spl2 | $spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f13827,f13776])).
fof(f13827,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl2 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2326])).
fof(f14803,plain,(
  $false | ($spl2 | $spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f13826,f13776])).
fof(f13826,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl2 | $spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f2316])).
fof(f14802,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f13809,f13776])).
fof(f13809,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,'T_0'))))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f1433])).
fof(f14801,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f13808,f13776])).
fof(f13808,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'T_0'))),s(bool,X0))))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f1337])).
fof(f14800,plain,(
  $false | ($spl76 | $spl102 | $spl104 | $spl108)),
  inference(subsumption_resolution,[],[f14799,f13776])).
fof(f14799,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13779,f13777])).
fof(f13779,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f257])).
fof(f14798,plain,(
  $false | ($spl102 | $spl104)),
  inference(subsumption_resolution,[],[f13778,f243])).
fof(f13778,plain,(
  ~p(s(bool,'T_0')) | ($spl102 | $spl104)),
  inference(backward_demodulation,[],[f13777,f246])).
fof(f13775,plain,(
  $false | $spl74),
  inference(subsumption_resolution,[],[f13773,f246])).
fof(f13773,plain,(
  p(s(bool,'F_0')) | $spl74),
  inference(backward_demodulation,[],[f13705,f13684])).
fof(f13684,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))),s(fun('type.quote.index',X0),X2))),s('type.quote.index',sK1))))) )),
  inference(forward_demodulation,[],[f13683,f356])).
fof(f13683,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))))),s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',sK1))))))) )),
  inference(forward_demodulation,[],[f13401,f301])).
fof(f301,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9] : ! [X7,X10,X11,X12] : s(X9,happ(s(fun('type.quote.index',X9),X10),s('type.quote.index',X7))) = s(X9,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X7))),s(fun('type.quote.index',X9),X10),s(fun('type.quote.index',X9),X11),s(X9,X12)))),
  file('thm.quote.index_distinct|split|0','thm.quote.index_case_def|split|0')).
fof(f13401,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))))),s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X1),s(X0,X3))))))) )),
  inference(unit_resulting_resolution,[],[f4057,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f4057,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X1),s(X0,X3)))) )),
  inference(superposition,[],[f300,f242])).
fof(f242,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f175])).
fof(f175,plain,(
  ? [X0,X1] : s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1] : s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ~s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ~s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X7))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ~s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X7))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X6)))),
  file('thm.quote.index_distinct|split|0','thm.quote.index_distinct|split|0')).
fof(f300,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X3),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X3),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X3),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : ! [X7,X10,X11,X12] : s(X9,happ(s(fun('type.quote.index',X9),X11),s('type.quote.index',X7))) = s(X9,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X7))),s(fun('type.quote.index',X9),X10),s(fun('type.quote.index',X9),X11),s(X9,X12)))),
  file('thm.quote.index_distinct|split|0','thm.quote.index_case_def|split|1')).
fof(f13769,plain,(
  $false | ($spl62 | $spl66)),
  inference(subsumption_resolution,[],[f13696,f13699])).
fof(f13768,plain,(
  $false | ($spl72 | $spl82)),
  inference(subsumption_resolution,[],[f13703,f13713])).
fof(f13766,plain,(
  ( ! [X874] : (s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) != s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),X874),s('type.quote.index',sK0)))) ) | $spl114),
  inference(cnf_transformation,[],[f13766_D])).
fof(f13766_D,plain,(
  ( ! [X874] : (s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) != s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),X874),s('type.quote.index',sK0)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f13767,plain,(
  ( ! [X873] : (p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),X873),s('type.quote.index',sK1))))))) ) | $spl116),
  inference(cnf_transformation,[],[f13767_D])).
fof(f13767_D,plain,(
  ( ! [X873] : (p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),X873),s('type.quote.index',sK1))))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f13765,plain,(
  ( ! [X874,X873] : (p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),X873),s('type.quote.index',sK1)))))) | s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) != s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),X874),s('type.quote.index',sK0)))) )),
  inference(forward_demodulation,[],[f13672,f301])).
fof(f13672,plain,(
  ( ! [X874,X875,X873] : (s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) != s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),X874),s('type.quote.index',sK0))) | p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index','const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index','type.quote.index'),X873),s(fun('type.quote.index','type.quote.index'),X874),s('type.quote.index',X875))))))) )),
  inference(superposition,[],[f3929,f4057])).
fof(f3929,plain,(
  ( ! [X0] : (s('type.quote.index',X0) != s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) | p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',X0))))) )),
  inference(subsumption_resolution,[],[f3874,f243])).
fof(f3874,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.quote.index',X0) != s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) | p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',X0))))) )),
  inference(superposition,[],[f295,f3475])).
fof(f3475,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index'))))),
  inference(unit_resulting_resolution,[],[f243,f3239,f305])).
fof(f305,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f3239,plain,(
  p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')))))),
  inference(unit_resulting_resolution,[],[f3182,f312])).
fof(f312,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0)))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X44,X45] : (p(s(bool,happ(s(fun(fun(X44,bool),bool),'?_1'),s(fun(X44,bool),X45)))) <=> ? [X41] : p(s(bool,happ(s(fun(X44,bool),X45),s(X44,X41)))))),
  file('thm.quote.index_distinct|split|0','?_HL_QUANT')).
fof(f3182,plain,(
  p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'?_1'),s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))))))),
  inference(unit_resulting_resolution,[],[f3100,f312])).
fof(f3100,plain,(
  p(s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'?_1'))))),
  inference(unit_resulting_resolution,[],[f3046,f313])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f3046,plain,(
  p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'?_1'),s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))),
  inference(unit_resulting_resolution,[],[f2519,f313])).
fof(f2519,plain,(
  p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),
  inference(forward_demodulation,[],[f2518,f338])).
fof(f2518,plain,(
  p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),
  inference(forward_demodulation,[],[f2517,f1679])).
fof(f1679,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,'F_0')),'type.quote.index'))),
  inference(unit_resulting_resolution,[],[f246,f1677])).
fof(f1677,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X1)),'type.quote.index'))) )),
  inference(superposition,[],[f1026,f328])).
fof(f1026,plain,(
  ( ! [X27] : (p(s(bool,happ(s(fun('type.quote.index',bool),X27),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),X27,'type.quote.index'))) )),
  inference(forward_demodulation,[],[f1011,f242])).
fof(f1011,plain,(
  ( ! [X27] : (p(s(bool,happ(s(fun('type.quote.index',bool),X27),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),X27,'type.quote.index'))) )),
  inference(superposition,[],[f293,f242])).
fof(f2517,plain,(
  p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,'F_0')),'type.quote.index')))))),
  inference(forward_demodulation,[],[f2447,f338])).
fof(f2447,plain,(
  p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,'F_0')),'type.quote.index')))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),
  inference(unit_resulting_resolution,[],[f1679,f335])).
fof(f13749,plain,(
  ( ! [X389] : (s(bool,'T_0') != s(bool,happ(s(fun('type.quote.index',bool),X389),s('type.quote.index',sK0)))) ) | $spl110),
  inference(cnf_transformation,[],[f13749_D])).
fof(f13749_D,plain,(
  ( ! [X389] : (s(bool,'T_0') != s(bool,happ(s(fun('type.quote.index',bool),X389),s('type.quote.index',sK0)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f13750,plain,(
  ( ! [X391,X388,X392] : (s(bool,'T_0') = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X391),s(fun('type.quote.index',bool),X392),s(bool,happ(s(fun('type.quote.index',bool),X388),s('type.quote.index',sK1)))))) ) | $spl112),
  inference(cnf_transformation,[],[f13750_D])).
fof(f13750_D,plain,(
  ( ! [X391,X388,X392] : (s(bool,'T_0') = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X391),s(fun('type.quote.index',bool),X392),s(bool,happ(s(fun('type.quote.index',bool),X388),s('type.quote.index',sK1)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f13748,plain,(
  ( ! [X389,X391,X388,X392] : (s(bool,'T_0') = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X391),s(fun('type.quote.index',bool),X392),s(bool,happ(s(fun('type.quote.index',bool),X388),s('type.quote.index',sK1))))) | s(bool,'T_0') != s(bool,happ(s(fun('type.quote.index',bool),X389),s('type.quote.index',sK0)))) )),
  inference(forward_demodulation,[],[f13604,f301])).
fof(f13604,plain,(
  ( ! [X389,X391,X388,X390,X392] : (s(bool,'T_0') != s(bool,happ(s(fun('type.quote.index',bool),X389),s('type.quote.index',sK0))) | s(bool,'T_0') = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X391),s(fun('type.quote.index',bool),X392),s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X388),s(fun('type.quote.index',bool),X389),s(bool,X390)))))) )),
  inference(superposition,[],[f6526,f4057])).
fof(f6526,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,'T_0') = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X33),s(fun('type.quote.index',bool),X34),s(bool,X35)))) )),
  inference(forward_demodulation,[],[f6490,f3188])).
fof(f3188,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'?_1')))),
  inference(unit_resulting_resolution,[],[f243,f3100,f305])).
fof(f6490,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) != s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X33),s(fun('type.quote.index',bool),X34),s(bool,X35)))) )),
  inference(superposition,[],[f3611,f296])).
fof(f296,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X9] : ! [X10,X11,X12] : s(X9,X12) = s(X9,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X9),X10),s(fun('type.quote.index',X9),X11),s(X9,X12)))),
  file('thm.quote.index_distinct|split|0','thm.quote.index_case_def|split|2')).
fof(f3611,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f3558,f243])).
fof(f3558,plain,(
  ( ! [X16] : (~p(s(bool,'T_0')) | s(bool,X16) != s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f398,f3188])).
fof(f13699,plain,(
  ( ! [X286] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X286),s('type.quote.index',sK0))))) ) | $spl66),
  inference(cnf_transformation,[],[f13699_D])).
fof(f13699_D,plain,(
  ( ! [X286] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X286),s('type.quote.index',sK0))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f13746,plain,(
  ( ! [X384,X381] : (s(bool,X384) != s(bool,happ(s(fun('type.quote.index',bool),X381),s('type.quote.index',sK1))) | p(s(bool,X384))) ) | $spl108),
  inference(cnf_transformation,[],[f13746_D])).
fof(f13746_D,plain,(
  ( ! [X384,X381] : (s(bool,X384) != s(bool,happ(s(fun('type.quote.index',bool),X381),s('type.quote.index',sK1))) | p(s(bool,X384))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f13745,plain,(
  ( ! [X382,X384,X381] : (s(bool,X384) != s(bool,happ(s(fun('type.quote.index',bool),X381),s('type.quote.index',sK1))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X382),s('type.quote.index',sK0)))) | p(s(bool,X384))) )),
  inference(forward_demodulation,[],[f13602,f301])).
fof(f13602,plain,(
  ( ! [X383,X382,X384,X381] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X382),s('type.quote.index',sK0)))) | p(s(bool,X384)) | s(bool,X384) != s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X381),s(fun('type.quote.index',bool),X382),s(bool,X383)))) )),
  inference(superposition,[],[f3694,f4057])).
fof(f13743,plain,(
  ( ! [X377,X375] : (p(s(bool,'|dollar|exists_1'(s(fun(X377,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.quote.index',bool),X375),s('type.quote.index',sK0))))))))) ) | $spl106),
  inference(cnf_transformation,[],[f13743_D])).
fof(f13743_D,plain,(
  ( ! [X377,X375] : (p(s(bool,'|dollar|exists_1'(s(fun(X377,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.quote.index',bool),X375),s('type.quote.index',sK0))))))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f13703,plain,(
  ( ! [X288] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X288),s('type.quote.index',sK1))))) ) | $spl72),
  inference(cnf_transformation,[],[f13703_D])).
fof(f13703_D,plain,(
  ( ! [X288] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X288),s('type.quote.index',sK1))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f13742,plain,(
  ( ! [X377,X374,X375] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X374),s('type.quote.index',sK1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X377,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.quote.index',bool),X375),s('type.quote.index',sK0))))))))) )),
  inference(forward_demodulation,[],[f13600,f301])).
fof(f13600,plain,(
  ( ! [X377,X374,X376,X375] : (p(s(bool,'|dollar|exists_1'(s(fun(X377,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.quote.index',bool),X375),s('type.quote.index',sK0)))))))) | ~p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X374),s(fun('type.quote.index',bool),X375),s(bool,X376))))) )),
  inference(superposition,[],[f1271,f4057])).
fof(f1271,plain,(
  ( ! [X4,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) | ~p(s(bool,X5))) )),
  inference(superposition,[],[f290,f328])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK2(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK2(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK2(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X9] : ! [X18,X16] : (! [X19] : (s(X9,X16) = s(X9,X19) => p(s(bool,happ(s(fun(X9,bool),X18),s(X9,X19))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X18)))))),
  file('thm.quote.index_distinct|split|0','thm.bool.JRH_INDUCT_UTIL')).
fof(f13696,plain,(
  ( ! [X283] : (p(s(bool,happ(s(fun('type.quote.index',bool),X283),s('type.quote.index',sK0))))) ) | $spl62),
  inference(cnf_transformation,[],[f13696_D])).
fof(f13696_D,plain,(
  ( ! [X283] : (p(s(bool,happ(s(fun('type.quote.index',bool),X283),s('type.quote.index',sK0))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13739,plain,(
  ( ! [X362,X365] : (s(bool,X365) = s(bool,happ(s(fun('type.quote.index',bool),X362),s('type.quote.index',sK1))) | s(bool,'T_0') = s(bool,X365)) ) | $spl104),
  inference(cnf_transformation,[],[f13739_D])).
fof(f13739_D,plain,(
  ( ! [X362,X365] : (s(bool,X365) = s(bool,happ(s(fun('type.quote.index',bool),X362),s('type.quote.index',sK1))) | s(bool,'T_0') = s(bool,X365)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f13738,plain,(
  ( ! [X362,X365,X363] : (s(bool,X365) = s(bool,happ(s(fun('type.quote.index',bool),X362),s('type.quote.index',sK1))) | p(s(bool,happ(s(fun('type.quote.index',bool),X363),s('type.quote.index',sK0)))) | s(bool,'T_0') = s(bool,X365)) )),
  inference(forward_demodulation,[],[f13597,f301])).
fof(f13597,plain,(
  ( ! [X364,X362,X365,X363] : (p(s(bool,happ(s(fun('type.quote.index',bool),X363),s('type.quote.index',sK0)))) | s(bool,X365) = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X362),s(fun('type.quote.index',bool),X363),s(bool,X364))) | s(bool,'T_0') = s(bool,X365)) )),
  inference(superposition,[],[f454,f4057])).
fof(f13737,plain,(
  ( ! [X358,X361] : (s(bool,X361) != s(bool,happ(s(fun('type.quote.index',bool),X358),s('type.quote.index',sK1))) | s(bool,'T_0') = s(bool,X361)) ) | $spl102),
  inference(cnf_transformation,[],[f13737_D])).
fof(f13737_D,plain,(
  ( ! [X358,X361] : (s(bool,X361) != s(bool,happ(s(fun('type.quote.index',bool),X358),s('type.quote.index',sK1))) | s(bool,'T_0') = s(bool,X361)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f13736,plain,(
  ( ! [X358,X359,X361] : (s(bool,X361) != s(bool,happ(s(fun('type.quote.index',bool),X358),s('type.quote.index',sK1))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X359),s('type.quote.index',sK0)))) | s(bool,'T_0') = s(bool,X361)) )),
  inference(forward_demodulation,[],[f13596,f301])).
fof(f13596,plain,(
  ( ! [X358,X360,X359,X361] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X359),s('type.quote.index',sK0)))) | s(bool,X361) != s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X358),s(fun('type.quote.index',bool),X359),s(bool,X360))) | s(bool,'T_0') = s(bool,X361)) )),
  inference(superposition,[],[f398,f4057])).
fof(f13735,plain,(
  ( ! [X354,X357] : (p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X357),s(fun('type.quote.index',bool),X354))),s('type.quote.index',sK1))))) ) | $spl100),
  inference(cnf_transformation,[],[f13735_D])).
fof(f13735_D,plain,(
  ( ! [X354,X357] : (p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X357),s(fun('type.quote.index',bool),X354))),s('type.quote.index',sK1))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f13734,plain,(
  ( ! [X354,X357,X355] : (p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X357),s(fun('type.quote.index',bool),X354))),s('type.quote.index',sK1)))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X355),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13733,f356])).
fof(f13733,plain,(
  ( ! [X354,X357,X355] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X357))),s(bool,happ(s(fun('type.quote.index',bool),X354),s('type.quote.index',sK1)))))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X355),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13595,f301])).
fof(f13595,plain,(
  ( ! [X356,X354,X357,X355] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X355),s('type.quote.index',sK0)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X357))),s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X354),s(fun('type.quote.index',bool),X355),s(bool,X356))))))) )),
  inference(superposition,[],[f322,f4057])).
fof(f322,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f13732,plain,(
  ( ! [X349,X347] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X347),s('type.quote.index',sK0))))),s(bool,X349))))) ) | $spl98),
  inference(cnf_transformation,[],[f13732_D])).
fof(f13732_D,plain,(
  ( ! [X349,X347] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X347),s('type.quote.index',sK0))))),s(bool,X349))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f13713,plain,(
  ( ! [X306] : (p(s(bool,happ(s(fun('type.quote.index',bool),X306),s('type.quote.index',sK1))))) ) | $spl82),
  inference(cnf_transformation,[],[f13713_D])).
fof(f13713_D,plain,(
  ( ! [X306] : (p(s(bool,happ(s(fun('type.quote.index',bool),X306),s('type.quote.index',sK1))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f13731,plain,(
  ( ! [X346,X349,X347] : (p(s(bool,happ(s(fun('type.quote.index',bool),X346),s('type.quote.index',sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X347),s('type.quote.index',sK0))))),s(bool,X349))))) )),
  inference(forward_demodulation,[],[f13593,f301])).
fof(f13593,plain,(
  ( ! [X346,X349,X347,X348] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X347),s('type.quote.index',sK0))))),s(bool,X349)))) | p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X346),s(fun('type.quote.index',bool),X347),s(bool,X348))))) )),
  inference(superposition,[],[f321,f4057])).
fof(f321,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f233])).
fof(f13730,plain,(
  ( ! [X345,X342] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X342),s('type.quote.index',sK1))))),s(bool,X345)))) | p(s(bool,X345))) ) | $spl96),
  inference(cnf_transformation,[],[f13730_D])).
fof(f13730_D,plain,(
  ( ! [X345,X342] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X342),s('type.quote.index',sK1))))),s(bool,X345)))) | p(s(bool,X345))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f13729,plain,(
  ( ! [X343,X345,X342] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X342),s('type.quote.index',sK1))))),s(bool,X345)))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X343),s('type.quote.index',sK0)))) | p(s(bool,X345))) )),
  inference(forward_demodulation,[],[f13592,f301])).
fof(f13592,plain,(
  ( ! [X344,X343,X345,X342] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X343),s('type.quote.index',sK0)))) | p(s(bool,X345)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X342),s(fun('type.quote.index',bool),X343),s(bool,X344))))),s(bool,X345))))) )),
  inference(superposition,[],[f320,f4057])).
fof(f13728,plain,(
  ( ! [X341,X338] : (p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X341),s(fun('type.quote.index',bool),X338))),s('type.quote.index',sK1))))) ) | $spl94),
  inference(cnf_transformation,[],[f13728_D])).
fof(f13728_D,plain,(
  ( ! [X341,X338] : (p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X341),s(fun('type.quote.index',bool),X338))),s('type.quote.index',sK1))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f13727,plain,(
  ( ! [X341,X339,X338] : (p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X341),s(fun('type.quote.index',bool),X338))),s('type.quote.index',sK1)))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X339),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13726,f356])).
fof(f13726,plain,(
  ( ! [X341,X339,X338] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X341))),s(bool,happ(s(fun('type.quote.index',bool),X338),s('type.quote.index',sK1)))))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X339),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13591,f301])).
fof(f13591,plain,(
  ( ! [X341,X339,X340,X338] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X339),s('type.quote.index',sK0)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X341))),s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X338),s(fun('type.quote.index',bool),X339),s(bool,X340))))))) )),
  inference(superposition,[],[f319,f4057])).
fof(f319,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f13725,plain,(
  ( ! [X337,X334] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X334),s('type.quote.index',sK1))))),s(bool,X337))))) ) | $spl92),
  inference(cnf_transformation,[],[f13725_D])).
fof(f13725_D,plain,(
  ( ! [X337,X334] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X334),s('type.quote.index',sK1))))),s(bool,X337))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f13724,plain,(
  ( ! [X335,X337,X334] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X334),s('type.quote.index',sK1))))),s(bool,X337)))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X335),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13590,f301])).
fof(f13590,plain,(
  ( ! [X335,X337,X334,X336] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X335),s('type.quote.index',sK0)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X334),s(fun('type.quote.index',bool),X335),s(bool,X336))))),s(bool,X337))))) )),
  inference(superposition,[],[f318,f4057])).
fof(f318,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f13723,plain,(
  ( ! [X333,X331] : (~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X333),s(fun('type.quote.index',bool),X331))),s('type.quote.index',sK0)))) | p(s(bool,X333))) ) | $spl90),
  inference(cnf_transformation,[],[f13723_D])).
fof(f13723_D,plain,(
  ( ! [X333,X331] : (~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X333),s(fun('type.quote.index',bool),X331))),s('type.quote.index',sK0)))) | p(s(bool,X333))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f13722,plain,(
  ( ! [X333,X331,X330] : (p(s(bool,happ(s(fun('type.quote.index',bool),X330),s('type.quote.index',sK1)))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X333),s(fun('type.quote.index',bool),X331))),s('type.quote.index',sK0)))) | p(s(bool,X333))) )),
  inference(forward_demodulation,[],[f13721,f301])).
fof(f13721,plain,(
  ( ! [X333,X331,X332,X330] : (~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X333),s(fun('type.quote.index',bool),X331))),s('type.quote.index',sK0)))) | p(s(bool,X333)) | p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X330),s(fun('type.quote.index',bool),X331),s(bool,X332))))) )),
  inference(forward_demodulation,[],[f13589,f356])).
fof(f13589,plain,(
  ( ! [X333,X331,X332,X330] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X333))),s(bool,happ(s(fun('type.quote.index',bool),X331),s('type.quote.index',sK0)))))) | p(s(bool,X333)) | p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X330),s(fun('type.quote.index',bool),X331),s(bool,X332))))) )),
  inference(superposition,[],[f317,f4057])).
fof(f13720,plain,(
  ( ! [X327,X329] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X327),s('type.quote.index',sK0))))),s(bool,X329)))) | p(s(bool,X329))) ) | $spl88),
  inference(cnf_transformation,[],[f13720_D])).
fof(f13720_D,plain,(
  ( ! [X327,X329] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X327),s('type.quote.index',sK0))))),s(bool,X329)))) | p(s(bool,X329))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f13719,plain,(
  ( ! [X327,X329,X326] : (p(s(bool,happ(s(fun('type.quote.index',bool),X326),s('type.quote.index',sK1)))) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X327),s('type.quote.index',sK0))))),s(bool,X329)))) | p(s(bool,X329))) )),
  inference(forward_demodulation,[],[f13588,f301])).
fof(f13588,plain,(
  ( ! [X327,X329,X326,X328] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X327),s('type.quote.index',sK0))))),s(bool,X329)))) | p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X326),s(fun('type.quote.index',bool),X327),s(bool,X328)))) | p(s(bool,X329))) )),
  inference(superposition,[],[f317,f4057])).
fof(f13718,plain,(
  ( ! [X325,X322] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X322),s('type.quote.index',sK1))))),s(bool,X325)))) | ~p(s(bool,X325))) ) | $spl86),
  inference(cnf_transformation,[],[f13718_D])).
fof(f13718_D,plain,(
  ( ! [X325,X322] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X322),s('type.quote.index',sK1))))),s(bool,X325)))) | ~p(s(bool,X325))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f13717,plain,(
  ( ! [X325,X323,X322] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X322),s('type.quote.index',sK1))))),s(bool,X325)))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X323),s('type.quote.index',sK0)))) | ~p(s(bool,X325))) )),
  inference(forward_demodulation,[],[f13587,f301])).
fof(f13587,plain,(
  ( ! [X325,X323,X324,X322] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X323),s('type.quote.index',sK0)))) | ~p(s(bool,X325)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X322),s(fun('type.quote.index',bool),X323),s(bool,X324))))),s(bool,X325))))) )),
  inference(superposition,[],[f316,f4057])).
fof(f316,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f13716,plain,(
  ( ! [X321,X319] : (~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X321),s(fun('type.quote.index',bool),X319))),s('type.quote.index',sK0))))) ) | $spl84),
  inference(cnf_transformation,[],[f13716_D])).
fof(f13716_D,plain,(
  ( ! [X321,X319] : (~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X321),s(fun('type.quote.index',bool),X319))),s('type.quote.index',sK0))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f13715,plain,(
  ( ! [X321,X318,X319] : (p(s(bool,happ(s(fun('type.quote.index',bool),X318),s('type.quote.index',sK1)))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X321),s(fun('type.quote.index',bool),X319))),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13714,f301])).
fof(f13714,plain,(
  ( ! [X321,X318,X320,X319] : (~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X321),s(fun('type.quote.index',bool),X319))),s('type.quote.index',sK0)))) | p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X318),s(fun('type.quote.index',bool),X319),s(bool,X320))))) )),
  inference(forward_demodulation,[],[f13586,f356])).
fof(f13586,plain,(
  ( ! [X321,X318,X320,X319] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,X321))),s(bool,happ(s(fun('type.quote.index',bool),X319),s('type.quote.index',sK0)))))) | p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X318),s(fun('type.quote.index',bool),X319),s(bool,X320))))) )),
  inference(superposition,[],[f315,f4057])).
fof(f13712,plain,(
  ( ! [X309,X307] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X307),s('type.quote.index',sK0))))),s(bool,X309))))) ) | $spl80),
  inference(cnf_transformation,[],[f13712_D])).
fof(f13712_D,plain,(
  ( ! [X309,X307] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X307),s('type.quote.index',sK0))))),s(bool,X309))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f13711,plain,(
  ( ! [X306,X309,X307] : (p(s(bool,happ(s(fun('type.quote.index',bool),X306),s('type.quote.index',sK1)))) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X307),s('type.quote.index',sK0))))),s(bool,X309))))) )),
  inference(forward_demodulation,[],[f13583,f301])).
fof(f13583,plain,(
  ( ! [X308,X306,X309,X307] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(bool,happ(s(fun('type.quote.index',bool),X307),s('type.quote.index',sK0))))),s(bool,X309)))) | p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X306),s(fun('type.quote.index',bool),X307),s(bool,X308))))) )),
  inference(superposition,[],[f314,f4057])).
fof(f13710,plain,(
  ( ! [X302,X305] : (s(bool,X305) = s(bool,happ(s(fun('type.quote.index',bool),X302),s('type.quote.index',sK1))) | ~p(s(bool,X305))) ) | $spl78),
  inference(cnf_transformation,[],[f13710_D])).
fof(f13710_D,plain,(
  ( ! [X302,X305] : (s(bool,X305) = s(bool,happ(s(fun('type.quote.index',bool),X302),s('type.quote.index',sK1))) | ~p(s(bool,X305))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f13709,plain,(
  ( ! [X302,X303,X305] : (s(bool,X305) = s(bool,happ(s(fun('type.quote.index',bool),X302),s('type.quote.index',sK1))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X303),s('type.quote.index',sK0)))) | ~p(s(bool,X305))) )),
  inference(forward_demodulation,[],[f13582,f301])).
fof(f13582,plain,(
  ( ! [X302,X304,X303,X305] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X303),s('type.quote.index',sK0)))) | ~p(s(bool,X305)) | s(bool,X305) = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X302),s(fun('type.quote.index',bool),X303),s(bool,X304)))) )),
  inference(superposition,[],[f305,f4057])).
fof(f13707,plain,(
  ( ! [X294,X297] : (s(bool,X297) = s(bool,happ(s(fun('type.quote.index',bool),X294),s('type.quote.index',sK1))) | p(s(bool,X297))) ) | $spl76),
  inference(cnf_transformation,[],[f13707_D])).
fof(f13707_D,plain,(
  ( ! [X294,X297] : (s(bool,X297) = s(bool,happ(s(fun('type.quote.index',bool),X294),s('type.quote.index',sK1))) | p(s(bool,X297))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f13708,plain,(
  ( ! [X298,X301,X299] : (s(bool,X301) = s(bool,happ(s(fun('type.quote.index',bool),X298),s('type.quote.index',sK1))) | p(s(bool,happ(s(fun('type.quote.index',bool),X299),s('type.quote.index',sK0)))) | p(s(bool,X301))) )),
  inference(forward_demodulation,[],[f13581,f301])).
fof(f13581,plain,(
  ( ! [X300,X298,X301,X299] : (p(s(bool,happ(s(fun('type.quote.index',bool),X299),s('type.quote.index',sK0)))) | p(s(bool,X301)) | s(bool,X301) = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X298),s(fun('type.quote.index',bool),X299),s(bool,X300)))) )),
  inference(superposition,[],[f302,f4057])).
fof(f13706,plain,(
  ( ! [X294,X295,X297] : (s(bool,X297) = s(bool,happ(s(fun('type.quote.index',bool),X294),s('type.quote.index',sK1))) | p(s(bool,happ(s(fun('type.quote.index',bool),X295),s('type.quote.index',sK0)))) | p(s(bool,X297))) )),
  inference(forward_demodulation,[],[f13580,f301])).
fof(f13580,plain,(
  ( ! [X294,X296,X295,X297] : (p(s(bool,happ(s(fun('type.quote.index',bool),X295),s('type.quote.index',sK0)))) | p(s(bool,X297)) | s(bool,X297) = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X294),s(fun('type.quote.index',bool),X295),s(bool,X296)))) )),
  inference(superposition,[],[f302,f4057])).
fof(f13705,plain,(
  ( ! [X291] : (s(bool,'F_0') = s(bool,happ(s(fun('type.quote.index',bool),X291),s('type.quote.index',sK1)))) ) | $spl74),
  inference(cnf_transformation,[],[f13705_D])).
fof(f13705_D,plain,(
  ( ! [X291] : (s(bool,'F_0') = s(bool,happ(s(fun('type.quote.index',bool),X291),s('type.quote.index',sK1)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f13704,plain,(
  ( ! [X292,X291] : (s(bool,'F_0') = s(bool,happ(s(fun('type.quote.index',bool),X291),s('type.quote.index',sK1))) | p(s(bool,happ(s(fun('type.quote.index',bool),X292),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13579,f301])).
fof(f13579,plain,(
  ( ! [X292,X293,X291] : (p(s(bool,happ(s(fun('type.quote.index',bool),X292),s('type.quote.index',sK0)))) | s(bool,'F_0') = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X291),s(fun('type.quote.index',bool),X292),s(bool,X293)))) )),
  inference(superposition,[],[f258,f4057])).
fof(f13702,plain,(
  ( ! [X289] : (s(bool,'F_0') != s(bool,happ(s(fun('type.quote.index',bool),X289),s('type.quote.index',sK0)))) ) | $spl70),
  inference(cnf_transformation,[],[f13702_D])).
fof(f13702_D,plain,(
  ( ! [X289] : (s(bool,'F_0') != s(bool,happ(s(fun('type.quote.index',bool),X289),s('type.quote.index',sK0)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f13701,plain,(
  ( ! [X288,X289] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X288),s('type.quote.index',sK1)))) | s(bool,'F_0') != s(bool,happ(s(fun('type.quote.index',bool),X289),s('type.quote.index',sK0)))) )),
  inference(forward_demodulation,[],[f13578,f301])).
fof(f13578,plain,(
  ( ! [X288,X290,X289] : (s(bool,'F_0') != s(bool,happ(s(fun('type.quote.index',bool),X289),s('type.quote.index',sK0))) | ~p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X288),s(fun('type.quote.index',bool),X289),s(bool,X290))))) )),
  inference(superposition,[],[f257,f4057])).
fof(f13700,plain,(
  ( ! [X285] : (s(bool,'T_0') = s(bool,happ(s(fun('type.quote.index',bool),X285),s('type.quote.index',sK1)))) ) | $spl68),
  inference(cnf_transformation,[],[f13700_D])).
fof(f13700_D,plain,(
  ( ! [X285] : (s(bool,'T_0') = s(bool,happ(s(fun('type.quote.index',bool),X285),s('type.quote.index',sK1)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f13698,plain,(
  ( ! [X285,X286] : (s(bool,'T_0') = s(bool,happ(s(fun('type.quote.index',bool),X285),s('type.quote.index',sK1))) | ~p(s(bool,happ(s(fun('type.quote.index',bool),X286),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13577,f301])).
fof(f13577,plain,(
  ( ! [X285,X287,X286] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X286),s('type.quote.index',sK0)))) | s(bool,'T_0') = s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X285),s(fun('type.quote.index',bool),X286),s(bool,X287)))) )),
  inference(superposition,[],[f254,f4057])).
fof(f254,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f196])).
fof(f13697,plain,(
  ( ! [X282] : (s(bool,'T_0') != s(bool,happ(s(fun('type.quote.index',bool),X282),s('type.quote.index',sK1)))) ) | $spl64),
  inference(cnf_transformation,[],[f13697_D])).
fof(f13697_D,plain,(
  ( ! [X282] : (s(bool,'T_0') != s(bool,happ(s(fun('type.quote.index',bool),X282),s('type.quote.index',sK1)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13695,plain,(
  ( ! [X282,X283] : (s(bool,'T_0') != s(bool,happ(s(fun('type.quote.index',bool),X282),s('type.quote.index',sK1))) | p(s(bool,happ(s(fun('type.quote.index',bool),X283),s('type.quote.index',sK0))))) )),
  inference(forward_demodulation,[],[f13576,f301])).
fof(f13576,plain,(
  ( ! [X282,X283,X284] : (p(s(bool,happ(s(fun('type.quote.index',bool),X283),s('type.quote.index',sK0)))) | s(bool,'T_0') != s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',bool),X282),s(fun('type.quote.index',bool),X283),s(bool,X284)))) )),
  inference(superposition,[],[f253,f4057])).
fof(f12313,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ARB_0') | $spl58),
  inference(cnf_transformation,[],[f12313_D])).
fof(f12313_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ARB_0') <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f12315,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'(bool),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'(bool),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.|48|_0'),s(bool,'F_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x01_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(bool,'F_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s('type.ind_type.recspace'(bool),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'F_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'))))))),s(fun('type.quote.index','type.ind_type.recspace'(bool)),sK13)))) | $spl60),
  inference(cnf_transformation,[],[f12315_D])).
fof(f12315_D,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'(bool),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'(bool),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.|48|_0'),s(bool,'F_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x01_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(bool,'F_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s('type.ind_type.recspace'(bool),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'F_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'))))))),s(fun('type.quote.index','type.ind_type.recspace'(bool)),sK13)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f12277,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'(bool),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'(bool),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.|48|_0'),s(bool,'F_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x01_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(bool,'F_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s('type.ind_type.recspace'(bool),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'F_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'))))))),s(fun('type.quote.index','type.ind_type.recspace'(bool)),sK13)))) | s(bool,'T_0') = s(bool,'const.bool.ARB_0')),
  inference(superposition,[],[f376,f252])).
fof(f376,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'(bool),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'(bool),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.|48|_0'),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x01_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s('type.ind_type.recspace'(bool),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.bool.ARB_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'))))))),s(fun('type.quote.index','type.ind_type.recspace'(bool)),sK13))))),
  inference(forward_demodulation,[],[f363,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),'const.ind_type.CONSTR_2'(s('type.num.num',X1),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(forward_demodulation,[],[f298,f288])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),'const.ind_type.CONSTR_2'(s('type.num.num',X1),s(X0,X2))) = s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),'const.ind_type.CONSTR_2'(s('type.num.num',X1),s(X0,X2))) = s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),'const.ind_type.CONSTR_2'(s('type.num.num',X1),s(X0,X2))) = s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X9] : ! [X8,X33] : s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),'const.ind_type.CONSTR_2'(s('type.num.num',X8),s(X9,X33))) = s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X33)))),
  file('thm.quote.index_distinct|split|0','ARITY_1_2_const.ind_type.CONSTR')).
fof(f298,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9] : ! [X8,X33,X34] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X33),s(fun('type.num.num','type.ind_type.recspace'(X9)),X34))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X33))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X34)))),
  file('thm.quote.index_distinct|split|0','ARITY_1_3_const.ind_type.CONSTR')).
fof(f363,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'(bool),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'(bool),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.|48|_0'),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x01_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s('type.ind_type.recspace'(bool),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'))))))),s(fun('type.quote.index','type.ind_type.recspace'(bool)),sK13))))),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'(bool),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'(bool),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.|48|_0'),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x01_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s('type.ind_type.recspace'(bool),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'))))))),s(fun('type.quote.index','type.ind_type.recspace'(bool)),sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f174])).
fof(f174,plain,(
  ? [X0] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'(bool),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'(bool),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.|48|_0'),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x01_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s('type.ind_type.recspace'(bool),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'))))))),s(fun('type.quote.index','type.ind_type.recspace'(bool)),X0))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ? [X37] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'(bool),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'(bool),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.|48|_0'),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11x01_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'(bool),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool))),'_dst_x000x011x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'),s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_dst_x0x1_2'(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(bool,'const.bool.ARB_0'))),s(fun('type.ind_type.recspace'(bool),fun('type.num.num','type.ind_type.recspace'(bool))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun(fun('type.num.num','type.ind_type.recspace'(bool)),fun('type.num.num','type.ind_type.recspace'(bool)))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'(bool),fun('type.ind_type.recspace'(bool),bool)),'=_0'),s('type.ind_type.recspace'(bool),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(bool)),'type.ind_type.recspace'(bool)),'const.ind_type.CONSTR_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'(bool)),'_dst_x0_1'(s('type.ind_type.recspace'(bool),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'(bool),bool),fun('type.ind_type.recspace'(bool),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'(bool),fun(fun('type.ind_type.recspace'(bool),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'(bool),'type.ind_type.recspace'(bool)),'_i_0_0'))))))),s(fun('type.quote.index','type.ind_type.recspace'(bool)),X37))))),
  file('thm.quote.index_distinct|split|0','thm.quote.index_TY_DEF')).
fof(f12239,plain,(
  $false | $spl41),
  inference(subsumption_resolution,[],[f12157,f12169])).
fof(f12169,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl41),
  inference(unit_resulting_resolution,[],[f328,f11913,f3694])).
fof(f11913,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK6(X0,'type.quote.index')))))) ) | $spl41),
  inference(unit_resulting_resolution,[],[f11784,f311])).
fof(f311,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X44,X45] : (p(s(bool,happ(s(fun(fun(X44,bool),bool),'!_1'),s(fun(X44,bool),X45)))) <=> ! [X41] : p(s(bool,happ(s(fun(X44,bool),X45),s(X44,X41)))))),
  file('thm.quote.index_distinct|split|0','!_HL_QUANT')).
fof(f11784,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),X0))))) ) | $spl41),
  inference(unit_resulting_resolution,[],[f8911,f313])).
fof(f12157,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK6(X0,'type.quote.index')))))),s(bool,X1))))) ) | $spl41),
  inference(unit_resulting_resolution,[],[f11913,f321])).
fof(f12170,plain,(
  $false | $spl41),
  inference(unit_resulting_resolution,[],[f243,f328,f11913,f3694])).
fof(f12144,plain,(
  $false | $spl41),
  inference(unit_resulting_resolution,[],[f328,f11913,f253])).
fof(f10478,plain,(
  $false | ($spl4 | $spl7)),
  inference(subsumption_resolution,[],[f10477,f10138])).
fof(f10138,plain,(
  s('type.quote.index',sK0) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))) | $spl4),
  inference(backward_demodulation,[],[f10137,f242])).
fof(f10137,plain,(
  s('type.quote.index',sK0) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) | $spl4),
  inference(backward_demodulation,[],[f1789,f1679])).
fof(f10477,plain,(
  s('type.quote.index',sK0) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))) | ($spl4 | $spl7)),
  inference(forward_demodulation,[],[f1791,f10137])).
fof(f10472,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10471,f10138])).
fof(f10471,plain,(
  s('type.quote.index',sK0) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10470,f10138])).
fof(f10470,plain,(
  s('type.quote.index',sK0) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2065,f10137])).
fof(f10469,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10468,f10138])).
fof(f10468,plain,(
  s('type.quote.index',sK0) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10467,f10138])).
fof(f10467,plain,(
  s('type.quote.index',sK0) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10466,f10138])).
fof(f10466,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2218,f10137])).
fof(f10465,plain,(
  $false | ($spl4 | $spl15)),
  inference(trivial_inequality_removal,[],[f10464])).
fof(f10464,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10463,f10138])).
fof(f10463,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10462,f10138])).
fof(f10462,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2219,f10137])).
fof(f10461,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10460,f10138])).
fof(f10460,plain,(
  s('type.quote.index',sK0) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10459,f10138])).
fof(f10459,plain,(
  s('type.quote.index',sK0) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10458,f10138])).
fof(f10458,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10457,f10138])).
fof(f10457,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2224,f10137])).
fof(f2224,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2218,f308])).
fof(f10456,plain,(
  $false | ($spl4 | $spl15)),
  inference(trivial_inequality_removal,[],[f10455])).
fof(f10455,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10454,f10138])).
fof(f10454,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10453,f10138])).
fof(f10453,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10452,f10138])).
fof(f10452,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2225,f10137])).
fof(f2225,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2218,f306])).
fof(f10451,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10450,f2548])).
fof(f2548,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f2547,f338])).
fof(f2547,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f2546,f286])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X9] : ! [X13,X14] : s(X9,X13) = s(X9,'const.bool.COND_3'(s(bool,'T_0'),s(X9,X13),s(X9,X14)))),
  file('thm.quote.index_distinct|split|0','thm.bool.COND_CLAUSES|split|0')).
fof(f2546,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2))))))) )),
  inference(forward_demodulation,[],[f2437,f338])).
fof(f2437,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f286,f335])).
fof(f10450,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index',sK0)))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10449,f10138])).
fof(f10449,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10448,f10138])).
fof(f10448,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2232,f10137])).
fof(f10447,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10446,f2548])).
fof(f10446,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index',sK0)))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10445,f10138])).
fof(f10445,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10444,f10138])).
fof(f10444,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10443,f10138])).
fof(f10443,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2234,f10137])).
fof(f10442,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10441,f2548])).
fof(f10441,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index',sK0)))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10440,f338])).
fof(f10440,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index',sK0)))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10439,f10138])).
fof(f10439,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10438,f10138])).
fof(f10438,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10437,f10138])).
fof(f10437,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2310,f10137])).
fof(f10436,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10435,f2548])).
fof(f10435,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index',sK0)))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10434,f338])).
fof(f10434,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index',sK0)))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10433,f10138])).
fof(f10433,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10432,f10138])).
fof(f10432,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index',sK0))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2312,f10137])).
fof(f10431,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10430,f2548])).
fof(f10430,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0)))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10429,f10138])).
fof(f10429,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10428,f10138])).
fof(f10428,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2378,f10137])).
fof(f10427,plain,(
  $false | ($spl4 | $spl15)),
  inference(trivial_inequality_removal,[],[f10426])).
fof(f10426,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10425,f10138])).
fof(f10425,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10424,f10138])).
fof(f10424,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2380,f10137])).
fof(f2380,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2219,f308])).
fof(f10423,plain,(
  $false | ($spl4 | $spl15)),
  inference(trivial_inequality_removal,[],[f10422])).
fof(f10422,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10421,f10138])).
fof(f10421,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10420,f10138])).
fof(f10420,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0))))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2381,f10137])).
fof(f2381,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f2219,f306])).
fof(f10419,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10418,f2548])).
fof(f10418,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0)))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10417,f338])).
fof(f10417,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0)))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10416,f10138])).
fof(f10416,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f10415,f10138])).
fof(f10415,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))))))))) | ($spl4 | $spl15)),
  inference(forward_demodulation,[],[f2383,f10137])).
fof(f10129,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10128,f2548])).
fof(f10128,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10127,f338])).
fof(f10127,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10079,f1790])).
fof(f1790,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | $spl6),
  inference(cnf_transformation,[],[f1790_D])).
fof(f1790_D,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f10079,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2383])).
fof(f10126,plain,(
  $false | ($spl6 | $spl15)),
  inference(trivial_inequality_removal,[],[f10125])).
fof(f10125,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10078,f1790])).
fof(f10078,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2381])).
fof(f10124,plain,(
  $false | ($spl6 | $spl15)),
  inference(trivial_inequality_removal,[],[f10123])).
fof(f10123,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10077,f1790])).
fof(f10077,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2380])).
fof(f10122,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10121,f2548])).
fof(f10121,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10076,f1790])).
fof(f10076,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2378])).
fof(f10120,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10119,f2548])).
fof(f10119,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10118,f338])).
fof(f10118,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10075,f1790])).
fof(f10075,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2312])).
fof(f10117,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10116,f2548])).
fof(f10116,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10115,f338])).
fof(f10115,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10114,f1790])).
fof(f10114,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10074,f1790])).
fof(f10074,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),'_dst_x1x0_2'(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2310])).
fof(f10113,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10112,f2548])).
fof(f10112,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10111,f1790])).
fof(f10111,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10073,f1790])).
fof(f10073,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2234])).
fof(f10110,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10109,f2548])).
fof(f10109,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10072,f1790])).
fof(f10072,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2232])).
fof(f10108,plain,(
  $false | ($spl6 | $spl15)),
  inference(trivial_inequality_removal,[],[f10107])).
fof(f10107,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10106,f1790])).
fof(f10106,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10071,f1790])).
fof(f10071,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2225])).
fof(f10105,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10104,f1790])).
fof(f10104,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10103,f1790])).
fof(f10103,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10070,f1790])).
fof(f10070,plain,(
  s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2224])).
fof(f10102,plain,(
  $false | ($spl6 | $spl15)),
  inference(trivial_inequality_removal,[],[f10101])).
fof(f10101,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10069,f1790])).
fof(f10069,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2219])).
fof(f10100,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10099,f1790])).
fof(f10099,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | ($spl6 | $spl15)),
  inference(forward_demodulation,[],[f10068,f1790])).
fof(f10068,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2218])).
fof(f10098,plain,(
  $false | ($spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10067,f1790])).
fof(f10067,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | ($spl6 | $spl15)),
  inference(backward_demodulation,[],[f1790,f2065])).
fof(f10064,plain,(
  ( ! [X19] : (s('type.quote.index',sK0) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X19)),'type.quote.index')) | s(bool,'T_0') = s(bool,X19)) ) | $spl56),
  inference(cnf_transformation,[],[f10064_D])).
fof(f10064_D,plain,(
  ( ! [X19] : (s('type.quote.index',sK0) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X19)),'type.quote.index')) | s(bool,'T_0') = s(bool,X19)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1791,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | $spl7),
  inference(cnf_transformation,[],[f1791_D])).
fof(f1791_D,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1898,plain,(
  ( ! [X19] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | s('type.quote.index',sK0) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X19)),'type.quote.index')) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f868,f1694])).
fof(f1694,plain,(
  ( ! [X0] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X0)),'type.quote.index')) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1680,f246])).
fof(f1680,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X0)),'type.quote.index')) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1677,f252])).
fof(f868,plain,(
  ( ! [X0] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0))) | s('type.quote.index',sK0) = s('type.quote.index',X0)) )),
  inference(superposition,[],[f308,f242])).
fof(f10063,plain,(
  ( ! [X20] : (s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'_i_0_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X20)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X20)) ) | $spl54),
  inference(cnf_transformation,[],[f10063_D])).
fof(f10063_D,plain,(
  ( ! [X20] : (s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'_i_0_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X20)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X20)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1899,plain,(
  ( ! [X20] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'_i_0_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X20)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f885,f1694])).
fof(f885,plain,(
  ( ! [X2] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X2))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'_i_0_0'),s('type.quote.index',X2)))) )),
  inference(superposition,[],[f868,f282])).
fof(f10062,plain,(
  ( ! [X21,X22] : (s('type.quote.index',sK0) = s('type.quote.index','const.bool.COND_3'(s(bool,'F_0'),s('type.quote.index',X22),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X21)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X21)) ) | $spl52),
  inference(cnf_transformation,[],[f10062_D])).
fof(f10062_D,plain,(
  ( ! [X21,X22] : (s('type.quote.index',sK0) = s('type.quote.index','const.bool.COND_3'(s(bool,'F_0'),s('type.quote.index',X22),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X21)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X21)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1900,plain,(
  ( ! [X21,X22] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | s('type.quote.index',sK0) = s('type.quote.index','const.bool.COND_3'(s(bool,'F_0'),s('type.quote.index',X22),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X21)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f887,f1694])).
fof(f887,plain,(
  ( ! [X6,X7] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X7))) | s('type.quote.index',sK0) = s('type.quote.index','const.bool.COND_3'(s(bool,'F_0'),s('type.quote.index',X6),s('type.quote.index',X7)))) )),
  inference(superposition,[],[f868,f285])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X9] : ! [X13,X14] : s(X9,X14) = s(X9,'const.bool.COND_3'(s(bool,'F_0'),s(X9,X13),s(X9,X14)))),
  file('thm.quote.index_distinct|split|0','thm.bool.COND_CLAUSES|split|1')).
fof(f10061,plain,(
  ( ! [X24,X23] : (s('type.quote.index',sK0) = s('type.quote.index','const.bool.COND_3'(s(bool,'T_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X23)),'type.quote.index')),s('type.quote.index',X24))) | s(bool,'T_0') = s(bool,X23)) ) | $spl50),
  inference(cnf_transformation,[],[f10061_D])).
fof(f10061_D,plain,(
  ( ! [X24,X23] : (s('type.quote.index',sK0) = s('type.quote.index','const.bool.COND_3'(s(bool,'T_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X23)),'type.quote.index')),s('type.quote.index',X24))) | s(bool,'T_0') = s(bool,X23)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1901,plain,(
  ( ! [X24,X23] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | s('type.quote.index',sK0) = s('type.quote.index','const.bool.COND_3'(s(bool,'T_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X23)),'type.quote.index')),s('type.quote.index',X24))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f888,f1694])).
fof(f888,plain,(
  ( ! [X8,X9] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X8))) | s('type.quote.index',sK0) = s('type.quote.index','const.bool.COND_3'(s(bool,'T_0'),s('type.quote.index',X8),s('type.quote.index',X9)))) )),
  inference(superposition,[],[f868,f286])).
fof(f10060,plain,(
  s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',sK0))) | $spl48),
  inference(cnf_transformation,[],[f10060_D])).
fof(f10060_D,plain,(
  s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',sK0))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2058,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',sK0)))),
  inference(superposition,[],[f883,f242])).
fof(f883,plain,(
  ( ! [X0] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0))))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',X0)))) )),
  inference(superposition,[],[f868,f247])).
fof(f247,plain,(
  ( ! [X0] : (s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0))) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0))) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X8))) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',X8)))),
  file('thm.quote.index_distinct|split|0','ARITY_0_1_const.quote.Right_idx')).
fof(f10056,plain,(
  ( ! [X12,X10,X11] : (s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun(X11,'type.quote.index'),'_dst_x0_1'(s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X10)),'type.quote.index')))),s(X11,X12))) | s(bool,'T_0') = s(bool,X10)) ) | $spl46),
  inference(cnf_transformation,[],[f10056_D])).
fof(f10056_D,plain,(
  ( ! [X12,X10,X11] : (s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun(X11,'type.quote.index'),'_dst_x0_1'(s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X10)),'type.quote.index')))),s(X11,X12))) | s(bool,'T_0') = s(bool,X10)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7190,plain,(
  ( ! [X12,X10,X11] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun(X11,'type.quote.index'),'_dst_x0_1'(s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X10)),'type.quote.index')))),s(X11,X12))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f886,f1694])).
fof(f886,plain,(
  ( ! [X4,X5,X3] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X4))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun(X3,'type.quote.index'),'_dst_x0_1'(s('type.quote.index',X4))),s(X3,X5)))) )),
  inference(superposition,[],[f868,f328])).
fof(f10053,plain,(
  ( ! [X6,X7] : (s('type.quote.index',sK0) = s('type.quote.index','const.ind_type.FCONS_3'(s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X6)),'type.quote.index')),s(fun('type.num.num','type.quote.index'),X7),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl44),
  inference(cnf_transformation,[],[f10053_D])).
fof(f10053_D,plain,(
  ( ! [X6,X7] : (s('type.quote.index',sK0) = s('type.quote.index','const.ind_type.FCONS_3'(s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X6)),'type.quote.index')),s(fun('type.num.num','type.quote.index'),X7),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7226,plain,(
  ( ! [X6,X7] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | s('type.quote.index',sK0) = s('type.quote.index','const.ind_type.FCONS_3'(s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X6)),'type.quote.index')),s(fun('type.num.num','type.quote.index'),X7),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f889,f1694])).
fof(f889,plain,(
  ( ! [X10,X11] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X10))) | s('type.quote.index',sK0) = s('type.quote.index','const.ind_type.FCONS_3'(s('type.quote.index',X10),s(fun('type.num.num','type.quote.index'),X11),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f868,f287])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9] : ! [X7,X10] : s(X9,X7) = s(X9,'const.ind_type.FCONS_3'(s(X9,X7),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.quote.index_distinct|split|0','thm.ind_type.FCONS|split|0')).
fof(f10050,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f9966,f9979])).
fof(f9979,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f328,f9727,f3694])).
fof(f9727,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK6(X0,'type.quote.index')))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8900,f311])).
fof(f9966,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.quote.index',bool),X0),s('type.quote.index',sK6(X0,'type.quote.index')))))),s(bool,X1))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f9727,f321])).
fof(f9980,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f243,f328,f9727,f3694])).
fof(f9953,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f328,f9727,f253])).
fof(f8911,plain,(
  ~p(s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'!_1')))) | $spl41),
  inference(cnf_transformation,[],[f8911_D])).
fof(f8911_D,plain,(
  ~p(s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'!_1')))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f8913,plain,(
  p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK7(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool)))))) | $spl42),
  inference(cnf_transformation,[],[f8913_D])).
fof(f8913_D,plain,(
  p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK7(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool)))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f8909,plain,(
  p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK7(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool)))))) | ~p(s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'!_1')))) | $spl7),
  inference(forward_demodulation,[],[f8842,f8253])).
fof(f8253,plain,(
  s(fun(fun('type.quote.index',bool),bool),'!_1') = s(fun(fun('type.quote.index',bool),bool),sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool))) | $spl7),
  inference(unit_resulting_resolution,[],[f8155,f293])).
fof(f8155,plain,(
  ~p(s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'!_1'),s(fun(fun('type.quote.index',bool),bool),'!_1')))) | $spl7),
  inference(unit_resulting_resolution,[],[f7785,f310])).
fof(f310,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f7785,plain,(
  ~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f2231,f310])).
fof(f2231,plain,(
  ~p(s(bool,happ(s(fun('type.quote.index',bool),happ(s(fun('type.quote.index',fun('type.quote.index',bool)),'=_0'),s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))),s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1791,f334])).
fof(f8842,plain,(
  ~p(s(bool,happ(s(fun(fun(fun('type.quote.index',bool),bool),bool),'?_1'),s(fun(fun('type.quote.index',bool),bool),'!_1')))) | p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool))),s(fun('type.quote.index',bool),sK7(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool)))))) | $spl7),
  inference(superposition,[],[f312,f8253])).
fof(f8900,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),X2))))) ) | $spl36),
  inference(cnf_transformation,[],[f8900_D])).
fof(f8900_D,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),X2))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f8901,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK3(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),X3,fun('type.quote.index',bool)))))) | p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),X3),s(fun('type.quote.index',bool),sK4(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),X3,fun('type.quote.index',bool))))))) ) | $spl38),
  inference(cnf_transformation,[],[f8901_D])).
fof(f8901_D,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK3(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),X3,fun('type.quote.index',bool)))))) | p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),X3),s(fun('type.quote.index',bool),sK4(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),X3,fun('type.quote.index',bool))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f8899,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK3(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),X3,fun('type.quote.index',bool)))))) | ~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),X2)))) | p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),X3),s(fun('type.quote.index',bool),sK4(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),X3,fun('type.quote.index',bool))))))) ) | $spl7),
  inference(forward_demodulation,[],[f8835,f8253])).
fof(f8835,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),X2)))) | p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),X3),s(fun('type.quote.index',bool),sK4(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),X3,fun('type.quote.index',bool)))))) | p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool))),s(fun('type.quote.index',bool),sK3(sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),X3,fun('type.quote.index',bool))))))) ) | $spl7),
  inference(superposition,[],[f291,f8253])).
fof(f291,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK3(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK3(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4])],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9] : ! [X17,X18] : (! [X19] : (p(s(bool,happ(s(fun(X9,bool),X18),s(X9,X19)))) => p(s(bool,happ(s(fun(X9,bool),X17),s(X9,X19))))) => (? [X19] : p(s(bool,happ(s(fun(X9,bool),X18),s(X9,X19)))) => ? [X19] : p(s(bool,happ(s(fun(X9,bool),X17),s(X9,X19))))))),
  file('thm.quote.index_distinct|split|0','thm.bool.MONO_EXISTS')).
fof(f8898,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK2(X1,sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool))))))) ) | $spl34),
  inference(cnf_transformation,[],[f8898_D])).
fof(f8898_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK2(X1,sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f8896,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(fun('type.quote.index',bool),bool),'!_1')))) | $spl32),
  inference(cnf_transformation,[],[f8896_D])).
fof(f8896_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(fun('type.quote.index',bool),bool),'!_1')))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f8897,plain,(
  ( ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(fun('type.quote.index',bool),bool),'!_1')))) | ~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK2(X1,sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool))))))) ) | $spl7),
  inference(forward_demodulation,[],[f8834,f8253])).
fof(f8834,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(fun('type.quote.index',bool),bool),'!_1'),s(fun('type.quote.index',bool),sK2(X1,sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool)))))) | p(s(bool,'|dollar|exists_1'(s(fun(fun('type.quote.index',bool),bool),sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool))))))) ) | $spl7),
  inference(superposition,[],[f290,f8253])).
fof(f8894,plain,(
  ( ! [X0] : (s(fun('type.quote.index',bool),X0) = s(fun('type.quote.index',bool),sK2(X0,sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool)))) ) | $spl30),
  inference(cnf_transformation,[],[f8894_D])).
fof(f8894_D,plain,(
  ( ! [X0] : (s(fun('type.quote.index',bool),X0) = s(fun('type.quote.index',bool),sK2(X0,sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f8833,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(fun('type.quote.index',bool),bool),'!_1')))) | s(fun('type.quote.index',bool),X0) = s(fun('type.quote.index',bool),sK2(X0,sK5('!_1','!_1',fun(fun('type.quote.index',bool),bool)),fun('type.quote.index',bool)))) ) | $spl7),
  inference(superposition,[],[f289,f8253])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK2(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f5525,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5254,f5430])).
fof(f5430,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f335])).
fof(f5254,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f258])).
fof(f5524,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5256,f5430])).
fof(f5256,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f246,f2040,f302])).
fof(f5523,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5257,f5430])).
fof(f5257,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f2040,f302])).
fof(f5522,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5272,f246])).
fof(f5272,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f380,f2040,f302])).
fof(f5521,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f5274,f2040])).
fof(f5274,plain,(
  ( ! [X0] : (p(s(bool,sK2(X0,'_i_0_0',bool)))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1289,f2040,f302])).
fof(f5520,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5276,f2040])).
fof(f5276,plain,(
  p(s(bool,sK5('F_0','_i_0_0',bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1075,f2040,f302])).
fof(f5519,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f5278,f2040])).
fof(f5278,plain,(
  ( ! [X0] : (p(s(bool,sK5(sK2(X0,'_i_0_0',bool),'_i_0_0',bool)))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1538,f2040,f302])).
fof(f5518,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5280,f2040])).
fof(f5280,plain,(
  p(s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1250,f2040,f302])).
fof(f5517,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(subsumption_resolution,[],[f5282,f2040])).
fof(f5282,plain,(
  p(s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool))) | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2876,f2040,f302])).
fof(f5516,plain,(
  $false | ($spl10 | $spl23)),
  inference(subsumption_resolution,[],[f5284,f2040])).
fof(f5284,plain,(
  p(s(bool,sK6('_i_0_0',bool))) | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2670,f2040,f302])).
fof(f5515,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5286,f5430])).
fof(f5286,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f246,f2040,f302])).
fof(f5514,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5287,f5430])).
fof(f5287,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f2040,f302])).
fof(f5513,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5302,f246])).
fof(f5302,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f380,f2040,f302])).
fof(f5512,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f5304,f2040])).
fof(f5304,plain,(
  ( ! [X0] : (p(s(bool,sK2(X0,'_i_0_0',bool)))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1289,f2040,f302])).
fof(f5511,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5306,f2040])).
fof(f5306,plain,(
  p(s(bool,sK5('F_0','_i_0_0',bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1075,f2040,f302])).
fof(f5510,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f5308,f2040])).
fof(f5308,plain,(
  ( ! [X0] : (p(s(bool,sK5(sK2(X0,'_i_0_0',bool),'_i_0_0',bool)))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1538,f2040,f302])).
fof(f5509,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5310,f2040])).
fof(f5310,plain,(
  p(s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1250,f2040,f302])).
fof(f5508,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(subsumption_resolution,[],[f5312,f2040])).
fof(f5312,plain,(
  p(s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool))) | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2876,f2040,f302])).
fof(f5507,plain,(
  $false | ($spl10 | $spl23)),
  inference(subsumption_resolution,[],[f5314,f2040])).
fof(f5314,plain,(
  p(s(bool,sK6('_i_0_0',bool))) | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2670,f2040,f302])).
fof(f5506,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5323,f2040])).
fof(f5323,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f321])).
fof(f5505,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5332,f380])).
fof(f5332,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f380,f2040,f454])).
fof(f5504,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f5333,f1289])).
fof(f5333,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK2(X0,'_i_0_0',bool))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1289,f2040,f454])).
fof(f5503,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5334,f1075])).
fof(f5334,plain,(
  s(bool,'T_0') = s(bool,sK5('F_0','_i_0_0',bool)) | $spl10),
  inference(unit_resulting_resolution,[],[f1075,f2040,f454])).
fof(f5502,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f5335,f1538])).
fof(f5335,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5(sK2(X0,'_i_0_0',bool),'_i_0_0',bool))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1538,f2040,f454])).
fof(f5501,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5336,f1250])).
fof(f5336,plain,(
  s(bool,'T_0') = s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool)) | $spl10),
  inference(unit_resulting_resolution,[],[f1250,f2040,f454])).
fof(f5500,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(subsumption_resolution,[],[f5337,f2876])).
fof(f5337,plain,(
  s(bool,'T_0') = s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool)) | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2876,f2040,f454])).
fof(f5499,plain,(
  $false | ($spl10 | $spl23)),
  inference(subsumption_resolution,[],[f5338,f2670])).
fof(f5338,plain,(
  s(bool,'T_0') = s(bool,sK6('_i_0_0',bool)) | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2670,f2040,f454])).
fof(f5498,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5339,f5430])).
fof(f5339,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f380,f2040,f454])).
fof(f5497,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f5341,f5430])).
fof(f5341,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,sK2(X1,'_i_0_0',bool))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1289,f2040,f454])).
fof(f5496,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5343,f5430])).
fof(f5343,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK5('F_0','_i_0_0',bool))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1075,f2040,f454])).
fof(f5495,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f5345,f5430])).
fof(f5345,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,sK5(sK2(X1,'_i_0_0',bool),'_i_0_0',bool))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1538,f2040,f454])).
fof(f5494,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5347,f5430])).
fof(f5347,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK5(sK5('F_0','_i_0_0',bool),'_i_0_0',bool))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1250,f2040,f454])).
fof(f5493,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(subsumption_resolution,[],[f5349,f5430])).
fof(f5349,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK5(sK6('_i_0_0',bool),'_i_0_0',bool))) ) | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2876,f2040,f454])).
fof(f5492,plain,(
  $false | ($spl10 | $spl23)),
  inference(subsumption_resolution,[],[f5351,f5430])).
fof(f5351,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('_i_0_0',bool))) ) | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2670,f2040,f454])).
fof(f5491,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5353,f5430])).
fof(f5353,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK5(X0,'_i_0_0',bool))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1018])).
fof(f5490,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5357,f5430])).
fof(f5357,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,sK2(X1,X2,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f289])).
fof(f5489,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5429,f5430])).
fof(f5429,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,sK5(X1,X2,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f293])).
fof(f5488,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5441,f2040])).
fof(f5441,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f321])).
fof(f5450,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f282,f2040,f3929])).
fof(f5445,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f243,f243,f2040,f316])).
fof(f5442,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f246,f2040,f321])).
fof(f5440,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f243,f2040,f322])).
fof(f5438,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f243,f2040,f318])).
fof(f5436,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f243,f2040,f319])).
fof(f5431,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f282,f2040,f335])).
fof(f5358,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1289,f2040,f289])).
fof(f5352,plain,(
  $false | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2670,f2670,f2040,f454])).
fof(f5350,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2876,f2876,f2040,f454])).
fof(f5348,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1250,f1250,f2040,f454])).
fof(f5346,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1538,f1538,f2040,f454])).
fof(f5344,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1075,f1075,f2040,f454])).
fof(f5342,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1289,f1289,f2040,f454])).
fof(f5340,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f380,f380,f2040,f454])).
fof(f5324,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f2040,f321])).
fof(f5315,plain,(
  $false | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2040,f2670,f2040,f302])).
fof(f5313,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2040,f2876,f2040,f302])).
fof(f5311,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1250,f2040,f302])).
fof(f5309,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f2040,f1538,f2040,f302])).
fof(f5307,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1075,f2040,f302])).
fof(f5305,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f2040,f1289,f2040,f302])).
fof(f5303,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f246,f380,f2040,f302])).
fof(f5301,plain,(
  $false | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2040,f2670,f2040,f302])).
fof(f5299,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2040,f2876,f2040,f302])).
fof(f5297,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1250,f2040,f302])).
fof(f5295,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f2040,f1538,f2040,f302])).
fof(f5293,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1075,f2040,f302])).
fof(f5291,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f2040,f1289,f2040,f302])).
fof(f5289,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f380,f2040,f302])).
fof(f5285,plain,(
  $false | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2040,f2670,f2040,f302])).
fof(f5283,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2040,f2876,f2040,f302])).
fof(f5281,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1250,f2040,f302])).
fof(f5279,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f2040,f1538,f2040,f302])).
fof(f5277,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1075,f2040,f302])).
fof(f5275,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f2040,f1289,f2040,f302])).
fof(f5273,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f246,f380,f2040,f302])).
fof(f5271,plain,(
  $false | ($spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2040,f2670,f2040,f302])).
fof(f5269,plain,(
  $false | ($spl2 | $spl10 | $spl23)),
  inference(unit_resulting_resolution,[],[f2040,f2876,f2040,f302])).
fof(f5267,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1250,f2040,f302])).
fof(f5265,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f2040,f1538,f2040,f302])).
fof(f5263,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f1075,f2040,f302])).
fof(f5261,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f2040,f1289,f2040,f302])).
fof(f5259,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2040,f380,f2040,f302])).
fof(f5255,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f380,f2040,f258])).
fof(f5253,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f281,f2040,f253])).
fof(f281,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X9] : ! [X19] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X9,X19)))),
  file('thm.quote.index_distinct|split|0','thm.bool.DATATYPE_TAG_THM')).
fof(f5251,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f243,f2040])).
fof(f2040,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl10),
  inference(cnf_transformation,[],[f2040_D])).
fof(f2040_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5248,plain,(
  ( ! [X10] : (p(s(bool,sK3('_i_0_0',X10,bool))) | p(s(bool,happ(s(fun(bool,bool),X10),s(bool,sK4('_i_0_0',X10,bool)))))) ) | $spl28),
  inference(cnf_transformation,[],[f5248_D])).
fof(f5248_D,plain,(
  ( ! [X10] : (p(s(bool,sK3('_i_0_0',X10,bool))) | p(s(bool,happ(s(fun(bool,bool),X10),s(bool,sK4('_i_0_0',X10,bool)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5247,plain,(
  ( ! [X10,X9] : (p(s(bool,sK3('_i_0_0',X10,bool))) | ~p(s(bool,X9)) | p(s(bool,happ(s(fun(bool,bool),X10),s(bool,sK4('_i_0_0',X10,bool)))))) )),
  inference(forward_demodulation,[],[f5228,f282])).
fof(f5228,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | p(s(bool,happ(s(fun(bool,bool),X10),s(bool,sK4('_i_0_0',X10,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK3('_i_0_0',X10,bool)))))) )),
  inference(superposition,[],[f291,f282])).
fof(f4754,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(cnf_transformation,[],[f4754_D])).
fof(f4754_D,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index'))))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4756,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) | $spl27),
  inference(cnf_transformation,[],[f4756_D])).
fof(f4756_D,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f4727,plain,(
  ( ! [X2] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index',sK7(sK7('?_1',fun('type.quote.index',bool)),'type.quote.index')) | p(s(bool,happ(s(fun('type.quote.index',bool),sK7('?_1',fun('type.quote.index',bool))),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f3929,f1694])).
fof(f2667,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl20),
  inference(cnf_transformation,[],[f2667_D])).
fof(f2667_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2669,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | $spl23),
  inference(cnf_transformation,[],[f2669_D])).
fof(f2669_D,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f2657,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f311,f282])).
fof(f2400,plain,(
  ( ! [X2] : (s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Left_idx_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(cnf_transformation,[],[f2400_D])).
fof(f2400_D,plain,(
  ( ! [X2] : (s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Left_idx_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2402,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | $spl19),
  inference(cnf_transformation,[],[f2402_D])).
fof(f2402_D,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f2388,plain,(
  ( ! [X2] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Left_idx_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f884,f1694])).
fof(f884,plain,(
  ( ! [X1] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Left_idx_0'),s('type.quote.index',X1)))) )),
  inference(superposition,[],[f868,f248])).
fof(f248,plain,(
  ( ! [X0] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Left_idx_0'),s('type.quote.index',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Left_idx_0'),s('type.quote.index',X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X8))) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Left_idx_0'),s('type.quote.index',X8)))),
  file('thm.quote.index_distinct|split|0','ARITY_0_1_const.quote.Left_idx')).
fof(f2063,plain,(
  ( ! [X2] : (s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(cnf_transformation,[],[f2063_D])).
fof(f2063_D,plain,(
  ( ! [X2] : (s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2065,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | $spl15),
  inference(cnf_transformation,[],[f2065_D])).
fof(f2065_D,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f2051,plain,(
  ( ! [X2] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))))) | s('type.quote.index',sK0) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'const.quote.Right_idx_0'),s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,X2)),'type.quote.index')))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f883,f1694])).
fof(f2039,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | $spl8),
  inference(cnf_transformation,[],[f2039_D])).
fof(f2039_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2030,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f313,f282])).
fof(f1789,plain,(
  s('type.quote.index',sK0) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,'F_0')),'type.quote.index')) | $spl4),
  inference(cnf_transformation,[],[f1789_D])).
fof(f1789_D,plain,(
  s('type.quote.index',sK0) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,'F_0')),'type.quote.index')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1768,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))))) | s('type.quote.index',sK0) = s('type.quote.index',sK5('const.quote.Right_idx_1'(s('type.quote.index',sK0)),'_dst_x0_1'(s(bool,'F_0')),'type.quote.index'))),
  inference(superposition,[],[f868,f1679])).
fof(f1279,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl0),
  inference(cnf_transformation,[],[f1279_D])).
fof(f1279_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1280,plain,(
  ( ! [X3] : (~p(s(bool,sK2(X3,'_i_0_0',bool)))) ) | $spl2),
  inference(cnf_transformation,[],[f1280_D])).
fof(f1280_D,plain,(
  ( ! [X3] : (~p(s(bool,sK2(X3,'_i_0_0',bool)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1270,plain,(
  ( ! [X3] : (~p(s(bool,sK2(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f290,f282])).
% SZS output end Proof for thm.quote
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 184
% Passive clauses: 3634
% Generated clauses: 18207
% Final active clauses: 74
% Final passive clauses: 608
% Input formulas: 69
% Initial clauses: 122
% 
% Duplicate literals: 11
% Trivial inequalities: 59
% Fw subsumption resolutions: 498
% Fw demodulations: 2953
% Bw demodulations: 2528
% Condensations: 51
% Global subsumptions: 2
% 
% Simple tautologies: 53
% Equational tautologies: 78
% Forward subsumptions: 10482
% Fw demodulations to eq. taut.: 76
% Bw demodulations to eq. taut.: 34
% 
% Unit resulting resolution: 2129
% Factoring: 1
% Forward superposition: 3545
% Backward superposition: 3806
% Self superposition: 134
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 43
% Split components: 86
% SAT solver clauses: 3991
% SAT solver unit clauses: 742
% SAT solver binary clauses: 2222
% 
% Sat splits: 43
% Sat splitting refutations: 267
% 
% Memory used [KB]: 31214
% Time elapsed: 1.411 s
% ------------------------------
% ------------------------------
% Success in time 2.74 s
