Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (110309)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 152
% Passive clauses: 3696
% Generated clauses: 18570
% Final active clauses: 144
% Final passive clauses: 3412
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 540
% Trivial inequalities: 304
% Fw subsumption resolutions: 885
% Fw demodulations: 1047
% Bw demodulations: 4
% 
% Simple tautologies: 310
% Equational tautologies: 22
% Forward subsumptions: 11442
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 8557
% Backward superposition: 7074
% Self superposition: 208
% Equality factoring: 54
% Equality resolution: 24
% 
% Split clauses: 219
% Split components: 438
% SAT solver clauses: 426
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 219
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 426
% TWLsolver calls for satisfiability: 52
% 
% Memory used [KB]: 19957
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (110331)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 122
% Passive clauses: 1263
% Generated clauses: 15468
% Final active clauses: 110
% Final passive clauses: 997
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 581
% Trivial inequalities: 198
% Fw subsumption resolutions: 556
% Fw demodulations: 1018
% Bw demodulations: 121
% Condensations: 889
% 
% Simple tautologies: 175
% Equational tautologies: 25
% Forward subsumptions: 8724
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 5788
% Backward superposition: 4558
% Self superposition: 210
% Equality factoring: 74
% Equality resolution: 17
% 
% Split clauses: 102
% Split components: 204
% SAT solver clauses: 1991
% SAT solver unit clauses: 235
% SAT solver binary clauses: 49
% 
% Sat splits: 102
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 1991
% TWLsolver calls for satisfiability: 257
% 
% Memory used [KB]: 18038
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f132586,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f132436])).
fof(f132436,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')),
  inference(superposition,[],[f307,f132226])).
fof(f132226,plain,(
  ( ! [X2] : (s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f30260,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X7] : ! [X12,X6] : s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X7),X12),s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|28','thm.canonical.canonical_sum_scalar2_def|split|2')).
fof(f30260,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.list.list'('type.quote.index'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(forward_demodulation,[],[f30225,f420])).
fof(f420,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar2_0'),s('type.list.list'('type.quote.index'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar2_0'),s('type.list.list'('type.quote.index'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X45,X46] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X45),s('type.canonical.canonical_sum'('type.num.num'),X46))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar2_0'),s('type.list.list'('type.quote.index'),X45))),s('type.canonical.canonical_sum'('type.num.num'),X46)))),
  file('thm.numRing.num_ring_thms|split|28','ARITY_0_2_const.numRing.num_canonical_sum_scalar2')).
fof(f30225,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar2_0'),s('type.list.list'('type.quote.index'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.list.list'('type.quote.index'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(superposition,[],[f367,f319])).
fof(f319,plain,(
  s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar2_0') = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f70])).
fof(f70,axiom,(
  s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar2_0') = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|28','thm.numRing.num_canonical_sum_scalar2_def')).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X7] : ! [X45,X46,X47] : s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X7),X45),s('type.list.list'('type.quote.index'),X46),s('type.canonical.canonical_sum'(X7),X47))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.list.list'('type.quote.index'),X46))),s('type.canonical.canonical_sum'(X7),X47)))),
  file('thm.numRing.num_ring_thms|split|28','ARITY_1_3_const.canonical.canonical_sum_scalar2')).
fof(f307,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),sK0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),sK0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f298])).
fof(f298,plain,(
  ? [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar2_2'(s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|28','thm.numRing.num_ring_thms|split|28')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 561
% Passive clauses: 10348
% Generated clauses: 135252
% Final active clauses: 238
% Final passive clauses: 1597
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 6674
% 
% Duplicate literals: 3760
% Trivial inequalities: 1232
% Fw subsumption resolutions: 3474
% Fw demodulations: 15339
% Bw demodulations: 2110
% Condensations: 5225
% Global subsumptions: 209
% 
% Simple tautologies: 633
% Equational tautologies: 271
% Forward subsumptions: 93319
% Fw demodulations to eq. taut.: 25
% Bw demodulations to eq. taut.: 212
% 
% Unit resulting resolution: 2798
% Forward superposition: 50289
% Backward superposition: 45273
% Self superposition: 777
% Equality factoring: 233
% Equality resolution: 50
% 
% Split clauses: 345
% Split components: 695
% SAT solver clauses: 11485
% SAT solver unit clauses: 1264
% SAT solver binary clauses: 2788
% 
% Sat splits: 345
% Sat splitting refutations: 236
% 
% Memory used [KB]: 162726
% Time elapsed: 7.553 s
% ------------------------------
% ------------------------------
% Success in time 8.85 s
