Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (76762)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 118
% Passive clauses: 4130
% Generated clauses: 13984
% Final active clauses: 111
% Final passive clauses: 3739
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 501
% Trivial inequalities: 228
% Fw subsumption resolutions: 1486
% Fw demodulations: 469
% Bw demodulations: 5
% 
% Simple tautologies: 300
% Equational tautologies: 88
% Forward subsumptions: 5773
% 
% Binary resolution: 282
% Factoring: 2
% Forward superposition: 4626
% Backward superposition: 5658
% Self superposition: 292
% Equality factoring: 190
% Equality resolution: 35
% 
% Split clauses: 465
% Split components: 972
% SAT solver clauses: 1062
% SAT solver unit clauses: 37
% SAT solver binary clauses: 372
% 
% Sat splits: 465
% 
% TWLsolver clauses: 1032
% TWLsolver calls for satisfiability: 131
% 
% Memory used [KB]: 11641
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (76775)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 134
% Passive clauses: 1290
% Generated clauses: 11318
% Final active clauses: 115
% Final passive clauses: 1021
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 178
% Trivial inequalities: 57
% Fw subsumption resolutions: 119
% Fw demodulations: 812
% Bw demodulations: 151
% Condensations: 368
% Global subsumptions: 12
% 
% Simple tautologies: 91
% Equational tautologies: 40
% Forward subsumptions: 6271
% Fw demodulations to eq. taut.: 5
% 
% Binary resolution: 136
% Factoring: 1
% Forward superposition: 4106
% Backward superposition: 3883
% Self superposition: 198
% Equality factoring: 50
% Equality resolution: 23
% 
% Split clauses: 25
% Split components: 50
% SAT solver clauses: 2206
% SAT solver unit clauses: 158
% SAT solver binary clauses: 71
% SAT solver learnt clauses: 24
% 
% Sat splits: 25
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2158
% TWLsolver calls for satisfiability: 420
% 
% Memory used [KB]: 24562
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 137 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f116207,plain,(
  $false),
  inference(subsumption_resolution,[],[f116206,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X1)))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X20,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X20)))),
  file('thm.ringNorm.polynom_distinct|split|1','thm.ringNorm.polynom_distinct|split|0')).
fof(f116206,plain,(
  s('type.ringNorm.polynom'('type.quote.index'),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'('type.quote.index'),'const.ringNorm.Pconst_1'(s('type.quote.index',sK3)))),
  inference(forward_demodulation,[],[f114463,f114462])).
fof(f114462,plain,(
  ( ! [X130,X131] : (s('type.ringNorm.polynom'(X130),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'(X130),happ(s(fun('type.quote.index','type.ringNorm.polynom'(X130)),X131),s('type.quote.index',sK3)))) )),
  inference(superposition,[],[f109235,f297])).
fof(f297,plain,(
  ( ! [X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(X0)),'const.ringNorm.Pvar_0'),s('type.quote.index',X1)))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(X0)),'const.ringNorm.Pvar_0'),s('type.quote.index',X1)))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(X0)),'const.ringNorm.Pvar_0'),s('type.quote.index',X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X10] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X10))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.quote.index','type.ringNorm.polynom'(X6)),'const.ringNorm.Pvar_0'),s('type.quote.index',X10)))),
  file('thm.ringNorm.polynom_distinct|split|1','ARITY_0_1_const.ringNorm.Pvar')).
fof(f109235,plain,(
  ( ! [X2,X0,X3] : (s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',sK3))) = s(X0,happ(s(fun('type.quote.index',X0),X3),s('type.quote.index',sK3)))) )),
  inference(superposition,[],[f34557,f34557])).
fof(f34557,plain,(
  ( ! [X6,X7,X9] : (s(X6,happ(s(fun('type.quote.index',X6),X7),s('type.quote.index',sK3))) = s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X6)),X9),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(superposition,[],[f11615,f342])).
fof(f342,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X14] : ! [X9,X15,X16,X17,X18,X19] : s(X14,happ(s(fun('type.quote.index',X14),X15),s('type.quote.index',X9))) = s(X14,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X9))),s(fun('type.quote.index',X14),X15),s(fun(X6,X14),X16),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X17),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X18),s(fun('type.ringNorm.polynom'(X6),X14),X19)))),
  file('thm.ringNorm.polynom_distinct|split|1','thm.ringNorm.polynom_case_def|split|0')).
fof(f11615,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),X0),X5)))) )),
  inference(superposition,[],[f345,f245])).
fof(f245,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|1','thm.ringNorm.polynom_distinct|split|1')).
fof(f345,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14,X6] : ! [X8,X7,X15,X16,X17,X18,X19] : s(X14,happ(s(fun('type.ringNorm.polynom'(X6),X14),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X17),s('type.ringNorm.polynom'(X6),X8))),s('type.ringNorm.polynom'(X6),X7))) = s(X14,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7))),s(fun('type.quote.index',X14),X15),s(fun(X6,X14),X16),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X17),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X18),s(fun('type.ringNorm.polynom'(X6),X14),X19)))),
  file('thm.ringNorm.polynom_distinct|split|1','thm.ringNorm.polynom_case_def|split|2')).
fof(f114463,plain,(
  ( ! [X132] : (s('type.ringNorm.polynom'('type.quote.index'),'const.ringNorm.Pconst_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'('type.quote.index'),happ(s(fun('type.quote.index','type.ringNorm.polynom'('type.quote.index')),X132),s('type.quote.index',sK3)))) )),
  inference(superposition,[],[f109235,f296])).
fof(f296,plain,(
  ( ! [X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun(X0,'type.ringNorm.polynom'(X0)),'const.ringNorm.Pconst_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun(X0,'type.ringNorm.polynom'(X0)),'const.ringNorm.Pconst_0'),s(X0,X1)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun(X0,'type.ringNorm.polynom'(X0)),'const.ringNorm.Pconst_0'),s(X0,X1)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X10] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X10))) = s('type.ringNorm.polynom'(X6),happ(s(fun(X6,'type.ringNorm.polynom'(X6)),'const.ringNorm.Pconst_0'),s(X6,X10)))),
  file('thm.ringNorm.polynom_distinct|split|1','ARITY_0_1_const.ringNorm.Pconst')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 495
% Passive clauses: 7988
% Generated clauses: 119409
% Final active clauses: 310
% Final passive clauses: 1813
% Input formulas: 69
% Initial clauses: 125
% Discarded non-redundant clauses: 7550
% 
% Duplicate literals: 1489
% Trivial inequalities: 701
% Fw subsumption resolutions: 3881
% Fw demodulations: 26275
% Bw demodulations: 2308
% Condensations: 2622
% Global subsumptions: 284
% 
% Simple tautologies: 462
% Equational tautologies: 602
% Forward subsumptions: 65113
% Fw demodulations to eq. taut.: 6050
% Bw demodulations to eq. taut.: 539
% 
% Unit resulting resolution: 24601
% Factoring: 5
% Forward superposition: 26278
% Backward superposition: 25320
% Self superposition: 407
% Equality factoring: 192
% Equality resolution: 77
% 
% Split clauses: 124
% Split components: 254
% SAT solver clauses: 8851
% SAT solver unit clauses: 606
% SAT solver binary clauses: 3395
% 
% Sat splits: 124
% Sat splitting refutations: 329
% 
% Memory used [KB]: 104390
% Time elapsed: 6.734 s
% ------------------------------
% ------------------------------
% Success in time 8.051 s
