Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (112912)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 87
% Passive clauses: 1532
% Generated clauses: 13546
% Final active clauses: 84
% Final passive clauses: 1345
% Input formulas: 19
% Initial clauses: 20
% 
% Duplicate literals: 938
% Trivial inequalities: 186
% Fw subsumption resolutions: 1847
% Fw demodulations: 344
% 
% Simple tautologies: 296
% Equational tautologies: 98
% Forward subsumptions: 6362
% 
% Binary resolution: 51
% Forward superposition: 6111
% Backward superposition: 3248
% Self superposition: 388
% Equality factoring: 283
% Equality resolution: 14
% 
% Split clauses: 744
% Split components: 1502
% SAT solver clauses: 1969
% SAT solver unit clauses: 244
% SAT solver binary clauses: 594
% 
% Sat splits: 744
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1961
% TWLsolver calls for satisfiability: 73
% 
% Memory used [KB]: 6780
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (112925)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 92
% Passive clauses: 290
% Generated clauses: 10895
% Final active clauses: 87
% Final passive clauses: 192
% Input formulas: 19
% Initial clauses: 20
% 
% Duplicate literals: 1239
% Trivial inequalities: 82
% Fw subsumption resolutions: 669
% Fw demodulations: 667
% Bw demodulations: 7
% Condensations: 1946
% 
% Simple tautologies: 170
% Equational tautologies: 43
% Forward subsumptions: 5002
% 
% Binary resolution: 14
% Forward superposition: 2863
% Backward superposition: 2813
% Self superposition: 209
% Equality factoring: 132
% Equality resolution: 4
% 
% Split clauses: 89
% Split components: 183
% SAT solver clauses: 832
% SAT solver unit clauses: 34
% SAT solver binary clauses: 67
% 
% Sat splits: 89
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 832
% TWLsolver calls for satisfiability: 95
% 
% Memory used [KB]: 8827
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f15463,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f15441])).
fof(f15441,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3)))))),
  inference(superposition,[],[f59,f1394])).
fof(f1394,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X4),X6),s('type.canonical.spolynom'(X4),X7))))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X6))),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X7)))))) )),
  inference(forward_demodulation,[],[f1393,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(forward_demodulation,[],[f287,f68])).
fof(f68,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X15,X16] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X15),s('type.canonical.spolynom'(X6),X16))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X6),X15))),s('type.canonical.spolynom'(X6),X16)))),
  file('thm.ringNorm.spolynom_normalize_def|split|3','ARITY_1_2_const.ringNorm.r_spolynom_normalize')).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(superposition,[],[f67,f63])).
fof(f63,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|3','thm.ringNorm.r_spolynom_normalize_def')).
fof(f67,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X15,X16] : s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X15),s('type.canonical.spolynom'(X6),X16))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X6),X15))),s('type.canonical.spolynom'(X6),X16)))),
  file('thm.ringNorm.spolynom_normalize_def|split|3','ARITY_1_2_const.canonical.spolynom_normalize')).
fof(f1393,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X4),X6),s('type.canonical.spolynom'(X4),X7))))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X6))),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X7)))))) )),
  inference(forward_demodulation,[],[f1392,f289])).
fof(f1392,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X4),X6),s('type.canonical.spolynom'(X4),X7))))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),X6))),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X7)))))) )),
  inference(forward_demodulation,[],[f1375,f289])).
fof(f1375,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X4),X6),s('type.canonical.spolynom'(X4),X7))))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),X6))),s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),X7)))))) )),
  inference(superposition,[],[f643,f77])).
fof(f77,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X12,X8,X7] : s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7))))) = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.spolynom'(X6),X8))),s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.spolynom'(X6),X7)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|3','thm.canonical.spolynom_normalize_def|split|3')).
fof(f643,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f589,f75])).
fof(f75,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X6),X15),s('type.canonical.canonical_sum'(X6),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X6),X15))),s('type.canonical.canonical_sum'(X6),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.spolynom_normalize_def|split|3','ARITY_1_3_const.ringNorm.r_canonical_sum_prod')).
fof(f589,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f73,f65])).
fof(f65,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|3','thm.ringNorm.r_canonical_sum_prod_def')).
fof(f73,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X6),X15),s('type.canonical.canonical_sum'(X6),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X6),X15))),s('type.canonical.canonical_sum'(X6),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.spolynom_normalize_def|split|3','ARITY_1_3_const.canonical.canonical_sum_prod')).
fof(f59,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK3)))))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f54])).
fof(f54,plain,(
  ? [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X9))))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X8))),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X9))))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X8))),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X9)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|3','thm.ringNorm.spolynom_normalize_def|split|3')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 119
% Passive clauses: 413
% Generated clauses: 15159
% Final active clauses: 90
% Final passive clauses: 223
% Input formulas: 19
% Initial clauses: 20
% 
% Duplicate literals: 1787
% Trivial inequalities: 147
% Fw subsumption resolutions: 752
% Fw demodulations: 467
% Bw demodulations: 71
% Condensations: 2957
% Global subsumptions: 4
% 
% Simple tautologies: 205
% Equational tautologies: 117
% Forward subsumptions: 8307
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 141
% Forward superposition: 4254
% Backward superposition: 4014
% Self superposition: 306
% Equality factoring: 198
% Equality resolution: 12
% 
% Split clauses: 71
% Split components: 148
% SAT solver clauses: 572
% SAT solver unit clauses: 123
% SAT solver binary clauses: 28
% 
% Sat splits: 71
% Sat splitting refutations: 11
% 
% Memory used [KB]: 9722
% Time elapsed: 0.962 s
% ------------------------------
% ------------------------------
% Success in time 2.282 s
