Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.list
% (70495)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 118
% Passive clauses: 3241
% Generated clauses: 24374
% Final active clauses: 113
% Final passive clauses: 2933
% Input formulas: 19
% Initial clauses: 36
% 
% Duplicate literals: 978
% Trivial inequalities: 369
% Fw subsumption resolutions: 3017
% Fw demodulations: 3817
% Bw demodulations: 2
% 
% Simple tautologies: 328
% Equational tautologies: 45
% Forward subsumptions: 11407
% 
% Binary resolution: 109
% Factoring: 2
% Forward superposition: 7686
% Backward superposition: 7852
% Self superposition: 255
% Equality factoring: 115
% Equality resolution: 24
% 
% Split clauses: 372
% Split components: 753
% SAT solver clauses: 844
% SAT solver unit clauses: 4
% SAT solver binary clauses: 43
% 
% Sat splits: 372
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 842
% TWLsolver calls for satisfiability: 86
% 
% Memory used [KB]: 17142
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (70527)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 83
% Passive clauses: 1017
% Generated clauses: 10876
% Final active clauses: 79
% Final passive clauses: 924
% Input formulas: 19
% Initial clauses: 36
% 
% Duplicate literals: 472
% Trivial inequalities: 69
% Fw subsumption resolutions: 625
% Fw demodulations: 1042
% Bw demodulations: 9
% Condensations: 891
% Global subsumptions: 5
% 
% Simple tautologies: 121
% Equational tautologies: 37
% Forward subsumptions: 5059
% 
% Binary resolution: 107
% Factoring: 2
% Forward superposition: 3697
% Backward superposition: 2724
% Self superposition: 171
% Equality factoring: 77
% Equality resolution: 7
% 
% Split clauses: 58
% Split components: 128
% SAT solver clauses: 1583
% SAT solver unit clauses: 45
% SAT solver binary clauses: 52
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 1
% 
% Sat splits: 58
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1562
% TWLsolver calls for satisfiability: 160
% 
% Memory used [KB]: 16758
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f2323,plain,(
  $false),
  inference(subsumption_resolution,[],[f2322,f94])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X13,X14] : ~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X14),s('type.list.list'(X6),X13)))),
  file('thm.list.FRONT_CONS|split|1','thm.list.NOT_CONS_NIL')).
fof(f2322,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))),
  inference(subsumption_resolution,[],[f2272,f68])).
fof(f68,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.FRONT_CONS|split|1',aHL_TRUTH)).
fof(f2272,plain,(
  ~p(s(bool,'T_0')) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))),
  inference(superposition,[],[f95,f2195])).
fof(f2195,plain,(
  s(bool,'T_0') = s(bool,sK4('const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)),sK1,sK0))),
  inference(unit_resulting_resolution,[],[f67,f329])).
fof(f329,plain,(
  ( ! [X6,X4,X5] : (s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X5),s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),X4))))) = s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X5),s('type.list.list'(X6),X4))))) | s(bool,'T_0') = s(bool,sK4(X4,X5,X6))) )),
  inference(forward_demodulation,[],[f309,f92])).
fof(f92,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X10,X11] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X10),s(X6,X11)))),
  file('thm.list.FRONT_CONS|split|1','thm.bool.COND_CLAUSES|split|1')).
fof(f309,plain,(
  ( ! [X6,X4,X5] : (s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X5),s('type.list.list'(X6),X4))))) = s('type.list.list'(X6),'const.bool.COND_3'(s(bool,'F_0'),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X5),s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),X4))))))) | s(bool,'T_0') = s(bool,sK4(X4,X5,X6))) )),
  inference(superposition,[],[f97,f71])).
fof(f71,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.FRONT_CONS|split|1',aHL_BOOL_CASES)).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,sK4(X2,X1,X0)),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2)))))))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,sK4(X2,X1,X0))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | p(s(bool,sK4(X2,X1,X0)))) & s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,sK4(X2,X1,X0)),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : ? [X3] : ((~p(s(bool,X3)) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | p(s(bool,X3))) & s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2))))))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : ? [X3] : (((~p(s(bool,X3)) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | p(s(bool,X3)))) & s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2))))))))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2] : ? [X3] : ((p(s(bool,X3)) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2))))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1,X2] : ? [X3] : ((p(s(bool,X3)) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2))))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X15,X12] : ? [X16] : ((p(s(bool,X16)) <=> s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X12)) & s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X15),s('type.list.list'(X6),X12))))) = s('type.list.list'(X6),'const.bool.COND_3'(s(bool,X16),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X15),s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),X12))))))))),
  file('thm.list.FRONT_CONS|split|1','thm.list.FRONT_DEF')).
fof(f67,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK1),s('type.list.list'(sK0),'const.list.FRONT_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))))))) != s('type.list.list'(sK0),'const.list.FRONT_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK1),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))))))),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK1),s('type.list.list'(sK0),'const.list.FRONT_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))))))) != s('type.list.list'(sK0),'const.list.FRONT_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK1),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f44])).
fof(f44,plain,(
  ? [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))) != s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))) = s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0] : ! [X1,X2,X3] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))) = s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X7),s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9))))))) = s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X7),s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9))))))) = s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))))))),
  file('thm.list.FRONT_CONS|split|1','thm.list.FRONT_CONS|split|1')).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK4(X2,X1,X0))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f65])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 42
% Passive clauses: 205
% Generated clauses: 2183
% Final active clauses: 35
% Final passive clauses: 151
% Input formulas: 19
% Initial clauses: 36
% 
% Duplicate literals: 185
% Trivial inequalities: 16
% Fw subsumption resolutions: 52
% Fw demodulations: 96
% Bw demodulations: 16
% Condensations: 248
% 
% Simple tautologies: 50
% Equational tautologies: 58
% Forward subsumptions: 1220
% 
% Unit resulting resolution: 103
% Factoring: 1
% Forward superposition: 520
% Backward superposition: 780
% Self superposition: 64
% Equality factoring: 52
% Equality resolution: 1
% 
% Split clauses: 15
% Split components: 35
% SAT solver clauses: 242
% SAT solver unit clauses: 43
% SAT solver binary clauses: 80
% 
% Sat splits: 15
% Sat splitting refutations: 5
% 
% Memory used [KB]: 2302
% Time elapsed: 0.089 s
% ------------------------------
% ------------------------------
% Success in time 1.407 s
