Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f17110,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14924,f3528,f440,f763,f441,f765,f761,f544,f767,f769,f453,f770,f442,f766,f764,f762,f17056])).
fof(f17056,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f17055,f766])).
fof(f17055,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f16974,f767])).
fof(f16974,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl0),
  inference(superposition,[],[f504,f16956])).
fof(f16956,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(forward_demodulation,[],[f16955,f462])).
fof(f462,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.real.eq_ints|split|2','thm.numeral.numeral_distrib|split|26')).
fof(f16955,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(forward_demodulation,[],[f16954,f569])).
fof(f569,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X7,X6] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6)))))),
  file('thm.real.eq_ints|split|2','thm.real.REAL_LE')).
fof(f16954,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | $spl0),
  inference(forward_demodulation,[],[f16872,f569])).
fof(f16872,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(superposition,[],[f484,f761])).
fof(f484,plain,(
  ( ! [X0] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X13] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13)))))),
  file('thm.real.eq_ints|split|2','thm.real.REAL_NEG_GE|48|')).
fof(f504,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X6))),
  file('thm.real.eq_ints|split|2','thm.arithmetic.LESS_EQ_|48|')).
fof(f762,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | $spl1),
  inference(cnf_transformation,[],[f762_D])).
fof(f762_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f764,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | $spl3),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f766,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl5),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f442,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))))),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) & s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f316])).
fof(f316,plain,(
  ? [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ? [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))))),
  inference(nnf_transformation,[],[f256])).
fof(f256,plain,(
  ? [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) <~> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ~! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.real.eq_ints|split|2','thm.real.eq_ints|split|2')).
fof(f770,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f453,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real.eq_ints|split|2',aHL_FALSITY)).
fof(f769,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(cnf_transformation,[],[f769_D])).
fof(f769_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f767,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f544,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f352])).
fof(f352,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.real.eq_ints|split|2','thm.bool.IMP_CLAUSES|split|2')).
fof(f761,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | $spl0),
  inference(cnf_transformation,[],[f761_D])).
fof(f761_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f765,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl4),
  inference(cnf_transformation,[],[f765_D])).
fof(f765_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f441,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))))),
  inference(cnf_transformation,[],[f317])).
fof(f763,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl2),
  inference(cnf_transformation,[],[f763_D])).
fof(f763_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f440,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))))),
  inference(cnf_transformation,[],[f317])).
fof(f3528,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f447,f3279])).
fof(f3279,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f3278,f763])).
fof(f3278,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f762,f765])).
fof(f447,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f100])).
fof(f100,axiom,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.real.eq_ints|split|2','thm.real.REAL_NEG_|48|')).
fof(f14924,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f14923,f764])).
fof(f14923,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | ($spl0 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f14842,f767])).
fof(f14842,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | ($spl0 | $spl4)),
  inference(superposition,[],[f504,f14804])).
fof(f14804,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',X0)))) ) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f14803,f462])).
fof(f14803,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',X0)))) ) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f14700,f569])).
fof(f14700,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f569,f3746])).
fof(f3746,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f447,f3529])).
fof(f3529,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f761,f765])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 769
% Passive clauses: 3458
% Generated clauses: 16651
% Final active clauses: 696
% Final passive clauses: 2331
% Input formulas: 117
% Initial clauses: 311
% 
% Duplicate literals: 208
% Trivial inequalities: 51
% Fw subsumption resolutions: 1415
% Fw demodulations: 1180
% Bw demodulations: 65
% 
% Simple tautologies: 201
% Equational tautologies: 28
% Forward subsumptions: 10057
% 
% Binary resolution: 4276
% Factoring: 1147
% Forward superposition: 3895
% Backward superposition: 3960
% Self superposition: 18
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 207
% Split components: 417
% SAT solver clauses: 367
% SAT solver unit clauses: 14
% SAT solver binary clauses: 59
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 207
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 365
% TWLsolver calls for satisfiability: 322
% 
% Memory used [KB]: 7931
% Time elapsed: 0.316 s
% ------------------------------
% ------------------------------
% Success in time 0.314 s
