Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.util_prob
% (68991)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 662
% Passive clauses: 10001
% Generated clauses: 18417
% Final active clauses: 656
% Final passive clauses: 8970
% Input formulas: 41
% Initial clauses: 186
% 
% Duplicate literals: 294
% Trivial inequalities: 28
% Fw subsumption resolutions: 1592
% Fw demodulations: 183
% Bw demodulations: 5
% 
% Simple tautologies: 223
% Equational tautologies: 56
% Forward subsumptions: 5704
% 
% Binary resolution: 11799
% Factoring: 1072
% Forward superposition: 1789
% Backward superposition: 1247
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 16
% 
% Split clauses: 350
% Split components: 707
% SAT solver clauses: 712
% SAT solver unit clauses: 27
% SAT solver binary clauses: 188
% 
% Sat splits: 350
% 
% TWLsolver clauses: 712
% TWLsolver calls for satisfiability: 408
% 
% Memory used [KB]: 9083
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (69017)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 416
% Passive clauses: 3201
% Generated clauses: 11999
% Final active clauses: 413
% Final passive clauses: 2785
% Input formulas: 41
% Initial clauses: 186
% 
% Duplicate literals: 204
% Trivial inequalities: 60
% Fw subsumption resolutions: 835
% Fw demodulations: 315
% Bw demodulations: 3
% Condensations: 295
% Global subsumptions: 13
% 
% Simple tautologies: 256
% Equational tautologies: 18
% Forward subsumptions: 3219
% 
% Binary resolution: 4409
% Factoring: 206
% Forward superposition: 1186
% Backward superposition: 1019
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 12
% 
% Split clauses: 42
% Split components: 84
% SAT solver clauses: 3795
% SAT solver unit clauses: 190
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 42
% 
% TWLsolver clauses: 3743
% TWLsolver calls for satisfiability: 1799
% 
% Memory used [KB]: 14328
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (69038)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1122
% Passive clauses: 8198
% Generated clauses: 87704
% Final active clauses: 553
% Final passive clauses: 418
% Input formulas: 41
% Initial clauses: 186
% 
% Duplicate literals: 5032
% Trivial inequalities: 197
% Fw subsumption resolutions: 6279
% Fw demodulations: 9719
% Bw demodulations: 5259
% Condensations: 8810
% Global subsumptions: 1184
% 
% Simple tautologies: 314
% Equational tautologies: 1211
% Forward subsumptions: 42488
% Bw demodulations to eq. taut.: 179
% 
% Unit resulting resolution: 16356
% Factoring: 1
% Forward superposition: 14039
% Backward superposition: 10544
% Self superposition: 257
% Equality factoring: 222
% Equality resolution: 9
% 
% Split clauses: 572
% Split components: 1195
% SAT solver clauses: 14271
% SAT solver unit clauses: 2719
% SAT solver binary clauses: 2812
% 
% Sat splits: 572
% Sat splitting refutations: 3743
% 
% Memory used [KB]: 64604
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% (69348)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2157
% Passive clauses: 44332
% Generated clauses: 191503
% Final active clauses: 2119
% Final passive clauses: 39785
% Input formulas: 41
% Initial clauses: 185
% 
% Selected by SInE selection: 40
% SInE iterations: 1
% 
% Duplicate literals: 6273
% Trivial inequalities: 266
% Fw subsumption resolutions: 21065
% Fw demodulations: 5870
% Bw demodulations: 7
% 
% Simple tautologies: 5929
% Equational tautologies: 307
% Forward subsumptions: 95205
% 
% Binary resolution: 144590
% Factoring: 461
% Forward superposition: 9919
% Backward superposition: 1637
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 33
% 
% Split clauses: 13994
% Split components: 33011
% SAT solver clauses: 15619
% SAT solver unit clauses: 51
% SAT solver binary clauses: 881
% 
% Sat splits: 13994
% Sat splitting refutations: 66
% 
% Memory used [KB]: 58208
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.util_prob
% (69542)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 446
% Passive clauses: 2775
% Generated clauses: 17448
% Final active clauses: 440
% Final passive clauses: 2215
% Input formulas: 41
% Initial clauses: 186
% 
% Duplicate literals: 495
% Trivial inequalities: 223
% Fw subsumption resolutions: 1994
% Fw demodulations: 373
% Bw demodulations: 4
% Global subsumptions: 16
% 
% Simple tautologies: 486
% Equational tautologies: 46
% Forward subsumptions: 10260
% 
% Unit resulting resolution: 4258
% Factoring: 18
% Forward superposition: 5301
% Backward superposition: 4445
% Self superposition: 86
% Equality factoring: 6
% Equality resolution: 22
% 
% Split clauses: 21
% Split components: 43
% SAT solver clauses: 3315
% SAT solver unit clauses: 692
% SAT solver binary clauses: 207
% 
% Sat splits: 21
% Sat splitting refutations: 501
% 
% Memory used [KB]: 9722
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.util_prob
% (69559)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1075
% Passive clauses: 5928
% Generated clauses: 22395
% Final active clauses: 640
% Final passive clauses: 3923
% Input formulas: 41
% Initial clauses: 186
% 
% Duplicate literals: 1102
% Trivial inequalities: 66
% Fw subsumption resolutions: 1603
% Bw subsumption resolutions: 364
% Fw demodulations: 1682
% Bw demodulations: 532
% Condensations: 1791
% Global subsumptions: 5
% 
% Simple tautologies: 1782
% Equational tautologies: 34
% Forward subsumptions: 8186
% Backward subsumptions: 591
% 
% Binary resolution: 8934
% Factoring: 39
% Forward superposition: 2283
% Backward superposition: 2650
% Self superposition: 42
% Equality factoring: 4
% Equality resolution: 13
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 6170
% SAT solver unit clauses: 395
% SAT solver binary clauses: 966
% 
% Sat splits: 1
% Sat splitting refutations: 233
% 
% Memory used [KB]: 10234
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f6196,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f247,f437,f246,f439,f245,f441,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f649,f606,f605,f604,f603,f596,f595,f594,f593,f651,f590,f589,f588,f587,f655,f586,f585,f584,f653,f583,f657,f580,f579,f578,f577,f659,f566,f565,f564,f563,f663,f562,f561,f560,f661,f559,f665,f558,f557,f556,f555,f673,f554,f553,f552,f669,f551,f677,f550,f667,f549,f675,f548,f671,f547,f679,f545,f544,f543,f542,f683,f541,f540,f681,f539,f538,f685,f533,f532,f531,f530,f687,f515,f514,f513,f512,f689,f507,f506,f505,f504,f697,f503,f502,f693,f501,f500,f701,f499,f691,f498,f695,f497,f699,f496,f703,f493,f492,f491,f490,f705,f487,f486,f485,f484,f707,f483,f482,f481,f480,f709,f479,f478,f477,f476,f475,f474,f473,f472,f471,f470,f719,f469,f468,f717,f467,f466,f715,f465,f711,f464,f713,f463,f462,f461,f460,f459,f458,f723,f457,f721,f456,f725,f454,f452,f447,f446,f445,f444,f647,f621,f622,f727,f743,f742,f749,f741,f751,f735,f734,f753,f733,f755,f908,f907,f904,f903,f902,f900,f899,f898,f945,f897,f943,f896,f895,f947,f887,f949,f885,f884,f883,f953,f882,f951,f881,f880,f955,f876,f957,f861,f959,f858,f857,f856,f963,f855,f961,f854,f853,f965,f850,f849,f848,f847,f846,f845,f844,f843,f842,f973,f841,f840,f839,f969,f838,f837,f836,f835,f834,f833,f977,f832,f967,f831,f830,f975,f829,f828,f827,f826,f971,f825,f824,f979,f822,f821,f820,f983,f819,f981,f818,f817,f985,f814,f987,f813,f808,f807,f806,f993,f805,f991,f804,f803,f995,f801,f800,f799,f999,f798,f997,f797,f796,f1001,f795,f794,f793,f1005,f792,f1003,f791,f790,f989,f789,f788,f787,f786,f785,f784,f783,f782,f781,f1015,f780,f1013,f779,f778,f777,f1011,f776,f1007,f775,f774,f1009,f768,f767,f765,f764,f760,f759,f941,f758,f757,f909,f910,f1018,f1017,f1174,f1173,f1170,f1169,f1168,f1166,f1165,f1164,f1212,f1163,f1210,f1162,f1161,f1214,f1153,f1216,f1151,f1150,f1149,f1220,f1148,f1218,f1147,f1146,f1222,f1142,f1224,f1127,f1226,f1124,f1123,f1122,f1230,f1121,f1228,f1120,f1119,f1232,f1116,f1115,f1114,f1113,f1112,f1111,f1110,f1109,f1108,f1240,f1107,f1106,f1105,f1236,f1104,f1103,f1102,f1101,f1100,f1099,f1244,f1098,f1234,f1097,f1096,f1242,f1095,f1094,f1093,f1092,f1238,f1091,f1090,f1246,f1088,f1087,f1086,f1250,f1085,f1248,f1084,f1083,f1252,f1080,f1254,f1079,f1074,f1073,f1072,f1260,f1071,f1258,f1070,f1069,f1262,f1067,f1066,f1065,f1266,f1064,f1264,f1063,f1062,f1268,f1061,f1060,f1059,f1272,f1058,f1270,f1057,f1056,f1256,f1055,f1054,f1053,f1052,f1051,f1050,f1049,f1048,f1047,f1282,f1046,f1280,f1045,f1044,f1043,f1278,f1042,f1274,f1041,f1040,f1276,f1034,f1033,f1031,f1030,f1026,f1025,f1208,f1024,f1023,f1176,f1177,f1285,f1284,f1384,f1385,f1388,f1393,f1392,f1391,f1396,f1405,f1397,f1403,f1398,f1401,f1597,f1596,f1595,f1599,f1609,f1600,f1607,f1601,f1605,f1654,f1660,f1662,f1706,f1712,f1714,f1723,f1729,f1731,f1765,f1767,f1772,f1774,f1779,f1781,f1786,f1790,f1795,f1799,f1828,f1827,f1831,f1826,f1825,f1833,f1840,f1839,f1843,f2099,f2105,f2107,f2176,f2182,f2187,f2193,f2199,f2204,f2227,f2231,f2237,f2241,f2247,f2251,f2272,f2280,f2268,f2283,f2296,f2295,f2394,f2289,f2288,f2398,f2465,f2464,f2567,f2449,f2448,f2578,f2858,f2864,f2867,f2886,f2888,f2894,f2896,f2907,f2909,f2997,f3003,f3010,f3030,f3034,f3042,f3046,f3060,f3059,f3066,f3075,f3074,f3078,f3093,f3092,f3096,f3089,f3088,f3098,f3107,f3106,f3110,f3375,f3381,f3383,f3451,f3457,f3462,f3468,f3474,f3479,f3503,f3507,f3513,f3517,f3523,f3527,f3959,f3971,f3958,f3973,f3957,f3975,f3956,f3977,f4159,f4158,f4266,f4141,f4140,f4279,f4369,f4371,f4503,f1387,f4504,f4550,f4556,f4563,f4588,f4592,f4600,f4604,f4651,f4655,f4660,f4665,f4663,f4667,f4674,f4669,f4672,f4677,f4681,f4686,f4691,f4689,f4694,f4702,f4696,f4700,f4708,f4707,f4712,f4720,f4713,f4718,f4725,f4729,f4733,f4736,f4735,f4739,f4746,f4740,f4744,f4751,f4750,f4755,f4763,f4756,f4761,f5025,f5029,f5035,f5038,f5042,f5047,f5053,f5052,f5057,f5065,f5058,f5063,f5070,f5074,f5078,f6189,f6192,f6191,f6193,f6195])).
fof(f6195,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f6193,plain,(
  ( ! [X0] : (s(sK12,sK17(X0,sK13,sK12)) = s(sK12,happ(s(fun('type.num.num',sK12),sK15(sK14,sK12)),s('type.num.num',sK16(sK17(X0,sK13,sK12),sK14,sK12))))) ) | $spl410),
  inference(cnf_transformation,[],[f6193_D])).
fof(f6193_D,plain,(
  ( ! [X0] : (s(sK12,sK17(X0,sK13,sK12)) = s(sK12,happ(s(fun('type.num.num',sK12),sK15(sK14,sK12)),s('type.num.num',sK16(sK17(X0,sK13,sK12),sK14,sK12))))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f6191,plain,(
  ~p(s(bool,'T_0')) | $spl409),
  inference(cnf_transformation,[],[f6191_D])).
fof(f6191_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl409),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl409])])).
fof(f6192,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(sK12,sK17(X0,sK13,sK12)) = s(sK12,happ(s(fun('type.num.num',sK12),sK15(sK14,sK12)),s('type.num.num',sK16(sK17(X0,sK13,sK12),sK14,sK12))))) ) | ($spl172 | $spl364)),
  inference(forward_demodulation,[],[f6042,f1017])).
fof(f6042,plain,(
  ( ! [X0] : (s(sK12,sK17(X0,sK13,sK12)) = s(sK12,happ(s(fun('type.num.num',sK12),sK15(sK14,sK12)),s('type.num.num',sK16(sK17(X0,sK13,sK12),sK14,sK12)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))))) ) | $spl364),
  inference(resolution,[],[f4504,f291])).
fof(f291,plain,(
  ( ! [X6,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),sK15(X1,X0)),s('type.num.num',sK16(X6,X1,X0)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),sK15(X1,X0)),s('type.num.num',sK16(X6,X1,X0)))))) & (! [X2] : (p(s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,sK17(X2,X1,X0)) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16,sK17])],[f174])).
fof(f174,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ? [X5] : ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X7))))) & (! [X2] : ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,X3) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ? [X2] : ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4))))) & (! [X2] : ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,X3) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X7)))) <=> ? [X19] : ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) => ? [X23] : s(X6,X11) = s(X6,happ(s(fun('type.num.num',X6),X19),s('type.num.num',X23)))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.util_prob.countable_def')).
fof(f6189,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl172 | $spl364)),
  inference(forward_demodulation,[],[f5624,f1017])).
fof(f5624,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) | ($spl1 | $spl364)),
  inference(unit_resulting_resolution,[],[f443,f4504,f291])).
fof(f443,plain,(
  ( ! [X0,X1] : (s(sK12,sK17(X0,sK13,sK12)) != s(sK12,happ(s(fun('type.num.num',sK12),X0),s('type.num.num',X1)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f293])).
fof(f293,plain,(
  ( ! [X4,X2,X0,X1] : (s(X0,sK17(X2,X1,X0)) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f5078,plain,(
  sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl406),
  inference(cnf_transformation,[],[f5078_D])).
fof(f5078_D,plain,(
  sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f5074,plain,(
  sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl360),
  inference(resolution,[],[f4279,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f209])).
fof(f209,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f5070,plain,(
  sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl360),
  inference(unit_resulting_resolution,[],[f4279,f350])).
fof(f5063,plain,(
  sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl402),
  inference(cnf_transformation,[],[f5063_D])).
fof(f5063_D,plain,(
  sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f5058,plain,(
  sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl358),
  inference(resolution,[],[f4266,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f5065,plain,(
  sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl404),
  inference(cnf_transformation,[],[f5065_D])).
fof(f5065_D,plain,(
  sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f5057,plain,(
  sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl358),
  inference(resolution,[],[f4266,f350])).
fof(f5052,plain,(
  sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl358),
  inference(unit_resulting_resolution,[],[f4266,f350])).
fof(f5053,plain,(
  sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl358),
  inference(unit_resulting_resolution,[],[f4266,f351])).
fof(f5047,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl400),
  inference(cnf_transformation,[],[f5047_D])).
fof(f5047_D,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f5042,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl354),
  inference(resolution,[],[f3975,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f5038,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl354),
  inference(unit_resulting_resolution,[],[f3975,f349])).
fof(f5035,plain,(
  sP2('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl398),
  inference(cnf_transformation,[],[f5035_D])).
fof(f5035_D,plain,(
  sP2('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f5029,plain,(
  sP2('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl350),
  inference(resolution,[],[f3971,f349])).
fof(f5025,plain,(
  sP2('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl350),
  inference(unit_resulting_resolution,[],[f3971,f349])).
fof(f4761,plain,(
  sP3('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl394),
  inference(cnf_transformation,[],[f4761_D])).
fof(f4761_D,plain,(
  sP3('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f4756,plain,(
  sP3('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl356),
  inference(resolution,[],[f3977,f350])).
fof(f4763,plain,(
  sP2('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl396),
  inference(cnf_transformation,[],[f4763_D])).
fof(f4763_D,plain,(
  sP2('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f4755,plain,(
  sP2('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl356),
  inference(resolution,[],[f3977,f349])).
fof(f4750,plain,(
  sP2('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl356),
  inference(unit_resulting_resolution,[],[f3977,f349])).
fof(f4751,plain,(
  sP3('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl356),
  inference(unit_resulting_resolution,[],[f3977,f350])).
fof(f4744,plain,(
  sP3('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl390),
  inference(cnf_transformation,[],[f4744_D])).
fof(f4744_D,plain,(
  sP3('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f4740,plain,(
  sP3('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl352),
  inference(resolution,[],[f3973,f350])).
fof(f4746,plain,(
  sP2('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl392),
  inference(cnf_transformation,[],[f4746_D])).
fof(f4746_D,plain,(
  sP2('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f4739,plain,(
  sP2('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl352),
  inference(resolution,[],[f3973,f349])).
fof(f4735,plain,(
  sP2('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl352),
  inference(unit_resulting_resolution,[],[f3973,f349])).
fof(f4736,plain,(
  sP3('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl352),
  inference(unit_resulting_resolution,[],[f3973,f350])).
fof(f4733,plain,(
  sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl388),
  inference(cnf_transformation,[],[f4733_D])).
fof(f4733_D,plain,(
  sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f4729,plain,(
  sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl312),
  inference(resolution,[],[f2578,f350])).
fof(f4725,plain,(
  sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl312),
  inference(unit_resulting_resolution,[],[f2578,f350])).
fof(f4718,plain,(
  sP4('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl384),
  inference(cnf_transformation,[],[f4718_D])).
fof(f4718_D,plain,(
  sP4('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f4713,plain,(
  sP4('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl310),
  inference(resolution,[],[f2567,f351])).
fof(f4720,plain,(
  sP3('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl386),
  inference(cnf_transformation,[],[f4720_D])).
fof(f4720_D,plain,(
  sP3('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f4712,plain,(
  sP3('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl310),
  inference(resolution,[],[f2567,f350])).
fof(f4707,plain,(
  sP3('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl310),
  inference(unit_resulting_resolution,[],[f2567,f350])).
fof(f4708,plain,(
  sP4('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl310),
  inference(unit_resulting_resolution,[],[f2567,f351])).
fof(f4700,plain,(
  sP4('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl380),
  inference(cnf_transformation,[],[f4700_D])).
fof(f4700_D,plain,(
  sP4('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f4696,plain,(
  sP4('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl308),
  inference(resolution,[],[f2398,f351])).
fof(f4702,plain,(
  sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl382),
  inference(cnf_transformation,[],[f4702_D])).
fof(f4702_D,plain,(
  sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f4694,plain,(
  sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl308),
  inference(resolution,[],[f2398,f349])).
fof(f4689,plain,(
  sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl308),
  inference(unit_resulting_resolution,[],[f2398,f349])).
fof(f4691,plain,(
  sP4('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl308),
  inference(unit_resulting_resolution,[],[f2398,f351])).
fof(f4686,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl378),
  inference(cnf_transformation,[],[f4686_D])).
fof(f4686_D,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f4681,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl306),
  inference(resolution,[],[f2394,f349])).
fof(f4677,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl306),
  inference(unit_resulting_resolution,[],[f2394,f349])).
fof(f4672,plain,(
  sP4('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl374),
  inference(cnf_transformation,[],[f4672_D])).
fof(f4672_D,plain,(
  sP4('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f4669,plain,(
  sP4('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl304),
  inference(resolution,[],[f2283,f351])).
fof(f4674,plain,(
  sP2('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl376),
  inference(cnf_transformation,[],[f4674_D])).
fof(f4674_D,plain,(
  sP2('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f4667,plain,(
  sP2('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl304),
  inference(resolution,[],[f2283,f349])).
fof(f4663,plain,(
  sP2('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl304),
  inference(unit_resulting_resolution,[],[f2283,f349])).
fof(f4665,plain,(
  sP4('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl304),
  inference(unit_resulting_resolution,[],[f2283,f351])).
fof(f4660,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl372),
  inference(cnf_transformation,[],[f4660_D])).
fof(f4660_D,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f4655,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl302),
  inference(resolution,[],[f2280,f349])).
fof(f4651,plain,(
  sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl302),
  inference(unit_resulting_resolution,[],[f2280,f349])).
fof(f4604,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl371),
  inference(cnf_transformation,[],[f4604_D])).
fof(f4604_D,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl371),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl371])])).
fof(f4600,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl231),
  inference(unit_resulting_resolution,[],[f1264,f351])).
fof(f4592,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl369),
  inference(cnf_transformation,[],[f4592_D])).
fof(f4592_D,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f4588,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl225),
  inference(unit_resulting_resolution,[],[f1258,f350])).
fof(f4563,plain,(
  sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl366),
  inference(cnf_transformation,[],[f4563_D])).
fof(f4563_D,plain,(
  sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f4556,plain,(
  sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl208),
  inference(resolution,[],[f1242,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f4550,plain,(
  sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl208),
  inference(unit_resulting_resolution,[],[f1242,f372])).
fof(f4504,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | $spl364),
  inference(cnf_transformation,[],[f4504_D])).
fof(f4504_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f1387,plain,(
  p(s(bool,'F_0')) | $spl254),
  inference(cnf_transformation,[],[f1387_D])).
fof(f1387_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f4503,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | ($spl4 | $spl86)),
  inference(forward_demodulation,[],[f4476,f727])).
fof(f4476,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,sK12)),s(fun(sK12,bool),sK14)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))))) ) | $spl4),
  inference(resolution,[],[f1175,f292])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f1175,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))) ) | $spl4),
  inference(resolution,[],[f441,f300])).
fof(f300,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK22(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK22(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.pred_set.SUBSET_DEF')).
fof(f4371,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl363),
  inference(cnf_transformation,[],[f4371_D])).
fof(f4371_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl363),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl363])])).
fof(f4369,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl51),
  inference(unit_resulting_resolution,[],[f691,f349])).
fof(f4279,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl360),
  inference(cnf_transformation,[],[f4279_D])).
fof(f4279_D,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f4140,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl250),
  inference(unit_resulting_resolution,[],[f249,f1284,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f98,f143,f142,f141,f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f141,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f142,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.sat.sat_2|49|_prove_ltb')).
fof(f249,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.COUNTABLE_SUBSET',aHL_TRUTH)).
fof(f4141,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl250),
  inference(unit_resulting_resolution,[],[f248,f1284,f369])).
fof(f248,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.bool.TRUTH')).
fof(f4266,plain,(
  sP5('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl358),
  inference(cnf_transformation,[],[f4266_D])).
fof(f4266_D,plain,(
  sP5('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f4158,plain,(
  sP5('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl250),
  inference(unit_resulting_resolution,[],[f249,f1284,f369])).
fof(f4159,plain,(
  sP5('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl250),
  inference(unit_resulting_resolution,[],[f248,f1284,f369])).
fof(f3977,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl356),
  inference(cnf_transformation,[],[f3977_D])).
fof(f3977_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f3956,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') | ($spl1 | $spl9)),
  inference(unit_resulting_resolution,[],[f649,f632])).
fof(f632,plain,(
  ( ! [X10,X11] : (sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X10,X11) | s(bool,X10) = s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f437,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f3975,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl354),
  inference(cnf_transformation,[],[f3975_D])).
fof(f3975_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f3957,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl7)),
  inference(unit_resulting_resolution,[],[f647,f632])).
fof(f3973,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl352),
  inference(cnf_transformation,[],[f3973_D])).
fof(f3973_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f3958,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | ($spl1 | $spl9)),
  inference(unit_resulting_resolution,[],[f649,f632])).
fof(f3971,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl350),
  inference(cnf_transformation,[],[f3971_D])).
fof(f3971_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f3959,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | ($spl1 | $spl7)),
  inference(unit_resulting_resolution,[],[f647,f632])).
fof(f3527,plain,(
  ~sP5('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl349),
  inference(cnf_transformation,[],[f3527_D])).
fof(f3527_D,plain,(
  ~sP5('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f3523,plain,(
  ~sP5('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl233),
  inference(unit_resulting_resolution,[],[f1266,f351])).
fof(f3517,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl347),
  inference(cnf_transformation,[],[f3517_D])).
fof(f3517_D,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl347),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl347])])).
fof(f3513,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl229),
  inference(unit_resulting_resolution,[],[f1262,f350])).
fof(f3507,plain,(
  ~sP5('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl345),
  inference(cnf_transformation,[],[f3507_D])).
fof(f3507_D,plain,(
  ~sP5('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl345),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl345])])).
fof(f3503,plain,(
  ~sP5('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl227),
  inference(unit_resulting_resolution,[],[f1260,f350])).
fof(f3479,plain,(
  sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl342),
  inference(cnf_transformation,[],[f3479_D])).
fof(f3479_D,plain,(
  sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f3474,plain,(
  sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl212),
  inference(resolution,[],[f1246,f372])).
fof(f3468,plain,(
  sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl212),
  inference(unit_resulting_resolution,[],[f1246,f372])).
fof(f3462,plain,(
  sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl340),
  inference(cnf_transformation,[],[f3462_D])).
fof(f3462_D,plain,(
  sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f3457,plain,(
  sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl210),
  inference(resolution,[],[f1244,f372])).
fof(f3451,plain,(
  sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl210),
  inference(unit_resulting_resolution,[],[f1244,f372])).
fof(f3383,plain,(
  sP10('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl338),
  inference(cnf_transformation,[],[f3383_D])).
fof(f3383_D,plain,(
  sP10('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f3381,plain,(
  sP10('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl182),
  inference(resolution,[],[f1216,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f236])).
fof(f236,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f3375,plain,(
  sP10('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl182),
  inference(unit_resulting_resolution,[],[f1216,f408])).
fof(f3110,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl336),
  inference(cnf_transformation,[],[f3110_D])).
fof(f3110_D,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f3106,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f254,f1068,f494,f1077,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f1077,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f140])).
fof(f494,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f141])).
fof(f1068,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f142])).
fof(f254,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.util_prob.COUNTABLE_SUBSET',aHL_FALSITY)).
fof(f3107,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f254,f488,f494,f1077,f355])).
fof(f488,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f213])).
fof(f3098,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl334),
  inference(cnf_transformation,[],[f3098_D])).
fof(f3098_D,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f3088,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f1068,f1075,f1076,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f1076,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f218])).
fof(f1075,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f3089,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f1068,f1075,f1076,f353])).
fof(f3096,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl332),
  inference(cnf_transformation,[],[f3096_D])).
fof(f3096_D,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f3092,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f254,f488,f1075,f1076,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f3093,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f254,f488,f494,f1076,f354])).
fof(f3078,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl330),
  inference(cnf_transformation,[],[f3078_D])).
fof(f3078_D,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f3074,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f249,f811,f802,f1075,f353])).
fof(f802,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f357])).
fof(f811,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f366])).
fof(f3075,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f248,f811,f802,f1075,f353])).
fof(f3066,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl328),
  inference(cnf_transformation,[],[f3066_D])).
fof(f3066_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f3059,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f249,f810,f809,f1068,f353])).
fof(f809,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f361])).
fof(f810,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f365])).
fof(f3060,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f248,f810,f809,f1068,f353])).
fof(f3046,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl327),
  inference(cnf_transformation,[],[f3046_D])).
fof(f3046_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f3042,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl153),
  inference(unit_resulting_resolution,[],[f997,f351])).
fof(f3034,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl325),
  inference(cnf_transformation,[],[f3034_D])).
fof(f3034_D,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl325),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl325])])).
fof(f3030,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl147),
  inference(unit_resulting_resolution,[],[f991,f350])).
fof(f3010,plain,(
  sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl322),
  inference(cnf_transformation,[],[f3010_D])).
fof(f3010_D,plain,(
  sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f3003,plain,(
  sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl130),
  inference(resolution,[],[f975,f372])).
fof(f2997,plain,(
  sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl130),
  inference(unit_resulting_resolution,[],[f975,f372])).
fof(f2909,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl321),
  inference(cnf_transformation,[],[f2909_D])).
fof(f2909_D,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl321),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl321])])).
fof(f2907,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl59),
  inference(unit_resulting_resolution,[],[f699,f349])).
fof(f2896,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl319),
  inference(cnf_transformation,[],[f2896_D])).
fof(f2896_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl319),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl319])])).
fof(f2894,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl55),
  inference(unit_resulting_resolution,[],[f695,f349])).
fof(f2888,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl317),
  inference(cnf_transformation,[],[f2888_D])).
fof(f2888_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl317),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl317])])).
fof(f2886,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl53),
  inference(unit_resulting_resolution,[],[f693,f349])).
fof(f2867,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl314),
  inference(cnf_transformation,[],[f2867_D])).
fof(f2867_D,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f2864,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl34),
  inference(resolution,[],[f675,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f2858,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl34),
  inference(unit_resulting_resolution,[],[f675,f391])).
fof(f2578,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl312),
  inference(cnf_transformation,[],[f2578_D])).
fof(f2578_D,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f2448,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1017,f369])).
fof(f2449,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f248,f1017,f369])).
fof(f2567,plain,(
  sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl310),
  inference(cnf_transformation,[],[f2567_D])).
fof(f2567_D,plain,(
  sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f2464,plain,(
  sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1017,f369])).
fof(f2465,plain,(
  sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f248,f1017,f369])).
fof(f2398,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl308),
  inference(cnf_transformation,[],[f2398_D])).
fof(f2398_D,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f2288,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl86),
  inference(unit_resulting_resolution,[],[f249,f727,f369])).
fof(f2289,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl86),
  inference(unit_resulting_resolution,[],[f248,f727,f369])).
fof(f2394,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl306),
  inference(cnf_transformation,[],[f2394_D])).
fof(f2394_D,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f2295,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl86),
  inference(unit_resulting_resolution,[],[f249,f727,f369])).
fof(f2296,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl86),
  inference(unit_resulting_resolution,[],[f248,f727,f369])).
fof(f2283,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl304),
  inference(cnf_transformation,[],[f2283_D])).
fof(f2283_D,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f2268,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f254,f647,f368])).
fof(f2280,plain,(
  sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl302),
  inference(cnf_transformation,[],[f2280_D])).
fof(f2280_D,plain,(
  sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f2272,plain,(
  sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl7),
  inference(unit_resulting_resolution,[],[f254,f647,f368])).
fof(f2251,plain,(
  ~sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl301),
  inference(cnf_transformation,[],[f2251_D])).
fof(f2251_D,plain,(
  ~sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl301),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl301])])).
fof(f2247,plain,(
  ~sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl155),
  inference(unit_resulting_resolution,[],[f999,f351])).
fof(f2241,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl299),
  inference(cnf_transformation,[],[f2241_D])).
fof(f2241_D,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl299),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl299])])).
fof(f2237,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl151),
  inference(unit_resulting_resolution,[],[f995,f350])).
fof(f2231,plain,(
  ~sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl297),
  inference(cnf_transformation,[],[f2231_D])).
fof(f2231_D,plain,(
  ~sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl297),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl297])])).
fof(f2227,plain,(
  ~sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl149),
  inference(unit_resulting_resolution,[],[f993,f350])).
fof(f2204,plain,(
  sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl294),
  inference(cnf_transformation,[],[f2204_D])).
fof(f2204_D,plain,(
  sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f2199,plain,(
  sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl134),
  inference(resolution,[],[f979,f372])).
fof(f2193,plain,(
  sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl134),
  inference(unit_resulting_resolution,[],[f979,f372])).
fof(f2187,plain,(
  sP6('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl292),
  inference(cnf_transformation,[],[f2187_D])).
fof(f2187_D,plain,(
  sP6('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f2182,plain,(
  sP6('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl132),
  inference(resolution,[],[f977,f372])).
fof(f2176,plain,(
  sP6('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl132),
  inference(unit_resulting_resolution,[],[f977,f372])).
fof(f2107,plain,(
  sP10('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl290),
  inference(cnf_transformation,[],[f2107_D])).
fof(f2107_D,plain,(
  sP10('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f2105,plain,(
  sP10('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl104),
  inference(resolution,[],[f949,f408])).
fof(f2099,plain,(
  sP10('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl104),
  inference(unit_resulting_resolution,[],[f949,f408])).
fof(f1843,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl288),
  inference(cnf_transformation,[],[f1843_D])).
fof(f1843_D,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f1839,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f254,f802,f494,f811,f355])).
fof(f1840,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f254,f488,f494,f811,f355])).
fof(f1833,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl286),
  inference(cnf_transformation,[],[f1833_D])).
fof(f1833_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f1825,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f802,f809,f810,f353])).
fof(f1826,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f802,f809,f810,f353])).
fof(f1831,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl284),
  inference(cnf_transformation,[],[f1831_D])).
fof(f1831_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f1827,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f254,f488,f809,f810,f354])).
fof(f1828,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f254,f488,f494,f810,f354])).
fof(f1799,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl283),
  inference(cnf_transformation,[],[f1799_D])).
fof(f1799_D,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f1795,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl67),
  inference(unit_resulting_resolution,[],[f707,f351])).
fof(f1790,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl281),
  inference(cnf_transformation,[],[f1790_D])).
fof(f1790_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') <=> ~$spl281),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl281])])).
fof(f1786,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f705,f350])).
fof(f1781,plain,(
  ~sP5('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl279),
  inference(cnf_transformation,[],[f1781_D])).
fof(f1781_D,plain,(
  ~sP5('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl279),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl279])])).
fof(f1779,plain,(
  ~sP5('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl63),
  inference(unit_resulting_resolution,[],[f703,f349])).
fof(f1774,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl277),
  inference(cnf_transformation,[],[f1774_D])).
fof(f1774_D,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f1772,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f701,f349])).
fof(f1767,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl275),
  inference(cnf_transformation,[],[f1767_D])).
fof(f1767_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f1765,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl57),
  inference(unit_resulting_resolution,[],[f697,f349])).
fof(f1731,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl272),
  inference(cnf_transformation,[],[f1731_D])).
fof(f1731_D,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f1729,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl38),
  inference(resolution,[],[f679,f391])).
fof(f1723,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl38),
  inference(unit_resulting_resolution,[],[f679,f391])).
fof(f1714,plain,(
  sP8('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl270),
  inference(cnf_transformation,[],[f1714_D])).
fof(f1714_D,plain,(
  sP8('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f1712,plain,(
  sP8('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl36),
  inference(resolution,[],[f677,f391])).
fof(f1706,plain,(
  sP8('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl36),
  inference(unit_resulting_resolution,[],[f677,f391])).
fof(f1662,plain,(
  sP10('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl268),
  inference(cnf_transformation,[],[f1662_D])).
fof(f1662_D,plain,(
  sP10('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f1660,plain,(
  sP10('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl16),
  inference(resolution,[],[f657,f408])).
fof(f1654,plain,(
  sP10('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f657,f408])).
fof(f1605,plain,(
  sP4('F_0','T_0','F_0') | $spl262),
  inference(cnf_transformation,[],[f1605_D])).
fof(f1605_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f1601,plain,(
  sP4('F_0','T_0','F_0') | $spl94),
  inference(resolution,[],[f755,f351])).
fof(f1607,plain,(
  sP3('F_0','T_0','F_0') | $spl264),
  inference(cnf_transformation,[],[f1607_D])).
fof(f1607_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f1600,plain,(
  sP3('F_0','T_0','F_0') | $spl94),
  inference(resolution,[],[f755,f350])).
fof(f1609,plain,(
  sP2('T_0','F_0','F_0') | $spl266),
  inference(cnf_transformation,[],[f1609_D])).
fof(f1609_D,plain,(
  sP2('T_0','F_0','F_0') <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f1599,plain,(
  sP2('T_0','F_0','F_0') | $spl94),
  inference(resolution,[],[f755,f349])).
fof(f1595,plain,(
  sP2('T_0','F_0','F_0') | $spl94),
  inference(unit_resulting_resolution,[],[f755,f349])).
fof(f1596,plain,(
  sP3('F_0','T_0','F_0') | $spl94),
  inference(unit_resulting_resolution,[],[f755,f350])).
fof(f1597,plain,(
  sP4('F_0','T_0','F_0') | $spl94),
  inference(unit_resulting_resolution,[],[f755,f351])).
fof(f1401,plain,(
  sP4('F_0','F_0','T_0') | $spl256),
  inference(cnf_transformation,[],[f1401_D])).
fof(f1401_D,plain,(
  sP4('F_0','F_0','T_0') <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f1398,plain,(
  sP4('F_0','F_0','T_0') | $spl90),
  inference(resolution,[],[f751,f351])).
fof(f1403,plain,(
  sP3('T_0','F_0','F_0') | $spl258),
  inference(cnf_transformation,[],[f1403_D])).
fof(f1403_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f1397,plain,(
  sP3('T_0','F_0','F_0') | $spl90),
  inference(resolution,[],[f751,f350])).
fof(f1405,plain,(
  sP2('F_0','T_0','F_0') | $spl260),
  inference(cnf_transformation,[],[f1405_D])).
fof(f1405_D,plain,(
  sP2('F_0','T_0','F_0') <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f1396,plain,(
  sP2('F_0','T_0','F_0') | $spl90),
  inference(resolution,[],[f751,f349])).
fof(f1391,plain,(
  sP2('F_0','T_0','F_0') | $spl90),
  inference(unit_resulting_resolution,[],[f751,f349])).
fof(f1392,plain,(
  sP3('T_0','F_0','F_0') | $spl90),
  inference(unit_resulting_resolution,[],[f751,f350])).
fof(f1393,plain,(
  sP4('F_0','F_0','T_0') | $spl90),
  inference(unit_resulting_resolution,[],[f751,f351])).
fof(f1388,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~p(s(bool,'F_0')) | p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))))),
  introduced(theory_axiom,[])).
fof(f1385,plain,(
  ( ! [X6,X7] : (sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X6,X7) | ~sP2(X7,X6,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~sP3(X6,X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) ) | $spl252),
  inference(cnf_transformation,[],[f1385_D])).
fof(f1385_D,plain,(
  ( ! [X6,X7] : (sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X6,X7) | ~sP2(X7,X6,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~sP3(X6,X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f1384,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X6,X7) | ~sP3(X6,X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~sP2(X7,X6,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) ) | ($spl1 | $spl86)),
  inference(forward_demodulation,[],[f1381,f727])).
fof(f1381,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) | sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X6,X7) | ~sP3(X6,X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~sP2(X7,X6,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) ) | $spl1),
  inference(resolution,[],[f489,f355])).
fof(f489,plain,(
  ( ! [X0,X1] : (sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f1284,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl250),
  inference(cnf_transformation,[],[f1284_D])).
fof(f1284_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f1285,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(global_subsumption,[],[f247,f246,f245,f441,f1174,f1173,f1170,f1169,f1168,f1167,f1166,f1165,f1164,f1163,f1162,f1161,f1160,f1159,f1156,f1154,f1153,f1152,f1151,f1150,f1149,f1148,f1147,f1146,f1144,f1143,f1142,f1141,f1140,f1137,f1135,f1134,f1131,f1129,f1128,f1127,f1124,f1123,f1122,f1121,f1120,f1119,f1118,f1117,f1116,f1115,f1114,f1113,f1112,f1111,f1110,f1109,f1108,f1107,f1106,f1105,f1104,f1103,f1102,f1101,f1100,f1099,f1098,f1097,f1096,f1095,f1094,f1093,f1092,f1091,f1090,f1089,f1088,f1087,f1086,f1085,f1084,f1083,f1078,f1077,f1076,f1075,f1074,f1073,f1072,f1071,f1070,f1069,f1068,f1067,f1066,f1065,f1064,f1063,f1062,f1061,f1060,f1059,f1058,f1057,f1056,f1055,f1054,f1053,f1052,f1051,f1050,f1049,f1048,f1047,f1046,f1045,f1044,f1043,f1042,f1041,f1040,f1039,f1038,f1034,f1033,f1031,f1030,f1026,f1025,f1024,f1023,f1175,f1176,f1177,f1180,f1185,f1186,f1187,f1190,f1197,f1200,f1201,f1203,f1204,f1205,f1206])).
fof(f1206,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(superposition,[],[f441,f255])).
fof(f255,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.util_prob.COUNTABLE_SUBSET',aHL_BOOL_CASES)).
fof(f1205,plain,(
  ( ! [X41,X40] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X40,X41) | ~p(s(bool,X40))) ) | $spl4),
  inference(resolution,[],[f441,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f133,f152,f151])).
fof(f151,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.sat.sat_27_prove_ltb')).
fof(f1204,plain,(
  ( ! [X39,X38] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X39,X38) | p(s(bool,X38))) ) | $spl4),
  inference(resolution,[],[f441,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f1203,plain,(
  ( ! [X37,X36] : (sP11(X37,X36,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X36)) | p(s(bool,X37))) ) | $spl4),
  inference(resolution,[],[f441,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f242])).
fof(f1201,plain,(
  ( ! [X33,X32] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X32,X33) | ~p(s(bool,X32))) ) | $spl4),
  inference(resolution,[],[f441,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f102,f149,f148])).
fof(f148,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.sat.sat_25_prove_ltb')).
fof(f1200,plain,(
  ( ! [X30,X31] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X31,X30) | ~p(s(bool,X30))) ) | $spl4),
  inference(resolution,[],[f441,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f1197,plain,(
  ( ! [X24,X25] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X24,X25) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f441,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f100,f146,f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) & p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.sat.sat_23_prove_ltb')).
fof(f1190,plain,(
  ( ! [X11] : (~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X11) | ~p(s(bool,X11))) ) | $spl4),
  inference(resolution,[],[f441,f338])).
fof(f338,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f91,f138])).
fof(f138,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f91,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X21,X22] : ((p(s(bool,X22)) <=> ~p(s(bool,X21))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.sat.sat_29_prove_ltb')).
fof(f1187,plain,(
  ( ! [X8] : (sP1(X8,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X8))) ) | $spl4),
  inference(resolution,[],[f441,f335])).
fof(f335,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f199])).
fof(f199,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f138])).
fof(f1186,plain,(
  ( ! [X7] : (sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X7) | p(s(bool,X7))) ) | $spl4),
  inference(resolution,[],[f441,f332])).
fof(f332,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f1185,plain,(
  ( ! [X6] : (~sP1(X6,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X6))) ) | $spl4),
  inference(resolution,[],[f441,f331])).
fof(f331,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f1180,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f441,f310])).
fof(f310,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.bool.IMP_ANTISYM_AX')).
fof(f1038,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f321])).
fof(f321,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f127,f136])).
fof(f136,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f127,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X15,X6] : ((~(p(s(bool,X6)) | p(s(bool,X15))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X15)) => p(s(bool,'F_0')))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.sat.sat_9_prove_ltb')).
fof(f1039,plain,(
  ( ! [X0] : (sP0('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f322])).
fof(f322,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f1078,plain,(
  ( ! [X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f218])).
fof(f1089,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f145])).
fof(f1117,plain,(
  ( ! [X0] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f226])).
fof(f1118,plain,(
  ( ! [X0] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f1128,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f231])).
fof(f231,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f148])).
fof(f1129,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f1131,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f401])).
fof(f1134,plain,(
  ( ! [X0] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f401])).
fof(f1135,plain,(
  ( ! [X0] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f401])).
fof(f1137,plain,(
  ( ! [X0] : (sP9('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f402])).
fof(f1140,plain,(
  ( ! [X0] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f402])).
fof(f1141,plain,(
  ( ! [X0] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f402])).
fof(f1143,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f1144,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f1152,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f239])).
fof(f239,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f1154,plain,(
  ( ! [X0] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f418])).
fof(f1156,plain,(
  ( ! [X0] : (sP11('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f419])).
fof(f1159,plain,(
  ( ! [X0] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f419])).
fof(f1160,plain,(
  ( ! [X0] : (sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f419])).
fof(f1167,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f242])).
fof(f1177,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(resolution,[],[f441,f267])).
fof(f267,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.bool.EQ_CLAUSES|split|0')).
fof(f1176,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(resolution,[],[f441,f265])).
fof(f265,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.bool.EQ_CLAUSES|split|1')).
fof(f1023,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f265])).
fof(f1024,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f267])).
fof(f1208,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl175),
  inference(cnf_transformation,[],[f1208_D])).
fof(f1208_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f1025,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f268])).
fof(f268,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.bool.EQ_CLAUSES|split|3')).
fof(f1026,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f270])).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.bool.EQ_CLAUSES|split|2')).
fof(f1030,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f310])).
fof(f1031,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f310])).
fof(f1033,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f310])).
fof(f1034,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f310])).
fof(f1276,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl243),
  inference(cnf_transformation,[],[f1276_D])).
fof(f1276_D,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f1040,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f331])).
fof(f1041,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f331])).
fof(f1274,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl241),
  inference(cnf_transformation,[],[f1274_D])).
fof(f1274_D,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl241),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl241])])).
fof(f1042,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f331])).
fof(f1278,plain,(
  ~sP1('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl245),
  inference(cnf_transformation,[],[f1278_D])).
fof(f1278_D,plain,(
  ~sP1('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f1043,plain,(
  ~sP1('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f331])).
fof(f1044,plain,(
  ~sP1('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f331])).
fof(f1045,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f331])).
fof(f1280,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl246),
  inference(cnf_transformation,[],[f1280_D])).
fof(f1280_D,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f1046,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f332])).
fof(f1282,plain,(
  sP1('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl248),
  inference(cnf_transformation,[],[f1282_D])).
fof(f1282_D,plain,(
  sP1('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f1047,plain,(
  sP1('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f335])).
fof(f1048,plain,(
  sP1('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f336])).
fof(f336,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f201])).
fof(f1049,plain,(
  sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f337])).
fof(f337,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f1050,plain,(
  ~sP1('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f338])).
fof(f1051,plain,(
  ~sP1('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f338])).
fof(f1052,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f338])).
fof(f1053,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f338])).
fof(f1054,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f338])).
fof(f1055,plain,(
  ~sP1('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f338])).
fof(f1256,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl223),
  inference(cnf_transformation,[],[f1256_D])).
fof(f1256_D,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f1056,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f210])).
fof(f1057,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f352])).
fof(f1270,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl237),
  inference(cnf_transformation,[],[f1270_D])).
fof(f1270_D,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f1058,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f352])).
fof(f1272,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl239),
  inference(cnf_transformation,[],[f1272_D])).
fof(f1272_D,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f1059,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f352])).
fof(f1060,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f352])).
fof(f1061,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f352])).
fof(f1268,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl235),
  inference(cnf_transformation,[],[f1268_D])).
fof(f1268_D,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f1062,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f1063,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f356])).
fof(f1264,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl231),
  inference(cnf_transformation,[],[f1264_D])).
fof(f1264_D,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f1064,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f356])).
fof(f1266,plain,(
  ~sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl233),
  inference(cnf_transformation,[],[f1266_D])).
fof(f1266_D,plain,(
  ~sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f1065,plain,(
  ~sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f356])).
fof(f1066,plain,(
  ~sP4('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f356])).
fof(f1067,plain,(
  ~sP4('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f356])).
fof(f1262,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl229),
  inference(cnf_transformation,[],[f1262_D])).
fof(f1262_D,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') <=> ~$spl229),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl229])])).
fof(f1069,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f1070,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f360])).
fof(f1258,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl225),
  inference(cnf_transformation,[],[f1258_D])).
fof(f1258_D,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f1071,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f360])).
fof(f1260,plain,(
  ~sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl227),
  inference(cnf_transformation,[],[f1260_D])).
fof(f1260_D,plain,(
  ~sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f1072,plain,(
  ~sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f360])).
fof(f1073,plain,(
  ~sP3('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f360])).
fof(f1074,plain,(
  ~sP3('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f360])).
fof(f1079,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | ($spl4 | $spl9)),
  inference(unit_resulting_resolution,[],[f649,f441,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f1254,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl221),
  inference(cnf_transformation,[],[f1254_D])).
fof(f1254_D,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f1080,plain,(
  ~sP5('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl4 | $spl9)),
  inference(unit_resulting_resolution,[],[f649,f441,f370])).
fof(f1252,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl219),
  inference(cnf_transformation,[],[f1252_D])).
fof(f1252_D,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f1083,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f224])).
fof(f1084,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f379])).
fof(f1248,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl215),
  inference(cnf_transformation,[],[f1248_D])).
fof(f1248_D,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f1085,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f379])).
fof(f1250,plain,(
  ~sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl217),
  inference(cnf_transformation,[],[f1250_D])).
fof(f1250_D,plain,(
  ~sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f1086,plain,(
  ~sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f379])).
fof(f1087,plain,(
  ~sP6('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f379])).
fof(f1088,plain,(
  ~sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f379])).
fof(f1246,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl212),
  inference(cnf_transformation,[],[f1246_D])).
fof(f1246_D,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f1090,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f249,f441,f385])).
fof(f1091,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f249,f441,f385])).
fof(f1238,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl204),
  inference(cnf_transformation,[],[f1238_D])).
fof(f1238_D,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f1092,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f249,f441,f385])).
fof(f1093,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f248,f441,f385])).
fof(f1094,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f248,f441,f385])).
fof(f1095,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f248,f441,f385])).
fof(f1242,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl208),
  inference(cnf_transformation,[],[f1242_D])).
fof(f1242_D,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f1096,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f441,f385])).
fof(f1097,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f441,f385])).
fof(f1234,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl200),
  inference(cnf_transformation,[],[f1234_D])).
fof(f1234_D,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f1098,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f441,f385])).
fof(f1244,plain,(
  sP7('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl210),
  inference(cnf_transformation,[],[f1244_D])).
fof(f1244_D,plain,(
  sP7('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f1099,plain,(
  sP7('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f249,f441,f385])).
fof(f1100,plain,(
  sP7('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f249,f441,f385])).
fof(f1101,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f249,f441,f385])).
fof(f1102,plain,(
  sP7('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f248,f441,f385])).
fof(f1103,plain,(
  sP7('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f248,f441,f385])).
fof(f1104,plain,(
  sP7('T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f248,f441,f385])).
fof(f1236,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl202),
  inference(cnf_transformation,[],[f1236_D])).
fof(f1236_D,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f1105,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f441,f385])).
fof(f1106,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f441,f385])).
fof(f1107,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f441,f385])).
fof(f1240,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl206),
  inference(cnf_transformation,[],[f1240_D])).
fof(f1240_D,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f1108,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f249,f441,f385])).
fof(f1109,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f249,f441,f385])).
fof(f1110,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f249,f441,f385])).
fof(f1111,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f248,f441,f385])).
fof(f1112,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f248,f248,f441,f385])).
fof(f1113,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f441,f248,f441,f385])).
fof(f1114,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f441,f441,f385])).
fof(f1115,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f441,f441,f385])).
fof(f1116,plain,(
  sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f441,f441,f385])).
fof(f1232,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl199),
  inference(cnf_transformation,[],[f1232_D])).
fof(f1232_D,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f1119,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f249,f441,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f226])).
fof(f1120,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f248,f441,f388])).
fof(f1228,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl195),
  inference(cnf_transformation,[],[f1228_D])).
fof(f1228_D,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f1121,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f441,f388])).
fof(f1230,plain,(
  ~sP7('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl197),
  inference(cnf_transformation,[],[f1230_D])).
fof(f1230_D,plain,(
  ~sP7('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f1122,plain,(
  ~sP7('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f249,f441,f388])).
fof(f1123,plain,(
  ~sP7('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f248,f441,f388])).
fof(f1124,plain,(
  ~sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f441,f388])).
fof(f1226,plain,(
  ~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl193),
  inference(cnf_transformation,[],[f1226_D])).
fof(f1226_D,plain,(
  ~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') <=> ~$spl193),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl193])])).
fof(f1127,plain,(
  ~sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f254,f441,f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f1224,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl191),
  inference(cnf_transformation,[],[f1224_D])).
fof(f1224_D,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') <=> ~$spl191),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl191])])).
fof(f1142,plain,(
  ~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f254,f441,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f1222,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl189),
  inference(cnf_transformation,[],[f1222_D])).
fof(f1222_D,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f1146,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f249,f441,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f240])).
fof(f1147,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f248,f441,f413])).
fof(f1218,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl185),
  inference(cnf_transformation,[],[f1218_D])).
fof(f1218_D,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f1148,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f441,f413])).
fof(f1220,plain,(
  ~sP10('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl187),
  inference(cnf_transformation,[],[f1220_D])).
fof(f1220_D,plain,(
  ~sP10('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f1149,plain,(
  ~sP10('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f254,f441,f413])).
fof(f1150,plain,(
  ~sP10('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f248,f254,f441,f413])).
fof(f1151,plain,(
  ~sP10('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f441,f254,f441,f413])).
fof(f1216,plain,(
  sP11('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl182),
  inference(cnf_transformation,[],[f1216_D])).
fof(f1216_D,plain,(
  sP11('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f1153,plain,(
  sP11('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f254,f441,f417])).
fof(f1214,plain,(
  ~sP11('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl181),
  inference(cnf_transformation,[],[f1214_D])).
fof(f1214_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f1161,plain,(
  ~sP11('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f249,f441,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f1162,plain,(
  ~sP11('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f248,f441,f420])).
fof(f1210,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl177),
  inference(cnf_transformation,[],[f1210_D])).
fof(f1210_D,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f1163,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f441,f420])).
fof(f1212,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl179),
  inference(cnf_transformation,[],[f1212_D])).
fof(f1212_D,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f1164,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f249,f441,f420])).
fof(f1165,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f254,f248,f441,f420])).
fof(f1166,plain,(
  ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f441,f441,f420])).
fof(f1168,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f254,f249,f441,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X11,X17,X12,X18] : ((s(bool,X11) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X12) = s(bool,X18))) => ((p(s(bool,X11)) => p(s(bool,X12))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.bool.IMP_CONG')).
fof(f1169,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f254,f248,f441,f429])).
fof(f1170,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f254,f441,f441,f429])).
fof(f1173,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f254,f441,f441,f429])).
fof(f1174,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f254,f254,f441,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f1017,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl172),
  inference(cnf_transformation,[],[f1017_D])).
fof(f1017_D,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f1018,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(global_subsumption,[],[f247,f246,f245,f439,f908,f907,f904,f903,f902,f901,f900,f899,f898,f897,f896,f895,f894,f893,f890,f888,f887,f886,f885,f884,f883,f882,f881,f880,f878,f877,f876,f875,f874,f871,f869,f868,f865,f863,f862,f861,f858,f857,f856,f855,f854,f853,f852,f851,f850,f849,f848,f847,f846,f845,f844,f843,f842,f841,f840,f839,f838,f837,f836,f835,f834,f833,f832,f831,f830,f829,f828,f827,f826,f825,f824,f823,f822,f821,f820,f819,f818,f817,f812,f811,f810,f809,f808,f807,f806,f805,f804,f803,f802,f801,f800,f799,f798,f797,f796,f795,f794,f793,f792,f791,f790,f789,f788,f787,f786,f785,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f773,f772,f768,f767,f765,f764,f760,f759,f758,f757,f909,f910,f913,f918,f919,f920,f923,f930,f933,f934,f936,f937,f938,f939])).
fof(f939,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(superposition,[],[f439,f255])).
fof(f938,plain,(
  ( ! [X39,X40] : (sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X39,X40) | ~p(s(bool,X39))) ) | $spl2),
  inference(resolution,[],[f439,f419])).
fof(f937,plain,(
  ( ! [X37,X38] : (sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X38,X37) | p(s(bool,X37))) ) | $spl2),
  inference(resolution,[],[f439,f418])).
fof(f936,plain,(
  ( ! [X35,X36] : (sP11(X36,X35,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl2),
  inference(resolution,[],[f439,f417])).
fof(f934,plain,(
  ( ! [X31,X32] : (sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X31,X32) | ~p(s(bool,X31))) ) | $spl2),
  inference(resolution,[],[f439,f402])).
fof(f933,plain,(
  ( ! [X30,X29] : (sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X30,X29) | ~p(s(bool,X29))) ) | $spl2),
  inference(resolution,[],[f439,f401])).
fof(f930,plain,(
  ( ! [X24,X23] : (sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X23,X24) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl2),
  inference(resolution,[],[f439,f385])).
fof(f923,plain,(
  ( ! [X10] : (~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X10) | ~p(s(bool,X10))) ) | $spl2),
  inference(resolution,[],[f439,f338])).
fof(f920,plain,(
  ( ! [X7] : (sP1(X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | p(s(bool,X7))) ) | $spl2),
  inference(resolution,[],[f439,f335])).
fof(f919,plain,(
  ( ! [X6] : (sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X6) | p(s(bool,X6))) ) | $spl2),
  inference(resolution,[],[f439,f332])).
fof(f918,plain,(
  ( ! [X5] : (~sP1(X5,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | ~p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f439,f331])).
fof(f913,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(resolution,[],[f439,f310])).
fof(f772,plain,(
  ( ! [X0] : (sP0(X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f321])).
fof(f773,plain,(
  ( ! [X0] : (sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f322])).
fof(f812,plain,(
  ( ! [X0,X1] : (sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f367])).
fof(f823,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f380])).
fof(f851,plain,(
  ( ! [X0] : (~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f386])).
fof(f852,plain,(
  ( ! [X0] : (~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f387])).
fof(f862,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f397])).
fof(f863,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f398])).
fof(f865,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f401])).
fof(f868,plain,(
  ( ! [X0] : (sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f401])).
fof(f869,plain,(
  ( ! [X0] : (sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f401])).
fof(f871,plain,(
  ( ! [X0] : (sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f402])).
fof(f874,plain,(
  ( ! [X0] : (sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f402])).
fof(f875,plain,(
  ( ! [X0] : (sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f402])).
fof(f877,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f404])).
fof(f878,plain,(
  ( ! [X0] : (~sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f405])).
fof(f886,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f415])).
fof(f888,plain,(
  ( ! [X0] : (sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f418])).
fof(f890,plain,(
  ( ! [X0] : (sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f419])).
fof(f893,plain,(
  ( ! [X0] : (sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f419])).
fof(f894,plain,(
  ( ! [X0] : (sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f419])).
fof(f901,plain,(
  ( ! [X0] : (~sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f422])).
fof(f910,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(resolution,[],[f439,f267])).
fof(f909,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(resolution,[],[f439,f265])).
fof(f757,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f265])).
fof(f758,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f267])).
fof(f941,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl97),
  inference(cnf_transformation,[],[f941_D])).
fof(f941_D,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f759,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f268])).
fof(f760,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f270])).
fof(f764,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f310])).
fof(f765,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f310])).
fof(f767,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f310])).
fof(f768,plain,(
  s(bool,'T_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f310])).
fof(f1009,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl165),
  inference(cnf_transformation,[],[f1009_D])).
fof(f1009_D,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f774,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f331])).
fof(f775,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f331])).
fof(f1007,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl163),
  inference(cnf_transformation,[],[f1007_D])).
fof(f1007_D,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f776,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f331])).
fof(f1011,plain,(
  ~sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl167),
  inference(cnf_transformation,[],[f1011_D])).
fof(f1011_D,plain,(
  ~sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f777,plain,(
  ~sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f331])).
fof(f778,plain,(
  ~sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f331])).
fof(f779,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f331])).
fof(f1013,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl168),
  inference(cnf_transformation,[],[f1013_D])).
fof(f1013_D,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f780,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f332])).
fof(f1015,plain,(
  sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl170),
  inference(cnf_transformation,[],[f1015_D])).
fof(f1015_D,plain,(
  sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f781,plain,(
  sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f335])).
fof(f782,plain,(
  sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f336])).
fof(f783,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f337])).
fof(f784,plain,(
  ~sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f338])).
fof(f785,plain,(
  ~sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f338])).
fof(f786,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f338])).
fof(f787,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f338])).
fof(f788,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f338])).
fof(f789,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f338])).
fof(f989,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl145),
  inference(cnf_transformation,[],[f989_D])).
fof(f989_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f790,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f352])).
fof(f791,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f352])).
fof(f1003,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl159),
  inference(cnf_transformation,[],[f1003_D])).
fof(f1003_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f792,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f352])).
fof(f1005,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl161),
  inference(cnf_transformation,[],[f1005_D])).
fof(f1005_D,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f793,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f352])).
fof(f794,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f352])).
fof(f795,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f352])).
fof(f1001,plain,(
  ~sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl157),
  inference(cnf_transformation,[],[f1001_D])).
fof(f1001_D,plain,(
  ~sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f796,plain,(
  ~sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f356])).
fof(f797,plain,(
  ~sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f356])).
fof(f997,plain,(
  ~sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl153),
  inference(cnf_transformation,[],[f997_D])).
fof(f997_D,plain,(
  ~sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f798,plain,(
  ~sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f356])).
fof(f999,plain,(
  ~sP4('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl155),
  inference(cnf_transformation,[],[f999_D])).
fof(f999_D,plain,(
  ~sP4('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f799,plain,(
  ~sP4('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f356])).
fof(f800,plain,(
  ~sP4('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f356])).
fof(f801,plain,(
  ~sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f356])).
fof(f995,plain,(
  ~sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl151),
  inference(cnf_transformation,[],[f995_D])).
fof(f995_D,plain,(
  ~sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f803,plain,(
  ~sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f360])).
fof(f804,plain,(
  ~sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f360])).
fof(f991,plain,(
  ~sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl147),
  inference(cnf_transformation,[],[f991_D])).
fof(f991_D,plain,(
  ~sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f805,plain,(
  ~sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f360])).
fof(f993,plain,(
  ~sP3('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl149),
  inference(cnf_transformation,[],[f993_D])).
fof(f993_D,plain,(
  ~sP3('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f806,plain,(
  ~sP3('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f360])).
fof(f807,plain,(
  ~sP3('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f360])).
fof(f808,plain,(
  ~sP3('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f360])).
fof(f813,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | ($spl2 | $spl9)),
  inference(unit_resulting_resolution,[],[f649,f439,f370])).
fof(f987,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl143),
  inference(cnf_transformation,[],[f987_D])).
fof(f987_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f814,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl2 | $spl9)),
  inference(unit_resulting_resolution,[],[f649,f439,f370])).
fof(f985,plain,(
  ~sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl141),
  inference(cnf_transformation,[],[f985_D])).
fof(f985_D,plain,(
  ~sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f817,plain,(
  ~sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f379])).
fof(f818,plain,(
  ~sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f379])).
fof(f981,plain,(
  ~sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl137),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  ~sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f819,plain,(
  ~sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f379])).
fof(f983,plain,(
  ~sP6('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl139),
  inference(cnf_transformation,[],[f983_D])).
fof(f983_D,plain,(
  ~sP6('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f820,plain,(
  ~sP6('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f379])).
fof(f821,plain,(
  ~sP6('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f379])).
fof(f822,plain,(
  ~sP6('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f379])).
fof(f979,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl134),
  inference(cnf_transformation,[],[f979_D])).
fof(f979_D,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f824,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f249,f439,f385])).
fof(f825,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f249,f439,f385])).
fof(f971,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl126),
  inference(cnf_transformation,[],[f971_D])).
fof(f971_D,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f826,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f249,f439,f385])).
fof(f827,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f248,f439,f385])).
fof(f828,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f248,f439,f385])).
fof(f829,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f248,f439,f385])).
fof(f975,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl130),
  inference(cnf_transformation,[],[f975_D])).
fof(f975_D,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f830,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f439,f385])).
fof(f831,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f439,f385])).
fof(f967,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl122),
  inference(cnf_transformation,[],[f967_D])).
fof(f967_D,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f832,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f439,f385])).
fof(f977,plain,(
  sP7('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl132),
  inference(cnf_transformation,[],[f977_D])).
fof(f977_D,plain,(
  sP7('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f833,plain,(
  sP7('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f249,f439,f385])).
fof(f834,plain,(
  sP7('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f249,f439,f385])).
fof(f835,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f249,f439,f385])).
fof(f836,plain,(
  sP7('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f248,f439,f385])).
fof(f837,plain,(
  sP7('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f248,f439,f385])).
fof(f838,plain,(
  sP7('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f248,f439,f385])).
fof(f969,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl124),
  inference(cnf_transformation,[],[f969_D])).
fof(f969_D,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f839,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f439,f385])).
fof(f840,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f439,f385])).
fof(f841,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f439,f385])).
fof(f973,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl128),
  inference(cnf_transformation,[],[f973_D])).
fof(f973_D,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f842,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f249,f439,f385])).
fof(f843,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f249,f439,f385])).
fof(f844,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f249,f439,f385])).
fof(f845,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f248,f439,f385])).
fof(f846,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f248,f248,f439,f385])).
fof(f847,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f248,f439,f385])).
fof(f848,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f439,f439,f385])).
fof(f849,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f439,f439,f385])).
fof(f850,plain,(
  sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f439,f439,f385])).
fof(f965,plain,(
  ~sP7('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl121),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  ~sP7('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f853,plain,(
  ~sP7('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f249,f439,f388])).
fof(f854,plain,(
  ~sP7('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f248,f439,f388])).
fof(f961,plain,(
  ~sP7('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl117),
  inference(cnf_transformation,[],[f961_D])).
fof(f961_D,plain,(
  ~sP7('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f855,plain,(
  ~sP7('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f439,f388])).
fof(f963,plain,(
  ~sP7('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl119),
  inference(cnf_transformation,[],[f963_D])).
fof(f963_D,plain,(
  ~sP7('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f856,plain,(
  ~sP7('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f249,f439,f388])).
fof(f857,plain,(
  ~sP7('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f248,f439,f388])).
fof(f858,plain,(
  ~sP7('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f439,f388])).
fof(f959,plain,(
  ~sP8('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl115),
  inference(cnf_transformation,[],[f959_D])).
fof(f959_D,plain,(
  ~sP8('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','F_0') <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f861,plain,(
  ~sP8('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f254,f439,f396])).
fof(f957,plain,(
  ~sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl113),
  inference(cnf_transformation,[],[f957_D])).
fof(f957_D,plain,(
  ~sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','F_0') <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f876,plain,(
  ~sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f254,f439,f403])).
fof(f955,plain,(
  ~sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl111),
  inference(cnf_transformation,[],[f955_D])).
fof(f955_D,plain,(
  ~sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f880,plain,(
  ~sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f249,f439,f413])).
fof(f881,plain,(
  ~sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f248,f439,f413])).
fof(f951,plain,(
  ~sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl107),
  inference(cnf_transformation,[],[f951_D])).
fof(f951_D,plain,(
  ~sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f882,plain,(
  ~sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f439,f413])).
fof(f953,plain,(
  ~sP10('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl109),
  inference(cnf_transformation,[],[f953_D])).
fof(f953_D,plain,(
  ~sP10('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f883,plain,(
  ~sP10('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f254,f439,f413])).
fof(f884,plain,(
  ~sP10('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f254,f439,f413])).
fof(f885,plain,(
  ~sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f254,f439,f413])).
fof(f949,plain,(
  sP11('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl104),
  inference(cnf_transformation,[],[f949_D])).
fof(f949_D,plain,(
  sP11('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f887,plain,(
  sP11('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f254,f439,f417])).
fof(f947,plain,(
  ~sP11('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl103),
  inference(cnf_transformation,[],[f947_D])).
fof(f947_D,plain,(
  ~sP11('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f895,plain,(
  ~sP11('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f249,f439,f420])).
fof(f896,plain,(
  ~sP11('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f248,f439,f420])).
fof(f943,plain,(
  ~sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl99),
  inference(cnf_transformation,[],[f943_D])).
fof(f943_D,plain,(
  ~sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f897,plain,(
  ~sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f439,f420])).
fof(f945,plain,(
  ~sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl101),
  inference(cnf_transformation,[],[f945_D])).
fof(f945_D,plain,(
  ~sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f898,plain,(
  ~sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f249,f439,f420])).
fof(f899,plain,(
  ~sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f254,f248,f439,f420])).
fof(f900,plain,(
  ~sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK14)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f439,f439,f420])).
fof(f902,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f254,f249,f439,f429])).
fof(f903,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f254,f248,f439,f429])).
fof(f904,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f254,f439,f439,f429])).
fof(f907,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f254,f439,f439,f429])).
fof(f908,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f254,f254,f439,f435])).
fof(f755,plain,(
  sP5('F_0','F_0','T_0') | $spl94),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f733,plain,(
  sP5('F_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f254,f649,f368])).
fof(f753,plain,(
  ~sP5('T_0','F_0','T_0') | $spl93),
  inference(cnf_transformation,[],[f753_D])).
fof(f753_D,plain,(
  ~sP5('T_0','F_0','T_0') <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f734,plain,(
  ~sP5('T_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f249,f649,f370])).
fof(f735,plain,(
  ~sP5('T_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f248,f649,f370])).
fof(f751,plain,(
  sP5('F_0','T_0','F_0') | $spl90),
  inference(cnf_transformation,[],[f751_D])).
fof(f751_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f741,plain,(
  sP5('F_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f254,f649,f368])).
fof(f749,plain,(
  ~sP5('T_0','T_0','F_0') | $spl89),
  inference(cnf_transformation,[],[f749_D])).
fof(f749_D,plain,(
  ~sP5('T_0','T_0','F_0') <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f742,plain,(
  ~sP5('T_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f249,f649,f370])).
fof(f743,plain,(
  ~sP5('T_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f248,f649,f370])).
fof(f727,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl86),
  inference(cnf_transformation,[],[f727_D])).
fof(f727_D,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f622,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(resolution,[],[f437,f271])).
fof(f271,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f160])).
fof(f621,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(resolution,[],[f437,f269])).
fof(f269,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f159])).
fof(f647,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl7),
  inference(cnf_transformation,[],[f647_D])).
fof(f647_D,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f444,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f264])).
fof(f264,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f157])).
fof(f445,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f266])).
fof(f266,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f158])).
fof(f446,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f269])).
fof(f447,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f271])).
fof(f452,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f307])).
fof(f307,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f125])).
fof(f454,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f307])).
fof(f725,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl85),
  inference(cnf_transformation,[],[f725_D])).
fof(f725_D,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f456,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f315])).
fof(f315,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f136])).
fof(f721,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl81),
  inference(cnf_transformation,[],[f721_D])).
fof(f721_D,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f457,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f315])).
fof(f723,plain,(
  ~sP0('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl83),
  inference(cnf_transformation,[],[f723_D])).
fof(f723_D,plain,(
  ~sP0('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f458,plain,(
  ~sP0('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f315])).
fof(f459,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f315])).
fof(f460,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f319])).
fof(f319,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f195])).
fof(f461,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f319])).
fof(f462,plain,(
  ~sP0('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f319])).
fof(f463,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f319])).
fof(f713,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl73),
  inference(cnf_transformation,[],[f713_D])).
fof(f713_D,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f464,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f330])).
fof(f330,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f200])).
fof(f711,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl71),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f465,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f330])).
fof(f715,plain,(
  ~sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl75),
  inference(cnf_transformation,[],[f715_D])).
fof(f715_D,plain,(
  ~sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f466,plain,(
  ~sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f330])).
fof(f467,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f330])).
fof(f717,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl76),
  inference(cnf_transformation,[],[f717_D])).
fof(f717_D,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f468,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f332])).
fof(f469,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f332])).
fof(f719,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl78),
  inference(cnf_transformation,[],[f719_D])).
fof(f719_D,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f470,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f335])).
fof(f471,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f335])).
fof(f472,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f336])).
fof(f473,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f336])).
fof(f474,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f337])).
fof(f475,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f337])).
fof(f476,plain,(
  ~sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f339])).
fof(f339,plain,(
  ( ! [X0,X1] : (~sP1(X1,X0) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f201])).
fof(f477,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f339])).
fof(f478,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f339])).
fof(f479,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f339])).
fof(f709,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl69),
  inference(cnf_transformation,[],[f709_D])).
fof(f709_D,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f480,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f249,f437,f352])).
fof(f481,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f249,f437,f352])).
fof(f482,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f248,f437,f352])).
fof(f483,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f248,f437,f352])).
fof(f707,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl67),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f484,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f249,f437,f356])).
fof(f485,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f249,f437,f356])).
fof(f486,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f248,f437,f356])).
fof(f487,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f248,f437,f356])).
fof(f705,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl65),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f490,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f249,f437,f360])).
fof(f491,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f249,f437,f360])).
fof(f492,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f248,f437,f360])).
fof(f493,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f248,f437,f360])).
fof(f703,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl63),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f496,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f218])).
fof(f699,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl59),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f497,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f364])).
fof(f695,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl55),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f498,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f437,f364])).
fof(f691,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl51),
  inference(cnf_transformation,[],[f691_D])).
fof(f691_D,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f499,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f437,f364])).
fof(f701,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl61),
  inference(cnf_transformation,[],[f701_D])).
fof(f701_D,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f500,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f364])).
fof(f501,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f364])).
fof(f693,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl53),
  inference(cnf_transformation,[],[f693_D])).
fof(f693_D,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f502,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f437,f364])).
fof(f503,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f437,f364])).
fof(f697,plain,(
  ~sP2('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl57),
  inference(cnf_transformation,[],[f697_D])).
fof(f697_D,plain,(
  ~sP2('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f504,plain,(
  ~sP2('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f364])).
fof(f505,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f364])).
fof(f506,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f437,f364])).
fof(f507,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f437,f364])).
fof(f689,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl49),
  inference(cnf_transformation,[],[f689_D])).
fof(f689_D,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f512,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f249,f437,f379])).
fof(f513,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f249,f437,f379])).
fof(f514,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f248,f437,f379])).
fof(f515,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f248,f437,f379])).
fof(f687,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl47),
  inference(cnf_transformation,[],[f687_D])).
fof(f687_D,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f530,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f249,f437,f388])).
fof(f531,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f249,f437,f388])).
fof(f532,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f248,f437,f388])).
fof(f533,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f248,f437,f388])).
fof(f685,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl45),
  inference(cnf_transformation,[],[f685_D])).
fof(f685_D,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f538,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f254,f437,f396])).
fof(f539,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f254,f437,f396])).
fof(f681,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl41),
  inference(cnf_transformation,[],[f681_D])).
fof(f681_D,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f540,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f437,f396])).
fof(f541,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f437,f396])).
fof(f683,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl43),
  inference(cnf_transformation,[],[f683_D])).
fof(f683_D,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f542,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f254,f437,f396])).
fof(f543,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f254,f437,f396])).
fof(f544,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f437,f396])).
fof(f545,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f437,f396])).
fof(f679,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl38),
  inference(cnf_transformation,[],[f679_D])).
fof(f679_D,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f547,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f671,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl30),
  inference(cnf_transformation,[],[f671_D])).
fof(f671_D,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f548,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f400])).
fof(f675,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl34),
  inference(cnf_transformation,[],[f675_D])).
fof(f675_D,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f549,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f437,f400])).
fof(f667,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl26),
  inference(cnf_transformation,[],[f667_D])).
fof(f667_D,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f550,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f437,f400])).
fof(f677,plain,(
  sP9('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl36),
  inference(cnf_transformation,[],[f677_D])).
fof(f677_D,plain,(
  sP9('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f551,plain,(
  sP9('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f400])).
fof(f669,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl28),
  inference(cnf_transformation,[],[f669_D])).
fof(f669_D,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f552,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f400])).
fof(f553,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f437,f400])).
fof(f554,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f437,f400])).
fof(f673,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl32),
  inference(cnf_transformation,[],[f673_D])).
fof(f673_D,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f555,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f254,f437,f400])).
fof(f556,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f437,f254,f437,f400])).
fof(f557,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f437,f437,f400])).
fof(f558,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f437,f437,f400])).
fof(f665,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl25),
  inference(cnf_transformation,[],[f665_D])).
fof(f665_D,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f559,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f249,f437,f403])).
fof(f661,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl21),
  inference(cnf_transformation,[],[f661_D])).
fof(f661_D,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f560,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f249,f437,f403])).
fof(f561,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f248,f437,f403])).
fof(f562,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f248,f437,f403])).
fof(f663,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl23),
  inference(cnf_transformation,[],[f663_D])).
fof(f663_D,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f563,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f249,f437,f403])).
fof(f564,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f249,f437,f403])).
fof(f565,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f248,f437,f403])).
fof(f566,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f437,f248,f437,f403])).
fof(f659,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl19),
  inference(cnf_transformation,[],[f659_D])).
fof(f659_D,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f577,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f249,f437,f413])).
fof(f578,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f249,f437,f413])).
fof(f579,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f248,f437,f413])).
fof(f580,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f248,f437,f413])).
fof(f657,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl16),
  inference(cnf_transformation,[],[f657_D])).
fof(f657_D,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f583,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f249,f437,f417])).
fof(f653,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl12),
  inference(cnf_transformation,[],[f653_D])).
fof(f653_D,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f584,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f437,f249,f437,f417])).
fof(f585,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f254,f248,f437,f417])).
fof(f586,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f437,f248,f437,f417])).
fof(f655,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl14),
  inference(cnf_transformation,[],[f655_D])).
fof(f655_D,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f587,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f254,f437,f417])).
fof(f588,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f254,f437,f417])).
fof(f589,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f437,f417])).
fof(f590,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f437,f417])).
fof(f651,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl11),
  inference(cnf_transformation,[],[f651_D])).
fof(f651_D,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f593,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f249,f437,f420])).
fof(f594,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f249,f437,f420])).
fof(f595,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f248,f437,f420])).
fof(f596,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f248,f437,f420])).
fof(f603,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f249,f249,f437,f429])).
fof(f604,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f248,f249,f437,f429])).
fof(f605,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f249,f248,f437,f429])).
fof(f606,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f248,f248,f437,f429])).
fof(f649,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f649_D])).
fof(f649_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f607,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f254,f249,f437,f429])).
fof(f608,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f254,f249,f437,f429])).
fof(f609,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f249,f437,f429])).
fof(f610,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f249,f437,f429])).
fof(f611,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f254,f248,f437,f429])).
fof(f612,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f254,f248,f437,f429])).
fof(f613,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f248,f437,f429])).
fof(f614,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f248,f437,f429])).
fof(f615,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f249,f437,f435])).
fof(f616,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f254,f248,f437,f435])).
fof(f617,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f249,f254,f437,f435])).
fof(f618,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f248,f254,f437,f435])).
fof(f619,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f249,f437,f437,f435])).
fof(f620,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f248,f437,f437,f435])).
fof(f441,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) | $spl4),
  inference(cnf_transformation,[],[f441_D])).
fof(f441_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f245,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) & p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) & ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f107])).
fof(f107,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X2)))) & ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X2))))) & ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X2))))) => p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X2))))) => p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X8))))) => p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X8))))) => p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X7)))))),
  file('thm.util_prob.COUNTABLE_SUBSET','thm.util_prob.COUNTABLE_SUBSET')).
fof(f439,plain,(
  p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) | $spl2),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f246,plain,(
  p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK14))))),
  inference(cnf_transformation,[],[f154])).
fof(f437,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) | $spl1),
  inference(cnf_transformation,[],[f437_D])).
fof(f437_D,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f247,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))))),
  inference(cnf_transformation,[],[f154])).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 487
% Passive clauses: 622
% Generated clauses: 5556
% Final active clauses: 482
% Final passive clauses: 264
% Input formulas: 41
% Initial clauses: 186
% 
% Duplicate literals: 98
% Trivial inequalities: 6
% Fw demodulations: 1148
% Bw demodulations: 12
% Global subsumptions: 6
% 
% Simple tautologies: 311
% Equational tautologies: 18
% Forward subsumptions: 2924
% Backward subsumptions: 2
% 
% Binary resolution: 1331
% Unit resulting resolution: 2501
% Forward superposition: 11
% Backward superposition: 244
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 1584
% SAT solver unit clauses: 6
% SAT solver binary clauses: 1328
% 
% Sat splits: 3
% 
% Memory used [KB]: 4349
% Time elapsed: 0.110 s
% ------------------------------
% ------------------------------
% Success in time 18.176 s
