Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (97959)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 172
% Passive clauses: 4265
% Generated clauses: 30947
% Final active clauses: 159
% Final passive clauses: 3921
% Input formulas: 69
% Initial clauses: 152
% 
% Duplicate literals: 948
% Trivial inequalities: 588
% Fw subsumption resolutions: 4340
% Fw demodulations: 5629
% Bw demodulations: 10
% 
% Simple tautologies: 593
% Equational tautologies: 67
% Forward subsumptions: 13731
% 
% Binary resolution: 249
% Factoring: 1
% Forward superposition: 8687
% Backward superposition: 9801
% Self superposition: 273
% Equality factoring: 138
% Equality resolution: 41
% 
% Split clauses: 319
% Split components: 639
% SAT solver clauses: 718
% SAT solver unit clauses: 49
% SAT solver binary clauses: 299
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 3
% 
% Sat splits: 319
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 716
% TWLsolver calls for satisfiability: 91
% 
% Memory used [KB]: 17014
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (97974)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 215
% Passive clauses: 1685
% Generated clauses: 21673
% Final active clauses: 154
% Final passive clauses: 936
% Input formulas: 69
% Initial clauses: 152
% 
% Duplicate literals: 419
% Trivial inequalities: 413
% Fw subsumption resolutions: 1578
% Fw demodulations: 4197
% Bw demodulations: 564
% Condensations: 567
% Global subsumptions: 17
% 
% Simple tautologies: 193
% Equational tautologies: 31
% Forward subsumptions: 10127
% Bw demodulations to eq. taut.: 9
% 
% Binary resolution: 207
% Factoring: 1
% Forward superposition: 5226
% Backward superposition: 5800
% Self superposition: 124
% Equality factoring: 41
% Equality resolution: 41
% 
% Split clauses: 85
% Split components: 170
% SAT solver clauses: 2033
% SAT solver unit clauses: 177
% SAT solver binary clauses: 63
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 3
% 
% Sat splits: 85
% Sat splitting refutations: 81
% 
% TWLsolver clauses: 1963
% TWLsolver calls for satisfiability: 843
% 
% Memory used [KB]: 18677
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f58490,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f263,f425,f423,f262,f424,f422,f430,f792,f798,f796,f4070,f4072,f6441,f6473,f6476,f6477,f6478,f6479,f6434,f6481,f6480,f6673,f6677,f6679,f10961,f10971,f10972,f10973,f10974,f10975,f13904,f13906,f13924,f13934,f13936,f13938,f13940,f13944,f13946,f13948,f13950,f13952,f13954,f13956,f13958,f13962,f13964,f13966,f13968,f13970,f13972,f13974,f13976,f13978,f13980,f13982,f14434,f14435,f14436,f14437,f14438,f14439,f14440,f14441,f17016,f17018,f17021,f16822,f17029,f17097,f17017,f17022,f17634,f17641,f17642,f17643,f26959,f26930,f26919,f6472,f27088,f32386,f32361,f32357,f17030,f32429,f34128,f34129,f34143,f34181,f34184,f34220,f34222,f34224,f34226,f34228,f34232,f34234,f34236,f34238,f34248,f34251,f34254,f34257,f34260,f34262,f34264,f34266,f34268,f34270,f34272,f34276,f34280,f34282,f34285,f34288,f34309,f34311,f34323,f34348,f34350,f34352,f34358,f34373,f34386,f34387,f34403,f34409,f34411,f34413,f34415,f34417,f34419,f34421,f34423,f34425,f34427,f34429,f34431,f34433,f34435,f34437,f34439,f34441,f34443,f34445,f34447,f34449,f34455,f34459,f34461,f34463,f34465,f34471,f34473,f34475,f34477,f34489,f34495,f34501,f34505,f34509,f34513,f34515,f34517,f34519,f34521,f34528,f34530,f34534,f34539,f34541,f34544,f34545,f34546,f34548,f34560,f34562,f34564,f34566,f34573,f34586,f34588,f34598,f49458,f49460,f49467,f49469,f49468,f49504,f49505,f49459,f49517,f55442,f55444,f4073,f55443,f58443,f58445])).
fof(f58445,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f58444,f336])).
fof(f336,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X27] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))),
  file('thm.list.LIST_REL_NIL|split|0','thm.num.NOT_SUC')).
fof(f58444,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),sK23('const.list.NIL_0',sK3,sK4,sK2,sK1)))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f58408,f273])).
fof(f273,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LENGTH|split|0')).
fof(f58408,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),sK23('const.list.NIL_0',sK3,sK4,sK2,sK1)))))) | ($spl0 | $spl3)),
  inference(superposition,[],[f345,f25167])).
fof(f25167,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK21('const.list.NIL_0',sK3,sK4,sK2,sK1)),s('type.list.list'(sK1),sK23('const.list.NIL_0',sK3,sK4,sK2,sK1)))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f4074,f425,f386])).
fof(f386,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3) | s('type.list.list'(X0),X4) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK21(X4,X3,X2,X1,X0)),s('type.list.list'(X0),sK23(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3)) | (s('type.list.list'(X0),X4) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK21(X4,X3,X2,X1,X0)),s('type.list.list'(X0),sK23(X4,X3,X2,X1,X0)))) & s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK22(X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK24(X4,X3,X2,X1,X0)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK21(X4,X3,X2,X1,X0)))),s(X1,sK22(X4,X3,X2,X1,X0))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),sK23(X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK24(X4,X3,X2,X1,X0))))))) & (((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X4) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3)) & ! [X5,X6,X7,X8] : (s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) | sP0(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22,sK23,sK24])],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3)) | ? [X9,X10,X11,X12] : (s('type.list.list'(X0),X4) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X9),s('type.list.list'(X0),X11))) & s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X10),s('type.list.list'(X1),X12))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X9))),s(X1,X10)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X11),s('type.list.list'(X1),X12)))))) & (((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X4) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3)) & ! [X5,X6,X7,X8] : (s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) | sP0(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X4,X3] : ((~sP0(X0,X1,X2,X4,X3) | (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) & (((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X3) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X4)) & ! [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) | s('type.list.list'(X1),X4) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) | sP0(X0,X1,X2,X4,X3)))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X4,X3] : ((~sP0(X0,X1,X2,X4,X3) | ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8))))))) & (((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X3) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X4)) & ! [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) | s('type.list.list'(X1),X4) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) | sP0(X0,X1,X2,X4,X3)))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X4,X3] : (sP0(X0,X1,X2,X4,X3) <=> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f4074,plain,(
  sP0(sK1,sK2,sK4,sK3,'const.list.NIL_0') | $spl0),
  inference(subsumption_resolution,[],[f4057,f264])).
fof(f264,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.LIST_REL_NIL|split|0',aHL_TRUTH)).
fof(f4057,plain,(
  ~p(s(bool,'T_0')) | sP0(sK1,sK2,sK4,sK3,'const.list.NIL_0') | $spl0),
  inference(superposition,[],[f395,f447])).
fof(f447,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | $spl0),
  inference(unit_resulting_resolution,[],[f422,f282])).
fof(f282,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.list.LIST_REL_NIL|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f395,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | sP0(X0,X1,X2,X4,X3)) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | sP0(X0,X1,X2,X4,X3)) & (~sP0(X0,X1,X2,X4,X3) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> sP0(X0,X1,X2,X4,X3))),
  inference(definition_folding,[],[f166,f194])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : ! [X9,X20,X21] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),X20),s('type.list.list'(X7),X21)))) <=> ((s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X20) & s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X21)) | ? [X22,X23,X14,X15] : (s('type.list.list'(X6),X20) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X22),s('type.list.list'(X6),X14))) & s('type.list.list'(X7),X21) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X23),s('type.list.list'(X7),X15))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X9),s(X6,X22))),s(X7,X23)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),X14),s('type.list.list'(X7),X15)))))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LIST_REL_cases')).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6] : ! [X44,X10] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X10))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X44),s('type.list.list'(X6),X10)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LENGTH|split|1')).
fof(f58443,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f58442,f437])).
fof(f437,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f426,f281])).
fof(f281,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f426,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f267,f264])).
fof(f267,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f50])).
fof(f50,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.LIST_REL_NIL|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f58442,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f58406,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0')))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ! [X2] : s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X7] : ! [X9] : s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X7),'const.list.NIL_0')))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LIST_REL_def|split|0')).
fof(f58406,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(sK1,bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(sK1),'const.list.NIL_0')))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f380,f25167])).
fof(f380,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X2)))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X2)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6,X7] : ! [X18,X19,X9] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X19),s('type.list.list'(X7),X18)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LIST_REL_def|split|2')).
fof(f55443,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11))) | s(bool,'F_0') = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) ) | $spl50),
  inference(cnf_transformation,[],[f55443_D])).
fof(f55443_D,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11))) | s(bool,'F_0') = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4073,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(cnf_transformation,[],[f4073_D])).
fof(f4073_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f55444,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X17)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),X15))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,X16)))) ) | $spl0),
  inference(subsumption_resolution,[],[f55373,f285])).
fof(f285,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.list.LIST_REL_NIL|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f55373,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X17)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),X15))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,X16))) | s(bool,'F_0') = s(bool,X17)) ) | $spl0),
  inference(superposition,[],[f54944,f879])).
fof(f879,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f785,f785])).
fof(f785,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f775,f264])).
fof(f775,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,X4) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f623,f447])).
fof(f623,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | s(bool,X12) = s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f622,f271])).
fof(f271,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X8] : s(bool,X8) = s(bool,'const.marker.Cong_1'(s(bool,X8)))),
  file('thm.list.LIST_REL_NIL|split|0','thm.marker.Cong_def')).
fof(f622,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | s(bool,X13) = s(bool,'const.marker.Cong_1'(s(bool,X12))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f620,f264])).
fof(f620,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | s(bool,X13) = s(bool,'const.marker.Cong_1'(s(bool,X12))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f364,f489])).
fof(f489,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f276,f271])).
fof(f276,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.LIST_REL_NIL|split|0',aHL_BOOL_CASES)).
fof(f364,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.list.LIST_REL_NIL|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f54944,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) ) | $spl0),
  inference(subsumption_resolution,[],[f54943,f282])).
fof(f54943,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) ) | $spl0),
  inference(forward_demodulation,[],[f54497,f271])).
fof(f54497,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))))) ) | $spl0),
  inference(superposition,[],[f348,f9980])).
fof(f9980,plain,(
  ( ! [X132,X130,X131] : (s(X130,X132) = s(X130,sK6(X132,X131,X130)) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'|dollar|exists_1'(s(fun(X130,bool),X131)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f9979,f426])).
fof(f9979,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,'F_0')) | s(X130,X132) = s(X130,sK6(X132,X131,X130)) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'|dollar|exists_1'(s(fun(X130,bool),X131)))))) ) | $spl0),
  inference(forward_demodulation,[],[f9556,f271])).
fof(f9556,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) | s(X130,X132) = s(X130,sK6(X132,X131,X130)) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'|dollar|exists_1'(s(fun(X130,bool),X131)))))) ) | $spl0),
  inference(superposition,[],[f347,f8912])).
fof(f8912,plain,(
  ( ! [X47] : (s(bool,X47) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X47)))) ) | $spl0),
  inference(superposition,[],[f8650,f271])).
fof(f8650,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(condensation,[],[f8649])).
fof(f8649,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f8316])).
fof(f8316,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f2308,f276])).
fof(f2308,plain,(
  ( ! [X19,X20] : (s(bool,X20) = s(bool,'const.marker.Cong_1'(s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | $spl0),
  inference(subsumption_resolution,[],[f2260,f426])).
fof(f2260,plain,(
  ( ! [X19,X20] : (p(s(bool,'F_0')) | s(bool,X20) = s(bool,'const.marker.Cong_1'(s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | $spl0),
  inference(superposition,[],[f546,f1930])).
fof(f1930,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(condensation,[],[f1929])).
fof(f1929,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1513])).
fof(f1513,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f1004,f276])).
fof(f1004,plain,(
  ( ! [X33,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X32)))) ) | $spl0),
  inference(superposition,[],[f879,f271])).
fof(f546,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f534,f426])).
fof(f534,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f361,f276])).
fof(f361,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK6(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X17,X10] : (! [X8] : (s(X6,X8) = s(X6,X10) => p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X8))))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X17)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.bool.JRH_INDUCT_UTIL')).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f55442,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X13)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11))) | s(bool,'F_0') = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) ) | $spl0),
  inference(subsumption_resolution,[],[f55372,f285])).
fof(f55372,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X13)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11))) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) ) | $spl0),
  inference(superposition,[],[f54944,f879])).
fof(f49517,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f343,f49459,f281])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : ! [X44,X10] : s(X6,X44) = s(X6,'const.list.HD_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X44),s('type.list.list'(X6),X10)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.HD')).
fof(f49459,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0))))) ) | $spl42),
  inference(cnf_transformation,[],[f49459_D])).
fof(f49459_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f49505,plain,(
  ( ! [X166] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK7('const.list.HD_0',X166,'type.list.list'(bool)))))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X166),s('type.list.list'(bool),sK8('const.list.HD_0',X166,'type.list.list'(bool))))))) ) | $spl48),
  inference(cnf_transformation,[],[f49505_D])).
fof(f49505_D,plain,(
  ( ! [X166] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK7('const.list.HD_0',X166,'type.list.list'(bool)))))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X166),s('type.list.list'(bool),sK8('const.list.HD_0',X166,'type.list.list'(bool))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f49504,plain,(
  ( ! [X165,X166] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK7('const.list.HD_0',X166,'type.list.list'(bool)))))) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X165)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X166),s('type.list.list'(bool),sK8('const.list.HD_0',X166,'type.list.list'(bool))))))) )),
  inference(forward_demodulation,[],[f49436,f338])).
fof(f338,plain,(
  ( ! [X0,X1] : (s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.HD_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.HD_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1] : s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.HD_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6] : ! [X34] : s(X6,'const.list.HD_1'(s('type.list.list'(X6),X34))) = s(X6,happ(s(fun('type.list.list'(X6),X6),'const.list.HD_0'),s('type.list.list'(X6),X34)))),
  file('thm.list.LIST_REL_NIL|split|0','ARITY_0_1_const.list.HD')).
fof(f49436,plain,(
  ( ! [X165,X166] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X165)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X166),s('type.list.list'(bool),sK8('const.list.HD_0',X166,'type.list.list'(bool)))))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),sK7('const.list.HD_0',X166,'type.list.list'(bool))))))) )),
  inference(superposition,[],[f349,f338])).
fof(f349,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK7(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK7(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X16,X17] : (! [X8] : (p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X8)))) => p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X8))))) => (? [X8] : p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X8)))) => ? [X8] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X8))))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.bool.MONO_EXISTS')).
fof(f49468,plain,(
  ( ! [X10,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK8(X10,X13,X9))))) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7(X10,X13,X9)))))) ) | $spl46),
  inference(cnf_transformation,[],[f49468_D])).
fof(f49468_D,plain,(
  ( ! [X10,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK8(X10,X13,X9))))) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7(X10,X13,X9)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f49469,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,X17)) | p(s(bool,happ(s(fun(X14,bool),X18),s(X14,sK8(X15,X18,X14))))) | p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK7(X15,X18,X14))))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,X16)))) ) | $spl0),
  inference(subsumption_resolution,[],[f49406,f285])).
fof(f49406,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,X17)) | p(s(bool,happ(s(fun(X14,bool),X18),s(X14,sK8(X15,X18,X14))))) | p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK7(X15,X18,X14))))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,X16))) | s(bool,'F_0') = s(bool,X17)) ) | $spl0),
  inference(superposition,[],[f349,f879])).
fof(f49467,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK8(X10,X13,X9))))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7(X10,X13,X9))))) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) ) | $spl0),
  inference(subsumption_resolution,[],[f49405,f285])).
fof(f49405,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK8(X10,X13,X9))))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7(X10,X13,X9))))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) ) | $spl0),
  inference(superposition,[],[f349,f879])).
fof(f49460,plain,(
  ( ! [X1] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK7('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool)))))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X1),s('type.list.list'(bool),sK8('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool))))))) ) | $spl44),
  inference(cnf_transformation,[],[f49460_D])).
fof(f49460_D,plain,(
  ( ! [X1] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK7('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool)))))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X1),s('type.list.list'(bool),sK8('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool))))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f49458,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK7('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool)))))) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X1),s('type.list.list'(bool),sK8('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool))))))) )),
  inference(forward_demodulation,[],[f49457,f4561])).
fof(f4561,plain,(
  ( ! [X0,X1] : (s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(X0),X1)))) )),
  inference(forward_demodulation,[],[f4475,f338])).
fof(f4475,plain,(
  ( ! [X0,X1] : (s(X0,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(X0),X1))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.HD_0'),s('type.list.list'(X0),X1)))) )),
  inference(superposition,[],[f346,f275])).
fof(f275,plain,(
  ( ! [X0] : (s(fun('type.list.list'(X0),X0),'const.list.HD_0') = s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s(fun('type.list.list'(X0),X0),'const.list.HD_0') = s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : s(fun('type.list.list'(X6),X6),'const.list.HD_0') = s(fun('type.list.list'(X6),X6),'const.list.EL_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.EL_restricted|split|0')).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6] : ! [X34,X35] : s(X6,'const.list.EL_2'(s('type.num.num',X34),s('type.list.list'(X6),X35))) = s(X6,happ(s(fun('type.list.list'(X6),X6),'const.list.EL_1'(s('type.num.num',X34))),s('type.list.list'(X6),X35)))),
  file('thm.list.LIST_REL_NIL|split|0','ARITY_1_2_const.list.EL')).
fof(f49457,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),sK7('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool)))))) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X1),s('type.list.list'(bool),sK8('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool))))))) )),
  inference(forward_demodulation,[],[f49456,f346])).
fof(f49456,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X1),s('type.list.list'(bool),sK8('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool)))))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num','const.num.|48|_0'))),s('type.list.list'(bool),sK7('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool))))))) )),
  inference(forward_demodulation,[],[f49308,f338])).
fof(f49308,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X0)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),X1),s('type.list.list'(bool),sK8('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool)))))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num','const.num.|48|_0'))),s('type.list.list'(bool),sK7('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X1,'type.list.list'(bool))))))) )),
  inference(superposition,[],[f349,f275])).
fof(f34598,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34597,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547,f34585,f34587])).
fof(f34587,plain,(
  ( ! [X0] : (s('type.list.list'(bool),X0) = s('type.list.list'(bool),sK9(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))) ) | $spl36),
  inference(subsumption_resolution,[],[f34063,f264])).
fof(f34063,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.list.list'(bool),X0) = s('type.list.list'(bool),sK9(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2134])).
fof(f34585,plain,(
  ( ! [X34] : (s('type.list.list'(bool),X34) = s('type.list.list'(bool),sK9(X34,'const.list.HD_0','type.list.list'(bool)))) ) | $spl36),
  inference(subsumption_resolution,[],[f34062,f264])).
fof(f34062,plain,(
  ( ! [X34] : (~p(s(bool,'T_0')) | s('type.list.list'(bool),X34) = s('type.list.list'(bool),sK9(X34,'const.list.HD_0','type.list.list'(bool)))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2133])).
fof(f34547,plain,(
  ( ! [X17,X16] : (s('type.list.list'(bool),X17) = s('type.list.list'(bool),sK9(X17,'const.list.EL_1'(s('type.num.num',X16)),'type.list.list'(bool)))) ) | $spl36),
  inference(subsumption_resolution,[],[f34010,f264])).
fof(f34010,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | s('type.list.list'(bool),X17) = s('type.list.list'(bool),sK9(X17,'const.list.EL_1'(s('type.num.num',X16)),'type.list.list'(bool)))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4496])).
fof(f34533,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X31),X33))))) = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(X31),'const.list.CONS_2'(s(X31,X32),s('type.list.list'(X31),X33)))),X34,'type.num.num'))) ) | $spl36),
  inference(subsumption_resolution,[],[f33994,f264])).
fof(f33994,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X31),X33))))) = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(X31),'const.list.CONS_2'(s(X31,X32),s('type.list.list'(X31),X33)))),X34,'type.num.num'))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f3277])).
fof(f34520,plain,(
  ( ! [X70,X71] : (s('type.num.num',X70) = s('type.num.num',sK9(sK5('const.num.SUC_1'(s('type.num.num',X70))),X71,'type.num.num'))) ) | $spl36),
  inference(subsumption_resolution,[],[f33984,f264])).
fof(f33984,plain,(
  ( ! [X70,X71] : (~p(s(bool,'T_0')) | s('type.num.num',X70) = s('type.num.num',sK9(sK5('const.num.SUC_1'(s('type.num.num',X70))),X71,'type.num.num'))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2154])).
fof(f34518,plain,(
  ( ! [X68,X69] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X68)))) = s('type.num.num',sK9('const.num.SUC_1'(s('type.num.num',sK5(X68))),X69,'type.num.num')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X68)) ) | $spl36),
  inference(subsumption_resolution,[],[f33983,f264])).
fof(f33983,plain,(
  ( ! [X68,X69] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X68)))) = s('type.num.num',sK9('const.num.SUC_1'(s('type.num.num',sK5(X68))),X69,'type.num.num')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X68)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2122])).
fof(f34516,plain,(
  ( ! [X66,X67] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(X66),'const.list.NIL_0')),X67,'type.num.num'))) ) | $spl36),
  inference(subsumption_resolution,[],[f33982,f264])).
fof(f33982,plain,(
  ( ! [X66,X67] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(X66),'const.list.NIL_0')),X67,'type.num.num'))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2153])).
fof(f34512,plain,(
  ( ! [X64,X65] : (s(fun('type.list.list'(X64),X64),'const.list.HD_0') = s(fun('type.list.list'(X64),X64),sK9('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X65,fun('type.list.list'(X64),X64)))) ) | $spl36),
  inference(subsumption_resolution,[],[f33978,f264])).
fof(f33978,plain,(
  ( ! [X64,X65] : (~p(s(bool,'T_0')) | s(fun('type.list.list'(X64),X64),'const.list.HD_0') = s(fun('type.list.list'(X64),X64),sK9('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X65,fun('type.list.list'(X64),X64)))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2152])).
fof(f34508,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,sK9('const.marker.Cong_1'(s(bool,X60)),X61,bool))) ) | $spl36),
  inference(subsumption_resolution,[],[f33957,f264])).
fof(f33957,plain,(
  ( ! [X61,X60] : (~p(s(bool,'T_0')) | s(bool,X60) = s(bool,sK9('const.marker.Cong_1'(s(bool,X60)),X61,bool))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2150])).
fof(f34504,plain,(
  ( ! [X10,X11,X9] : (s(X9,X11) = s(X9,sK9(X11,happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10)),X9))) ) | $spl36),
  inference(subsumption_resolution,[],[f33955,f264])).
fof(f33955,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | s(X9,X11) = s(X9,sK9(X11,happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10)),X9))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2094])).
fof(f34500,plain,(
  ( ! [X132,X130,X133,X131] : (s(X130,X133) = s(X130,sK9(X133,happ(s(fun('type.list.list'(fun(X130,bool)),fun(X130,bool)),'const.list.EL_1'(s('type.num.num',X131))),s('type.list.list'(fun(X130,bool)),X132)),X130))) ) | $spl36),
  inference(subsumption_resolution,[],[f33953,f264])).
fof(f33953,plain,(
  ( ! [X132,X130,X133,X131] : (~p(s(bool,'T_0')) | s(X130,X133) = s(X130,sK9(X133,happ(s(fun('type.list.list'(fun(X130,bool)),fun(X130,bool)),'const.list.EL_1'(s('type.num.num',X131))),s('type.list.list'(fun(X130,bool)),X132)),X130))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4541])).
fof(f34494,plain,(
  ( ! [X333,X335,X332,X334,X336] : (s(X332,X336) = s(X332,sK9(X336,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X333))),s('type.list.list'(fun(X332,bool)),'const.list.CONS_2'(s(fun(X332,bool),X334),s('type.list.list'(fun(X332,bool)),X335)))),X332))) ) | $spl36),
  inference(subsumption_resolution,[],[f33950,f264])).
fof(f33950,plain,(
  ( ! [X333,X335,X332,X334,X336] : (~p(s(bool,'T_0')) | s(X332,X336) = s(X332,sK9(X336,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X333))),s('type.list.list'(fun(X332,bool)),'const.list.CONS_2'(s(fun(X332,bool),X334),s('type.list.list'(fun(X332,bool)),X335)))),X332))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f11411])).
fof(f34488,plain,(
  ( ! [X198,X196,X194,X197,X195] : (s(X194,X198) = s(X194,sK9(X198,'const.ind_type.FCONS_3'(s(fun(X194,bool),X195),s(fun('type.num.num',fun(X194,bool)),X196),s('type.num.num','const.num.SUC_1'(s('type.num.num',X197)))),X194))) ) | $spl36),
  inference(subsumption_resolution,[],[f33945,f264])).
fof(f33945,plain,(
  ( ! [X198,X196,X194,X197,X195] : (~p(s(bool,'T_0')) | s(X194,X198) = s(X194,sK9(X198,'const.ind_type.FCONS_3'(s(fun(X194,bool),X195),s(fun('type.num.num',fun(X194,bool)),X196),s('type.num.num','const.num.SUC_1'(s('type.num.num',X197)))),X194))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f5582])).
fof(f34476,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.list.HD_1'(s('type.list.list'(X8),X9))) = s(X8,sK9(happ(s(fun('type.list.list'(X8),X8),'const.list.HD_0'),s('type.list.list'(X8),X9)),X10,X8))) ) | $spl36),
  inference(subsumption_resolution,[],[f33938,f264])).
fof(f33938,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | s(X8,'const.list.HD_1'(s('type.list.list'(X8),X9))) = s(X8,sK9(happ(s(fun('type.list.list'(X8),X8),'const.list.HD_0'),s('type.list.list'(X8),X9)),X10,X8))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2137])).
fof(f34474,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X4),X6))) = s(X4,sK9(happ(s(fun('type.list.list'(X4),X4),'const.list.EL_1'(s('type.num.num',X5))),s('type.list.list'(X4),X6)),X7,X4))) ) | $spl36),
  inference(subsumption_resolution,[],[f33937,f264])).
fof(f33937,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | s(X4,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X4),X6))) = s(X4,sK9(happ(s(fun('type.list.list'(X4),X4),'const.list.EL_1'(s('type.num.num',X5))),s('type.list.list'(X4),X6)),X7,X4))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4569])).
fof(f34472,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(X21,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(X21),X24))) = s(X21,sK9('const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.list.list'(X21),'const.list.CONS_2'(s(X21,X23),s('type.list.list'(X21),X24)))),X25,X21))) ) | $spl36),
  inference(subsumption_resolution,[],[f33936,f264])).
fof(f33936,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,'T_0')) | s(X21,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(X21),X24))) = s(X21,sK9('const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.list.list'(X21),'const.list.CONS_2'(s(X21,X23),s('type.list.list'(X21),X24)))),X25,X21))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f11461])).
fof(f34470,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8))) = s(X5,sK9('const.ind_type.FCONS_3'(s(X5,X6),s(fun('type.num.num',X5),X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))),X9,X5))) ) | $spl36),
  inference(subsumption_resolution,[],[f33935,f264])).
fof(f33935,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8))) = s(X5,sK9('const.ind_type.FCONS_3'(s(X5,X6),s(fun('type.num.num',X5),X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))),X9,X5))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f5619])).
fof(f34464,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,X5) = s(X4,sK9('const.ind_type.FCONS_3'(s(X4,X5),s(fun('type.num.num',X4),X6),s('type.num.num','const.num.|48|_0')),X7,X4))) ) | $spl36),
  inference(subsumption_resolution,[],[f33930,f264])).
fof(f33930,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | s(X4,X5) = s(X4,sK9('const.ind_type.FCONS_3'(s(X4,X5),s(fun('type.num.num',X4),X6),s('type.num.num','const.num.|48|_0')),X7,X4))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2136])).
fof(f34462,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,sK9('const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X3,X0))) ) | $spl36),
  inference(subsumption_resolution,[],[f33929,f264])).
fof(f33929,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,sK9('const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X3,X0))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2135])).
fof(f34460,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X8) = s(X5,sK9(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))) ) | $spl36),
  inference(subsumption_resolution,[],[f33928,f264])).
fof(f33928,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | s(X5,X8) = s(X5,sK9(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2093])).
fof(f34458,plain,(
  ( ! [X4,X2,X3,X1] : (s(X1,X4) = s(X1,sK9(X4,'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3)))),X1))) ) | $spl36),
  inference(subsumption_resolution,[],[f33927,f264])).
fof(f33927,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X4) = s(X1,sK9(X4,'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3)))),X1))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f2092])).
fof(f34454,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK9(X2,X1,X0))) ) | $spl36),
  inference(subsumption_resolution,[],[f33925,f264])).
fof(f33925,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X2) = s(X0,sK9(X2,X1,X0))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f354])).
fof(f34448,plain,(
  ( ! [X2,X0,X1] : (sP0('type.list.list'(fun(X0,bool)),X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X2,X1)) ) | $spl36),
  inference(subsumption_resolution,[],[f33921,f264])).
fof(f33921,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP0('type.list.list'(fun(X0,bool)),X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X2,X1)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4046])).
fof(f34446,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP0(X15,X16,happ(s(fun('type.list.list'(fun(X15,fun(X16,bool))),fun(X15,fun(X16,bool))),'const.list.HD_0'),s('type.list.list'(fun(X15,fun(X16,bool))),X17)),X19,X18)) ) | $spl36),
  inference(subsumption_resolution,[],[f33920,f264])).
fof(f33920,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | sP0(X15,X16,happ(s(fun('type.list.list'(fun(X15,fun(X16,bool))),fun(X15,fun(X16,bool))),'const.list.HD_0'),s('type.list.list'(fun(X15,fun(X16,bool))),X17)),X19,X18)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4049])).
fof(f34444,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (sP0(X138,X139,happ(s(fun('type.list.list'(fun(X138,fun(X139,bool))),fun(X138,fun(X139,bool))),'const.list.EL_1'(s('type.num.num',X140))),s('type.list.list'(fun(X138,fun(X139,bool))),X141)),X143,X142)) ) | $spl36),
  inference(subsumption_resolution,[],[f33919,f264])).
fof(f33919,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (~p(s(bool,'T_0')) | sP0(X138,X139,happ(s(fun('type.list.list'(fun(X138,fun(X139,bool))),fun(X138,fun(X139,bool))),'const.list.EL_1'(s('type.num.num',X140))),s('type.list.list'(fun(X138,fun(X139,bool))),X141)),X143,X142)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4543])).
fof(f34442,plain,(
  ( ! [X344,X346,X343,X345,X347,X342,X348] : (sP0(X342,X343,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X344))),s('type.list.list'(fun(X342,fun(X343,bool))),'const.list.CONS_2'(s(fun(X342,fun(X343,bool)),X345),s('type.list.list'(fun(X342,fun(X343,bool))),X346)))),X348,X347)) ) | $spl36),
  inference(subsumption_resolution,[],[f33918,f264])).
fof(f33918,plain,(
  ( ! [X344,X346,X343,X345,X347,X342,X348] : (~p(s(bool,'T_0')) | sP0(X342,X343,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X344))),s('type.list.list'(fun(X342,fun(X343,bool))),'const.list.CONS_2'(s(fun(X342,fun(X343,bool)),X345),s('type.list.list'(fun(X342,fun(X343,bool))),X346)))),X348,X347)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f11413])).
fof(f34440,plain,(
  ( ! [X210,X208,X206,X204,X209,X207,X205] : (sP0(X204,X205,'const.ind_type.FCONS_3'(s(fun(X204,fun(X205,bool)),X206),s(fun('type.num.num',fun(X204,fun(X205,bool))),X207),s('type.num.num','const.num.SUC_1'(s('type.num.num',X208)))),X210,X209)) ) | $spl36),
  inference(subsumption_resolution,[],[f33917,f264])).
fof(f33917,plain,(
  ( ! [X210,X208,X206,X204,X209,X207,X205] : (~p(s(bool,'T_0')) | sP0(X204,X205,'const.ind_type.FCONS_3'(s(fun(X204,fun(X205,bool)),X206),s(fun('type.num.num',fun(X204,fun(X205,bool))),X207),s('type.num.num','const.num.SUC_1'(s('type.num.num',X208)))),X210,X209)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f5584])).
fof(f34438,plain,(
  ( ! [X14,X12,X15,X13,X16] : (sP0(X12,X14,X15,X16,happ(s(fun('type.list.list'('type.list.list'(X12)),'type.list.list'(X12)),'const.list.HD_0'),s('type.list.list'('type.list.list'(X12)),X13)))) ) | $spl36),
  inference(subsumption_resolution,[],[f33914,f264])).
fof(f33914,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'T_0')) | sP0(X12,X14,X15,X16,happ(s(fun('type.list.list'('type.list.list'(X12)),'type.list.list'(X12)),'const.list.HD_0'),s('type.list.list'('type.list.list'(X12)),X13)))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4052])).
fof(f34436,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (sP0(X154,X157,X158,X159,happ(s(fun('type.list.list'('type.list.list'(X154)),'type.list.list'(X154)),'const.list.EL_1'(s('type.num.num',X155))),s('type.list.list'('type.list.list'(X154)),X156)))) ) | $spl36),
  inference(subsumption_resolution,[],[f33913,f264])).
fof(f33913,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (~p(s(bool,'T_0')) | sP0(X154,X157,X158,X159,happ(s(fun('type.list.list'('type.list.list'(X154)),'type.list.list'(X154)),'const.list.EL_1'(s('type.num.num',X155))),s('type.list.list'('type.list.list'(X154)),X156)))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4546])).
fof(f34434,plain,(
  ( ! [X389,X391,X387,X393,X388,X390,X392] : (sP0(X387,X391,X392,X393,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X388))),s('type.list.list'('type.list.list'(X387)),'const.list.CONS_2'(s('type.list.list'(X387),X389),s('type.list.list'('type.list.list'(X387)),X390)))))) ) | $spl36),
  inference(subsumption_resolution,[],[f33912,f264])).
fof(f33912,plain,(
  ( ! [X389,X391,X387,X393,X388,X390,X392] : (~p(s(bool,'T_0')) | sP0(X387,X391,X392,X393,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X388))),s('type.list.list'('type.list.list'(X387)),'const.list.CONS_2'(s('type.list.list'(X387),X389),s('type.list.list'('type.list.list'(X387)),X390)))))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f11421])).
fof(f34432,plain,(
  ( ! [X231,X229,X233,X230,X228,X234,X232] : (sP0(X228,X232,X233,X234,'const.ind_type.FCONS_3'(s('type.list.list'(X228),X229),s(fun('type.num.num','type.list.list'(X228)),X230),s('type.num.num','const.num.SUC_1'(s('type.num.num',X231)))))) ) | $spl36),
  inference(subsumption_resolution,[],[f33911,f264])).
fof(f33911,plain,(
  ( ! [X231,X229,X233,X230,X228,X234,X232] : (~p(s(bool,'T_0')) | sP0(X228,X232,X233,X234,'const.ind_type.FCONS_3'(s('type.list.list'(X228),X229),s(fun('type.num.num','type.list.list'(X228)),X230),s('type.num.num','const.num.SUC_1'(s('type.num.num',X231)))))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f5588])).
fof(f34430,plain,(
  ( ! [X14,X12,X15,X13,X16] : (sP0(X14,X12,X15,happ(s(fun('type.list.list'('type.list.list'(X12)),'type.list.list'(X12)),'const.list.HD_0'),s('type.list.list'('type.list.list'(X12)),X13)),X16)) ) | $spl36),
  inference(subsumption_resolution,[],[f33910,f264])).
fof(f33910,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'T_0')) | sP0(X14,X12,X15,happ(s(fun('type.list.list'('type.list.list'(X12)),'type.list.list'(X12)),'const.list.HD_0'),s('type.list.list'('type.list.list'(X12)),X13)),X16)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4055])).
fof(f34428,plain,(
  ( ! [X152,X151,X149,X153,X150,X148] : (sP0(X151,X148,X152,happ(s(fun('type.list.list'('type.list.list'(X148)),'type.list.list'(X148)),'const.list.EL_1'(s('type.num.num',X149))),s('type.list.list'('type.list.list'(X148)),X150)),X153)) ) | $spl36),
  inference(subsumption_resolution,[],[f33909,f264])).
fof(f33909,plain,(
  ( ! [X152,X151,X149,X153,X150,X148] : (~p(s(bool,'T_0')) | sP0(X151,X148,X152,happ(s(fun('type.list.list'('type.list.list'(X148)),'type.list.list'(X148)),'const.list.EL_1'(s('type.num.num',X149))),s('type.list.list'('type.list.list'(X148)),X150)),X153)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4545])).
fof(f34426,plain,(
  ( ! [X397,X399,X395,X396,X398,X394,X400] : (sP0(X398,X394,X399,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X395))),s('type.list.list'('type.list.list'(X394)),'const.list.CONS_2'(s('type.list.list'(X394),X396),s('type.list.list'('type.list.list'(X394)),X397)))),X400)) ) | $spl36),
  inference(subsumption_resolution,[],[f33908,f264])).
fof(f33908,plain,(
  ( ! [X397,X399,X395,X396,X398,X394,X400] : (~p(s(bool,'T_0')) | sP0(X398,X394,X399,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X395))),s('type.list.list'('type.list.list'(X394)),'const.list.CONS_2'(s('type.list.list'(X394),X396),s('type.list.list'('type.list.list'(X394)),X397)))),X400)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f11422])).
fof(f34424,plain,(
  ( ! [X222,X227,X225,X223,X221,X226,X224] : (sP0(X225,X221,X226,'const.ind_type.FCONS_3'(s('type.list.list'(X221),X222),s(fun('type.num.num','type.list.list'(X221)),X223),s('type.num.num','const.num.SUC_1'(s('type.num.num',X224)))),X227)) ) | $spl36),
  inference(subsumption_resolution,[],[f33907,f264])).
fof(f33907,plain,(
  ( ! [X222,X227,X225,X223,X221,X226,X224] : (~p(s(bool,'T_0')) | sP0(X225,X221,X226,'const.ind_type.FCONS_3'(s('type.list.list'(X221),X222),s(fun('type.num.num','type.list.list'(X221)),X223),s('type.num.num','const.num.SUC_1'(s('type.num.num',X224)))),X227)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f5587])).
fof(f34422,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP0(X9,X6,X10,'const.ind_type.FCONS_3'(s('type.list.list'(X6),X7),s(fun('type.num.num','type.list.list'(X6)),X8),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl36),
  inference(subsumption_resolution,[],[f33905,f264])).
fof(f33905,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | sP0(X9,X6,X10,'const.ind_type.FCONS_3'(s('type.list.list'(X6),X7),s(fun('type.num.num','type.list.list'(X6)),X8),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4054])).
fof(f34420,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP0(X3,X0,X4,'const.list.HD_1'(s('type.list.list'('type.list.list'(X0)),'const.list.CONS_2'(s('type.list.list'(X0),X1),s('type.list.list'('type.list.list'(X0)),X2)))),X5)) ) | $spl36),
  inference(subsumption_resolution,[],[f33904,f264])).
fof(f33904,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | sP0(X3,X0,X4,'const.list.HD_1'(s('type.list.list'('type.list.list'(X0)),'const.list.CONS_2'(s('type.list.list'(X0),X1),s('type.list.list'('type.list.list'(X0)),X2)))),X5)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4053])).
fof(f34418,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP0(X6,X9,X10,X11,'const.ind_type.FCONS_3'(s('type.list.list'(X6),X7),s(fun('type.num.num','type.list.list'(X6)),X8),s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(subsumption_resolution,[],[f33903,f264])).
fof(f33903,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | sP0(X6,X9,X10,X11,'const.ind_type.FCONS_3'(s('type.list.list'(X6),X7),s(fun('type.num.num','type.list.list'(X6)),X8),s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4051])).
fof(f34416,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP0(X0,X3,X4,X5,'const.list.HD_1'(s('type.list.list'('type.list.list'(X0)),'const.list.CONS_2'(s('type.list.list'(X0),X1),s('type.list.list'('type.list.list'(X0)),X2)))))) ) | $spl36),
  inference(subsumption_resolution,[],[f33902,f264])).
fof(f33902,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | sP0(X0,X3,X4,X5,'const.list.HD_1'(s('type.list.list'('type.list.list'(X0)),'const.list.CONS_2'(s('type.list.list'(X0),X1),s('type.list.list'('type.list.list'(X0)),X2)))))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4050])).
fof(f34414,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (sP0(X9,X10,'const.ind_type.FCONS_3'(s(fun(X9,fun(X10,bool)),X11),s(fun('type.num.num',fun(X9,fun(X10,bool))),X12),s('type.num.num','const.num.|48|_0')),X14,X13)) ) | $spl36),
  inference(subsumption_resolution,[],[f33901,f264])).
fof(f33901,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (~p(s(bool,'T_0')) | sP0(X9,X10,'const.ind_type.FCONS_3'(s(fun(X9,fun(X10,bool)),X11),s(fun('type.num.num',fun(X9,fun(X10,bool))),X12),s('type.num.num','const.num.|48|_0')),X14,X13)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4048])).
fof(f34412,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (sP0(X3,X4,'const.list.HD_1'(s('type.list.list'(fun(X3,fun(X4,bool))),'const.list.CONS_2'(s(fun(X3,fun(X4,bool)),X5),s('type.list.list'(fun(X3,fun(X4,bool))),X6)))),X8,X7)) ) | $spl36),
  inference(subsumption_resolution,[],[f33900,f264])).
fof(f33900,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (~p(s(bool,'T_0')) | sP0(X3,X4,'const.list.HD_1'(s('type.list.list'(fun(X3,fun(X4,bool))),'const.list.CONS_2'(s(fun(X3,fun(X4,bool)),X5),s('type.list.list'(fun(X3,fun(X4,bool))),X6)))),X8,X7)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f4047])).
fof(f34410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X4,X3)) ) | $spl36),
  inference(subsumption_resolution,[],[f33898,f264])).
fof(f33898,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | sP0(X0,X1,X2,X4,X3)) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f395])).
fof(f34408,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7)))) ) | $spl36),
  inference(subsumption_resolution,[],[f34407,f264])).
fof(f34407,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'T_0')) | sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7)))) ) | $spl36),
  inference(forward_demodulation,[],[f34406,f17030])).
fof(f34406,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP0(X0,X1,X2,X3,X4) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7)))) ) | $spl36),
  inference(subsumption_resolution,[],[f33897,f264])).
fof(f33897,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'T_0')) | sP0(X0,X1,X2,X3,X4) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7)))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f394])).
fof(f34402,plain,(
  ( ! [X4,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) ) | $spl36),
  inference(subsumption_resolution,[],[f33893,f264])).
fof(f33893,plain,(
  ( ! [X4,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f381])).
fof(f34127,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(subsumption_resolution,[],[f33395,f17030])).
fof(f33395,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl36),
  inference(backward_demodulation,[],[f17030,f285])).
fof(f5587,plain,(
  ( ! [X222,X227,X225,X223,X221,X226,X224] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X225,fun(X221,bool)),X226),s('type.list.list'(X225),X227),s('type.list.list'(X221),happ(s(fun('type.num.num','type.list.list'(X221)),X223),s('type.num.num',X224)))))) | sP0(X225,X221,X226,'const.ind_type.FCONS_3'(s('type.list.list'(X221),X222),s(fun('type.num.num','type.list.list'(X221)),X223),s('type.num.num','const.num.SUC_1'(s('type.num.num',X224)))),X227)) )),
  inference(superposition,[],[f395,f356])).
fof(f11422,plain,(
  ( ! [X397,X399,X395,X396,X398,X394,X400] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X398,fun(X394,bool)),X399),s('type.list.list'(X398),X400),s('type.list.list'(X394),'const.list.EL_2'(s('type.num.num',X395),s('type.list.list'('type.list.list'(X394)),X397)))))) | sP0(X398,X394,X399,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X395))),s('type.list.list'('type.list.list'(X394)),'const.list.CONS_2'(s('type.list.list'(X394),X396),s('type.list.list'('type.list.list'(X394)),X397)))),X400)) )),
  inference(superposition,[],[f395,f357])).
fof(f4545,plain,(
  ( ! [X152,X151,X149,X153,X150,X148] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X151,fun(X148,bool)),X152),s('type.list.list'(X151),X153),s('type.list.list'(X148),'const.list.EL_2'(s('type.num.num',X149),s('type.list.list'('type.list.list'(X148)),X150)))))) | sP0(X151,X148,X152,happ(s(fun('type.list.list'('type.list.list'(X148)),'type.list.list'(X148)),'const.list.EL_1'(s('type.num.num',X149))),s('type.list.list'('type.list.list'(X148)),X150)),X153)) )),
  inference(superposition,[],[f395,f346])).
fof(f4055,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X14,fun(X12,bool)),X15),s('type.list.list'(X14),X16),s('type.list.list'(X12),'const.list.HD_1'(s('type.list.list'('type.list.list'(X12)),X13)))))) | sP0(X14,X12,X15,happ(s(fun('type.list.list'('type.list.list'(X12)),'type.list.list'(X12)),'const.list.HD_0'),s('type.list.list'('type.list.list'(X12)),X13)),X16)) )),
  inference(superposition,[],[f395,f338])).
fof(f5588,plain,(
  ( ! [X231,X229,X233,X230,X228,X234,X232] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X228,fun(X232,bool)),X233),s('type.list.list'(X228),happ(s(fun('type.num.num','type.list.list'(X228)),X230),s('type.num.num',X231))),s('type.list.list'(X232),X234)))) | sP0(X228,X232,X233,X234,'const.ind_type.FCONS_3'(s('type.list.list'(X228),X229),s(fun('type.num.num','type.list.list'(X228)),X230),s('type.num.num','const.num.SUC_1'(s('type.num.num',X231)))))) )),
  inference(superposition,[],[f395,f356])).
fof(f11421,plain,(
  ( ! [X389,X391,X387,X393,X388,X390,X392] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X387,fun(X391,bool)),X392),s('type.list.list'(X387),'const.list.EL_2'(s('type.num.num',X388),s('type.list.list'('type.list.list'(X387)),X390))),s('type.list.list'(X391),X393)))) | sP0(X387,X391,X392,X393,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X388))),s('type.list.list'('type.list.list'(X387)),'const.list.CONS_2'(s('type.list.list'(X387),X389),s('type.list.list'('type.list.list'(X387)),X390)))))) )),
  inference(superposition,[],[f395,f357])).
fof(f4546,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X154,fun(X157,bool)),X158),s('type.list.list'(X154),'const.list.EL_2'(s('type.num.num',X155),s('type.list.list'('type.list.list'(X154)),X156))),s('type.list.list'(X157),X159)))) | sP0(X154,X157,X158,X159,happ(s(fun('type.list.list'('type.list.list'(X154)),'type.list.list'(X154)),'const.list.EL_1'(s('type.num.num',X155))),s('type.list.list'('type.list.list'(X154)),X156)))) )),
  inference(superposition,[],[f395,f346])).
fof(f4052,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X12,fun(X14,bool)),X15),s('type.list.list'(X12),'const.list.HD_1'(s('type.list.list'('type.list.list'(X12)),X13))),s('type.list.list'(X14),X16)))) | sP0(X12,X14,X15,X16,happ(s(fun('type.list.list'('type.list.list'(X12)),'type.list.list'(X12)),'const.list.HD_0'),s('type.list.list'('type.list.list'(X12)),X13)))) )),
  inference(superposition,[],[f395,f338])).
fof(f373,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),sK13(X3,X2,X1,X0)),s('type.list.list'(X1),sK14(X3,X2,X1,X0))))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK11(X3,X2,X1,X0)))),s(X1,sK12(X3,X2,X1,X0))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),sK13(X3,X2,X1,X0)),s('type.list.list'(X1),sK14(X3,X2,X1,X0))))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),sK13(X3,X2,X1,X0)))),s('type.list.list'(X1),sK14(X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK11(X3,X2,X1,X0)),s('type.list.list'(X0),sK13(X3,X2,X1,X0)))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK12(X3,X2,X1,X0)),s('type.list.list'(X1),sK14(X3,X2,X1,X0)))))))) | ! [X4,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13,sK14])],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X6,X7,X8,X9] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9)))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X8))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X7),s('type.list.list'(X1),X9))))))) | ! [X4,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))))),
  inference(rectify,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X4,X5,X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X6),s('type.list.list'(X1),X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7)))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7))))))) | ! [X8,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X6),s('type.list.list'(X1),X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) | ! [X8,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(ennf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) & ! [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X6),s('type.list.list'(X1),X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) => ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ! [X2,X3] : ((p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) & ! [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X6),s('type.list.list'(X1),X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) => ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6,X7] : ! [X9,X24] : ((p(s(bool,happ(s(fun('type.list.list'(X7),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X7),bool)),X24),s('type.list.list'(X6),'const.list.NIL_0'))),s('type.list.list'(X7),'const.list.NIL_0')))) & ! [X22,X23,X14,X15] : ((p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X9),s(X6,X22))),s(X7,X23)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),X14),s('type.list.list'(X7),X15)))) & p(s(bool,happ(s(fun('type.list.list'(X7),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X7),bool)),X24),s('type.list.list'(X6),X14))),s('type.list.list'(X7),X15))))) => p(s(bool,happ(s(fun('type.list.list'(X7),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X7),bool)),X24),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X22),s('type.list.list'(X6),X14))))),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X23),s('type.list.list'(X7),X15)))))))) => ! [X20,X21] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),X20),s('type.list.list'(X7),X21)))) => p(s(bool,happ(s(fun('type.list.list'(X7),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X7),bool)),X24),s('type.list.list'(X6),X20))),s('type.list.list'(X7),X21))))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LIST_REL_strongind')).
fof(f391,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),sK23(X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK24(X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f392,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),sK23(X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK24(X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f5584,plain,(
  ( ! [X210,X208,X206,X204,X209,X207,X205] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X204,fun(X205,bool)),happ(s(fun('type.num.num',fun(X204,fun(X205,bool))),X207),s('type.num.num',X208))),s('type.list.list'(X204),X209),s('type.list.list'(X205),X210)))) | sP0(X204,X205,'const.ind_type.FCONS_3'(s(fun(X204,fun(X205,bool)),X206),s(fun('type.num.num',fun(X204,fun(X205,bool))),X207),s('type.num.num','const.num.SUC_1'(s('type.num.num',X208)))),X210,X209)) )),
  inference(superposition,[],[f395,f356])).
fof(f11413,plain,(
  ( ! [X344,X346,X343,X345,X347,X342,X348] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X342,fun(X343,bool)),'const.list.EL_2'(s('type.num.num',X344),s('type.list.list'(fun(X342,fun(X343,bool))),X346))),s('type.list.list'(X342),X347),s('type.list.list'(X343),X348)))) | sP0(X342,X343,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X344))),s('type.list.list'(fun(X342,fun(X343,bool))),'const.list.CONS_2'(s(fun(X342,fun(X343,bool)),X345),s('type.list.list'(fun(X342,fun(X343,bool))),X346)))),X348,X347)) )),
  inference(superposition,[],[f395,f357])).
fof(f4543,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X138,fun(X139,bool)),'const.list.EL_2'(s('type.num.num',X140),s('type.list.list'(fun(X138,fun(X139,bool))),X141))),s('type.list.list'(X138),X142),s('type.list.list'(X139),X143)))) | sP0(X138,X139,happ(s(fun('type.list.list'(fun(X138,fun(X139,bool))),fun(X138,fun(X139,bool))),'const.list.EL_1'(s('type.num.num',X140))),s('type.list.list'(fun(X138,fun(X139,bool))),X141)),X143,X142)) )),
  inference(superposition,[],[f395,f346])).
fof(f4049,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X15,fun(X16,bool)),'const.list.HD_1'(s('type.list.list'(fun(X15,fun(X16,bool))),X17))),s('type.list.list'(X15),X18),s('type.list.list'(X16),X19)))) | sP0(X15,X16,happ(s(fun('type.list.list'(fun(X15,fun(X16,bool))),fun(X15,fun(X16,bool))),'const.list.HD_0'),s('type.list.list'(fun(X15,fun(X16,bool))),X17)),X19,X18)) )),
  inference(superposition,[],[f395,f338])).
fof(f4046,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.list.list'(fun(X0,bool)),fun(X0,bool)),'const.list.HD_0'),s('type.list.list'('type.list.list'(fun(X0,bool))),X1),s('type.list.list'(X0),X2)))) | sP0('type.list.list'(fun(X0,bool)),X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X2,X1)) )),
  inference(superposition,[],[f395,f275])).
fof(f5619,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8))) = s(X5,sK9('const.ind_type.FCONS_3'(s(X5,X6),s(fun('type.num.num',X5),X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))),X9,X5)) | ~p(s(bool,happ(s(fun(X5,bool),X9),s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8))))))) )),
  inference(forward_demodulation,[],[f5530,f356])).
fof(f5530,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X9),s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8)))))) | s(X5,'const.ind_type.FCONS_3'(s(X5,X6),s(fun('type.num.num',X5),X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) = s(X5,sK9('const.ind_type.FCONS_3'(s(X5,X6),s(fun('type.num.num',X5),X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))),X9,X5))) )),
  inference(superposition,[],[f354,f356])).
fof(f11461,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(X21,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(X21),X24))) = s(X21,sK9('const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.list.list'(X21),'const.list.CONS_2'(s(X21,X23),s('type.list.list'(X21),X24)))),X25,X21)) | ~p(s(bool,happ(s(fun(X21,bool),X25),s(X21,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(X21),X24))))))) )),
  inference(forward_demodulation,[],[f11329,f357])).
fof(f11329,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(X21,bool),X25),s(X21,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(X21),X24)))))) | s(X21,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.list.list'(X21),'const.list.CONS_2'(s(X21,X23),s('type.list.list'(X21),X24))))) = s(X21,sK9('const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.list.list'(X21),'const.list.CONS_2'(s(X21,X23),s('type.list.list'(X21),X24)))),X25,X21))) )),
  inference(superposition,[],[f354,f357])).
fof(f4569,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X4),X6))) = s(X4,sK9(happ(s(fun('type.list.list'(X4),X4),'const.list.EL_1'(s('type.num.num',X5))),s('type.list.list'(X4),X6)),X7,X4)) | ~p(s(bool,happ(s(fun(X4,bool),X7),s(X4,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X4),X6))))))) )),
  inference(forward_demodulation,[],[f4493,f346])).
fof(f4493,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),X7),s(X4,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X4),X6)))))) | s(X4,happ(s(fun('type.list.list'(X4),X4),'const.list.EL_1'(s('type.num.num',X5))),s('type.list.list'(X4),X6))) = s(X4,sK9(happ(s(fun('type.list.list'(X4),X4),'const.list.EL_1'(s('type.num.num',X5))),s('type.list.list'(X4),X6)),X7,X4))) )),
  inference(superposition,[],[f354,f346])).
fof(f2535,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK6(X4,'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3)))),X1)))))) )),
  inference(forward_demodulation,[],[f2510,f343])).
fof(f2510,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK6(X4,'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3)))),X1))))) | p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3))))))))) )),
  inference(superposition,[],[f348,f343])).
fof(f2536,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6)))) | ~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK6(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5)))))) )),
  inference(forward_demodulation,[],[f2511,f344])).
fof(f2511,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK6(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))))) | p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f348,f344])).
fof(f5285,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X8)))) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK9(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5)))))) )),
  inference(forward_demodulation,[],[f5250,f344])).
fof(f5250,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK9(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))))) | ~p(s(bool,happ(s(fun(X5,bool),'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0'))),s(X5,X8))))) )),
  inference(superposition,[],[f355,f344])).
fof(f5284,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X4)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK9(X4,'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3)))),X1)))))) )),
  inference(forward_demodulation,[],[f5249,f343])).
fof(f5249,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK9(X4,'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3)))),X1))))) | ~p(s(bool,happ(s(fun(X1,bool),'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3))))),s(X1,X4))))) )),
  inference(superposition,[],[f355,f343])).
fof(f394,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))) | sP0(X0,X1,X2,X3,X4) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7)))) )),
  inference(cnf_transformation,[],[f260])).
fof(f382,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(X4,X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK20(X4,X3,X2,X1,X0)),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',sK20(X4,X3,X2,X1,X0)),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(rectify,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))))) & ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(ennf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X7] : ! [X9,X25,X26] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),X25),s('type.list.list'(X7),X26)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X25))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),X26))) & ! [X27] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X25)))))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X9),s(X6,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(X6),X25))))),s(X7,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(X7),X26)))))))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LIST_REL_EL_EQN')).
fof(f384,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK20(X4,X3,X2,X1,X0)),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',sK20(X4,X3,X2,X1,X0)),s('type.list.list'(X1),X4)))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f372,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK11(X3,X2,X1,X0)))),s(X1,sK12(X3,X2,X1,X0))))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f376,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK15(X3,X2,X1,X0)))),s(X1,sK16(X3,X2,X1,X0))))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK15(X3,X2,X1,X0)))),s(X1,sK16(X3,X2,X1,X0))))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),sK17(X3,X2,X1,X0)))),s('type.list.list'(X1),sK18(X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK15(X3,X2,X1,X0)),s('type.list.list'(X0),sK17(X3,X2,X1,X0)))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK16(X3,X2,X1,X0)),s('type.list.list'(X1),sK18(X3,X2,X1,X0)))))))) | ! [X4,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16,sK17,sK18])],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X6,X7,X8,X9] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9)))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X8))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X7),s('type.list.list'(X1),X9))))))) | ! [X4,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))))),
  inference(rectify,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X4,X5,X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7)))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7))))))) | ! [X8,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) | ! [X8,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) & ! [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) => ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ! [X2,X3] : ((p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) & ! [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) => ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6,X7] : ! [X9,X24] : ((p(s(bool,happ(s(fun('type.list.list'(X7),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X7),bool)),X24),s('type.list.list'(X6),'const.list.NIL_0'))),s('type.list.list'(X7),'const.list.NIL_0')))) & ! [X22,X23,X14,X15] : ((p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X9),s(X6,X22))),s(X7,X23)))) & p(s(bool,happ(s(fun('type.list.list'(X7),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X7),bool)),X24),s('type.list.list'(X6),X14))),s('type.list.list'(X7),X15))))) => p(s(bool,happ(s(fun('type.list.list'(X7),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X7),bool)),X24),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X22),s('type.list.list'(X6),X14))))),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X23),s('type.list.list'(X7),X15)))))))) => ! [X20,X21] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),X20),s('type.list.list'(X7),X21)))) => p(s(bool,happ(s(fun('type.list.list'(X7),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X7),bool)),X24),s('type.list.list'(X6),X20))),s('type.list.list'(X7),X21))))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LIST_REL_ind')).
fof(f389,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK21(X4,X3,X2,X1,X0)))),s(X1,sK22(X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f390,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK21(X4,X3,X2,X1,X0)))),s(X1,sK22(X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f5664,plain,(
  ( ! [X187,X185,X188,X186,X184] : (~p(s(bool,happ(s(fun(X184,bool),happ(s(fun('type.num.num',fun(X184,bool)),X186),s('type.num.num',X187))),s(X184,X188)))) | p(s(bool,happ(s(fun(X184,bool),happ(s(fun('type.num.num',fun(X184,bool)),X186),s('type.num.num',X187))),s(X184,sK9(X188,'const.ind_type.FCONS_3'(s(fun(X184,bool),X185),s(fun('type.num.num',fun(X184,bool)),X186),s('type.num.num','const.num.SUC_1'(s('type.num.num',X187)))),X184)))))) )),
  inference(forward_demodulation,[],[f5580,f356])).
fof(f5580,plain,(
  ( ! [X187,X185,X188,X186,X184] : (p(s(bool,happ(s(fun(X184,bool),happ(s(fun('type.num.num',fun(X184,bool)),X186),s('type.num.num',X187))),s(X184,sK9(X188,'const.ind_type.FCONS_3'(s(fun(X184,bool),X185),s(fun('type.num.num',fun(X184,bool)),X186),s('type.num.num','const.num.SUC_1'(s('type.num.num',X187)))),X184))))) | ~p(s(bool,happ(s(fun(X184,bool),'const.ind_type.FCONS_3'(s(fun(X184,bool),X185),s(fun('type.num.num',fun(X184,bool)),X186),s('type.num.num','const.num.SUC_1'(s('type.num.num',X187))))),s(X184,X188))))) )),
  inference(superposition,[],[f355,f356])).
fof(f5582,plain,(
  ( ! [X198,X196,X194,X197,X195] : (~p(s(bool,happ(s(fun(X194,bool),happ(s(fun('type.num.num',fun(X194,bool)),X196),s('type.num.num',X197))),s(X194,X198)))) | s(X194,X198) = s(X194,sK9(X198,'const.ind_type.FCONS_3'(s(fun(X194,bool),X195),s(fun('type.num.num',fun(X194,bool)),X196),s('type.num.num','const.num.SUC_1'(s('type.num.num',X197)))),X194))) )),
  inference(superposition,[],[f354,f356])).
fof(f5665,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,'|dollar|exists_1'(s(fun(X189,bool),happ(s(fun('type.num.num',fun(X189,bool)),X191),s('type.num.num',X192)))))) | ~p(s(bool,happ(s(fun(X189,bool),happ(s(fun('type.num.num',fun(X189,bool)),X191),s('type.num.num',X192))),s(X189,sK6(X193,'const.ind_type.FCONS_3'(s(fun(X189,bool),X190),s(fun('type.num.num',fun(X189,bool)),X191),s('type.num.num','const.num.SUC_1'(s('type.num.num',X192)))),X189)))))) )),
  inference(forward_demodulation,[],[f5581,f356])).
fof(f5581,plain,(
  ( ! [X191,X189,X192,X190,X193] : (~p(s(bool,happ(s(fun(X189,bool),happ(s(fun('type.num.num',fun(X189,bool)),X191),s('type.num.num',X192))),s(X189,sK6(X193,'const.ind_type.FCONS_3'(s(fun(X189,bool),X190),s(fun('type.num.num',fun(X189,bool)),X191),s('type.num.num','const.num.SUC_1'(s('type.num.num',X192)))),X189))))) | p(s(bool,'|dollar|exists_1'(s(fun(X189,bool),'const.ind_type.FCONS_3'(s(fun(X189,bool),X190),s(fun('type.num.num',fun(X189,bool)),X191),s('type.num.num','const.num.SUC_1'(s('type.num.num',X192))))))))) )),
  inference(superposition,[],[f348,f356])).
fof(f11411,plain,(
  ( ! [X333,X335,X332,X334,X336] : (~p(s(bool,happ(s(fun(X332,bool),'const.list.EL_2'(s('type.num.num',X333),s('type.list.list'(fun(X332,bool)),X335))),s(X332,X336)))) | s(X332,X336) = s(X332,sK9(X336,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X333))),s('type.list.list'(fun(X332,bool)),'const.list.CONS_2'(s(fun(X332,bool),X334),s('type.list.list'(fun(X332,bool)),X335)))),X332))) )),
  inference(superposition,[],[f354,f357])).
fof(f11532,plain,(
  ( ! [X325,X323,X324,X326,X322] : (~p(s(bool,happ(s(fun(X322,bool),'const.list.EL_2'(s('type.num.num',X323),s('type.list.list'(fun(X322,bool)),X325))),s(X322,X326)))) | p(s(bool,happ(s(fun(X322,bool),'const.list.EL_2'(s('type.num.num',X323),s('type.list.list'(fun(X322,bool)),X325))),s(X322,sK9(X326,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X323))),s('type.list.list'(fun(X322,bool)),'const.list.CONS_2'(s(fun(X322,bool),X324),s('type.list.list'(fun(X322,bool)),X325)))),X322)))))) )),
  inference(forward_demodulation,[],[f11409,f357])).
fof(f11409,plain,(
  ( ! [X325,X323,X324,X326,X322] : (p(s(bool,happ(s(fun(X322,bool),'const.list.EL_2'(s('type.num.num',X323),s('type.list.list'(fun(X322,bool)),X325))),s(X322,sK9(X326,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X323))),s('type.list.list'(fun(X322,bool)),'const.list.CONS_2'(s(fun(X322,bool),X324),s('type.list.list'(fun(X322,bool)),X325)))),X322))))) | ~p(s(bool,happ(s(fun(X322,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X323))),s('type.list.list'(fun(X322,bool)),'const.list.CONS_2'(s(fun(X322,bool),X324),s('type.list.list'(fun(X322,bool)),X325))))),s(X322,X326))))) )),
  inference(superposition,[],[f355,f357])).
fof(f5287,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),'const.list.EL_2'(s('type.num.num',X13),s('type.list.list'(fun(X12,bool)),X14))),s(X12,X15)))) | p(s(bool,happ(s(fun(X12,bool),'const.list.EL_2'(s('type.num.num',X13),s('type.list.list'(fun(X12,bool)),X14))),s(X12,sK9(X15,happ(s(fun('type.list.list'(fun(X12,bool)),fun(X12,bool)),'const.list.EL_1'(s('type.num.num',X13))),s('type.list.list'(fun(X12,bool)),X14)),X12)))))) )),
  inference(forward_demodulation,[],[f5252,f346])).
fof(f5252,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(X12,bool),'const.list.EL_2'(s('type.num.num',X13),s('type.list.list'(fun(X12,bool)),X14))),s(X12,sK9(X15,happ(s(fun('type.list.list'(fun(X12,bool)),fun(X12,bool)),'const.list.EL_1'(s('type.num.num',X13))),s('type.list.list'(fun(X12,bool)),X14)),X12))))) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.list.list'(fun(X12,bool)),fun(X12,bool)),'const.list.EL_1'(s('type.num.num',X13))),s('type.list.list'(fun(X12,bool)),X14))),s(X12,X15))))) )),
  inference(superposition,[],[f355,f346])).
fof(f4541,plain,(
  ( ! [X132,X130,X133,X131] : (~p(s(bool,happ(s(fun(X130,bool),'const.list.EL_2'(s('type.num.num',X131),s('type.list.list'(fun(X130,bool)),X132))),s(X130,X133)))) | s(X130,X133) = s(X130,sK9(X133,happ(s(fun('type.list.list'(fun(X130,bool)),fun(X130,bool)),'const.list.EL_1'(s('type.num.num',X131))),s('type.list.list'(fun(X130,bool)),X132)),X130))) )),
  inference(superposition,[],[f354,f346])).
fof(f11533,plain,(
  ( ! [X327,X329,X331,X328,X330] : (p(s(bool,'|dollar|exists_1'(s(fun(X327,bool),'const.list.EL_2'(s('type.num.num',X328),s('type.list.list'(fun(X327,bool)),X330)))))) | ~p(s(bool,happ(s(fun(X327,bool),'const.list.EL_2'(s('type.num.num',X328),s('type.list.list'(fun(X327,bool)),X330))),s(X327,sK6(X331,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X328))),s('type.list.list'(fun(X327,bool)),'const.list.CONS_2'(s(fun(X327,bool),X329),s('type.list.list'(fun(X327,bool)),X330)))),X327)))))) )),
  inference(forward_demodulation,[],[f11410,f357])).
fof(f11410,plain,(
  ( ! [X327,X329,X331,X328,X330] : (~p(s(bool,happ(s(fun(X327,bool),'const.list.EL_2'(s('type.num.num',X328),s('type.list.list'(fun(X327,bool)),X330))),s(X327,sK6(X331,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X328))),s('type.list.list'(fun(X327,bool)),'const.list.CONS_2'(s(fun(X327,bool),X329),s('type.list.list'(fun(X327,bool)),X330)))),X327))))) | p(s(bool,'|dollar|exists_1'(s(fun(X327,bool),'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X328))),s('type.list.list'(fun(X327,bool)),'const.list.CONS_2'(s(fun(X327,bool),X329),s('type.list.list'(fun(X327,bool)),X330))))))))) )),
  inference(superposition,[],[f348,f357])).
fof(f4607,plain,(
  ( ! [X127,X128,X126,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X126,bool),'const.list.EL_2'(s('type.num.num',X127),s('type.list.list'(fun(X126,bool)),X128)))))) | ~p(s(bool,happ(s(fun(X126,bool),'const.list.EL_2'(s('type.num.num',X127),s('type.list.list'(fun(X126,bool)),X128))),s(X126,sK6(X129,happ(s(fun('type.list.list'(fun(X126,bool)),fun(X126,bool)),'const.list.EL_1'(s('type.num.num',X127))),s('type.list.list'(fun(X126,bool)),X128)),X126)))))) )),
  inference(forward_demodulation,[],[f4540,f346])).
fof(f4540,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,happ(s(fun(X126,bool),'const.list.EL_2'(s('type.num.num',X127),s('type.list.list'(fun(X126,bool)),X128))),s(X126,sK6(X129,happ(s(fun('type.list.list'(fun(X126,bool)),fun(X126,bool)),'const.list.EL_1'(s('type.num.num',X127))),s('type.list.list'(fun(X126,bool)),X128)),X126))))) | p(s(bool,'|dollar|exists_1'(s(fun(X126,bool),happ(s(fun('type.list.list'(fun(X126,bool)),fun(X126,bool)),'const.list.EL_1'(s('type.num.num',X127))),s('type.list.list'(fun(X126,bool)),X128))))))) )),
  inference(superposition,[],[f348,f346])).
fof(f5286,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),X10))),s(X9,X11)))) | p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),X10))),s(X9,sK9(X11,happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10)),X9)))))) )),
  inference(forward_demodulation,[],[f5251,f338])).
fof(f5251,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),X10))),s(X9,sK9(X11,happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10)),X9))))) | ~p(s(bool,happ(s(fun(X9,bool),happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10))),s(X9,X11))))) )),
  inference(superposition,[],[f355,f338])).
fof(f2537,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),X10)))))) | ~p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),X10))),s(X9,sK6(X11,happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10)),X9)))))) )),
  inference(forward_demodulation,[],[f2512,f338])).
fof(f2512,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),X10))),s(X9,sK6(X11,happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10)),X9))))) | p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10))))))) )),
  inference(superposition,[],[f348,f338])).
fof(f374,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),sK13(X3,X2,X1,X0)))),s('type.list.list'(X1),sK14(X3,X2,X1,X0))))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f375,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK11(X3,X2,X1,X0)),s('type.list.list'(X0),sK13(X3,X2,X1,X0)))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK12(X3,X2,X1,X0)),s('type.list.list'(X1),sK14(X3,X2,X1,X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f377,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),sK17(X3,X2,X1,X0)))),s('type.list.list'(X1),sK18(X3,X2,X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f378,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK15(X3,X2,X1,X0)),s('type.list.list'(X0),sK17(X3,X2,X1,X0)))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK16(X3,X2,X1,X0)),s('type.list.list'(X1),sK18(X3,X2,X1,X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f3277,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X31),X33))))) = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(X31),'const.list.CONS_2'(s(X31,X32),s('type.list.list'(X31),X33)))),X34,'type.num.num')) | ~p(s(bool,happ(s(fun('type.num.num',bool),X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X31),X33))))))))) )),
  inference(forward_demodulation,[],[f3266,f345])).
fof(f3266,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun('type.num.num',bool),X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X31),X33)))))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X31),'const.list.CONS_2'(s(X31,X32),s('type.list.list'(X31),X33))))) = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(X31),'const.list.CONS_2'(s(X31,X32),s('type.list.list'(X31),X33)))),X34,'type.num.num'))) )),
  inference(superposition,[],[f354,f345])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(X4,X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f5583,plain,(
  ( ! [X202,X200,X199,X203,X201] : (p(s(bool,'|dollar|exists_1'(s(fun(X199,bool),happ(s(fun('type.num.num',fun(X199,bool)),X201),s('type.num.num',X202)))))) | s(X199,X203) = s(X199,sK6(X203,'const.ind_type.FCONS_3'(s(fun(X199,bool),X200),s(fun('type.num.num',fun(X199,bool)),X201),s('type.num.num','const.num.SUC_1'(s('type.num.num',X202)))),X199))) )),
  inference(superposition,[],[f347,f356])).
fof(f6116,plain,(
  ( ! [X288,X285,X287,X289,X284,X286] : (s(bool,'F_0') = s(bool,sK9('const.list.LIST_REL_3'(s(fun(X284,fun(X285,bool)),X286),s('type.list.list'(X284),'const.list.NIL_0'),s('type.list.list'(X285),'const.list.CONS_2'(s(X285,X287),s('type.list.list'(X285),X288)))),X289,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X289),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f6053,f380])).
fof(f6053,plain,(
  ( ! [X288,X285,X287,X289,X284,X286] : (~p(s(bool,happ(s(fun(bool,bool),X289),s(bool,'F_0')))) | s(bool,'const.list.LIST_REL_3'(s(fun(X284,fun(X285,bool)),X286),s('type.list.list'(X284),'const.list.NIL_0'),s('type.list.list'(X285),'const.list.CONS_2'(s(X285,X287),s('type.list.list'(X285),X288))))) = s(bool,sK9('const.list.LIST_REL_3'(s(fun(X284,fun(X285,bool)),X286),s('type.list.list'(X284),'const.list.NIL_0'),s('type.list.list'(X285),'const.list.CONS_2'(s(X285,X287),s('type.list.list'(X285),X288)))),X289,bool))) )),
  inference(superposition,[],[f354,f380])).
fof(f2366,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') = s(bool,sK9('const.list.LIST_REL_3'(s(fun(X92,fun(X93,bool)),X94),s('type.list.list'(X92),'const.list.NIL_0'),s('type.list.list'(X93),'const.list.NIL_0')),X95,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f2354,f368])).
fof(f2354,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0')))) | s(bool,'const.list.LIST_REL_3'(s(fun(X92,fun(X93,bool)),X94),s('type.list.list'(X92),'const.list.NIL_0'),s('type.list.list'(X93),'const.list.NIL_0'))) = s(bool,sK9('const.list.LIST_REL_3'(s(fun(X92,fun(X93,bool)),X94),s('type.list.list'(X92),'const.list.NIL_0'),s('type.list.list'(X93),'const.list.NIL_0')),X95,bool))) )),
  inference(superposition,[],[f354,f368])).
fof(f3884,plain,(
  ( ! [X97,X98,X96] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X96),s('type.num.num',X97))) = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X96))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X97)))),X98,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,'const.prim_rec.<_2'(s('type.num.num',X96),s('type.num.num',X97))))))) )),
  inference(forward_demodulation,[],[f3843,f360])).
fof(f3843,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,'const.prim_rec.<_2'(s('type.num.num',X96),s('type.num.num',X97)))))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X96))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X97))))) = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X96))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X97)))),X98,bool))) )),
  inference(superposition,[],[f354,f360])).
fof(f1485,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),'const.list.HD_1'(s('type.list.list'(fun(X39,bool)),X40)))))) | s(X39,X41) = s(X39,sK6(X41,happ(s(fun('type.list.list'(fun(X39,bool)),fun(X39,bool)),'const.list.HD_0'),s('type.list.list'(fun(X39,bool)),X40)),X39))) )),
  inference(superposition,[],[f347,f338])).
fof(f4542,plain,(
  ( ! [X134,X136,X135,X137] : (p(s(bool,'|dollar|exists_1'(s(fun(X134,bool),'const.list.EL_2'(s('type.num.num',X135),s('type.list.list'(fun(X134,bool)),X136)))))) | s(X134,X137) = s(X134,sK6(X137,happ(s(fun('type.list.list'(fun(X134,bool)),fun(X134,bool)),'const.list.EL_1'(s('type.num.num',X135))),s('type.list.list'(fun(X134,bool)),X136)),X134))) )),
  inference(superposition,[],[f347,f346])).
fof(f11412,plain,(
  ( ! [X341,X337,X339,X340,X338] : (p(s(bool,'|dollar|exists_1'(s(fun(X337,bool),'const.list.EL_2'(s('type.num.num',X338),s('type.list.list'(fun(X337,bool)),X340)))))) | s(X337,X341) = s(X337,sK6(X341,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X338))),s('type.list.list'(fun(X337,bool)),'const.list.CONS_2'(s(fun(X337,bool),X339),s('type.list.list'(fun(X337,bool)),X340)))),X337))) )),
  inference(superposition,[],[f347,f357])).
fof(f794,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6)))) | s(X5,X8) = s(X5,sK6(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))) )),
  inference(superposition,[],[f347,f344])).
fof(f793,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) | s(X1,X4) = s(X1,sK6(X4,'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3)))),X1))) )),
  inference(superposition,[],[f347,f343])).
fof(f2513,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK6(X0,X1,bool))) )),
  inference(superposition,[],[f348,f276])).
fof(f4497,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.list.EL_2'(s('type.num.num',X18),s('type.list.list'(bool),sK6(X19,'const.list.EL_1'(s('type.num.num',X18)),'type.list.list'(bool)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X18))))))) )),
  inference(superposition,[],[f348,f346])).
fof(f350,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f241])).
fof(f353,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK9(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X17,X28] : (? [X8] : (s(X6,X8) = s(X6,X28) & p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X8))))) <=> p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X28)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.bool.UNWIND_THM2')).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK9(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f2092,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X4)))) | s(X1,X4) = s(X1,sK9(X4,'const.list.HD_1'(s('type.list.list'(fun(X1,bool)),'const.list.CONS_2'(s(fun(X1,bool),X2),s('type.list.list'(fun(X1,bool)),X3)))),X1))) )),
  inference(superposition,[],[f354,f343])).
fof(f2093,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X8)))) | s(X5,X8) = s(X5,sK9(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))) )),
  inference(superposition,[],[f354,f344])).
fof(f2135,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,sK9('const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X3,X0)) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f2095,f343])).
fof(f2095,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X1)))) | s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s(X0,sK9('const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))),X3,X0))) )),
  inference(superposition,[],[f354,f343])).
fof(f2136,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,X5) = s(X4,sK9('const.ind_type.FCONS_3'(s(X4,X5),s(fun('type.num.num',X4),X6),s('type.num.num','const.num.|48|_0')),X7,X4)) | ~p(s(bool,happ(s(fun(X4,bool),X7),s(X4,X5))))) )),
  inference(forward_demodulation,[],[f2096,f344])).
fof(f2096,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),X7),s(X4,X5)))) | s(X4,'const.ind_type.FCONS_3'(s(X4,X5),s(fun('type.num.num',X4),X6),s('type.num.num','const.num.|48|_0'))) = s(X4,sK9('const.ind_type.FCONS_3'(s(X4,X5),s(fun('type.num.num',X4),X6),s('type.num.num','const.num.|48|_0')),X7,X4))) )),
  inference(superposition,[],[f354,f344])).
fof(f5299,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X2,X1,X0))))) )),
  inference(subsumption_resolution,[],[f5266,f426])).
fof(f5266,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X2,X1,X0))))) )),
  inference(superposition,[],[f355,f276])).
fof(f2137,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.list.HD_1'(s('type.list.list'(X8),X9))) = s(X8,sK9(happ(s(fun('type.list.list'(X8),X8),'const.list.HD_0'),s('type.list.list'(X8),X9)),X10,X8)) | ~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,'const.list.HD_1'(s('type.list.list'(X8),X9))))))) )),
  inference(forward_demodulation,[],[f2097,f338])).
fof(f2097,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,'const.list.HD_1'(s('type.list.list'(X8),X9)))))) | s(X8,happ(s(fun('type.list.list'(X8),X8),'const.list.HD_0'),s('type.list.list'(X8),X9))) = s(X8,sK9(happ(s(fun('type.list.list'(X8),X8),'const.list.HD_0'),s('type.list.list'(X8),X9)),X10,X8))) )),
  inference(superposition,[],[f354,f338])).
fof(f2094,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),'const.list.HD_1'(s('type.list.list'(fun(X9,bool)),X10))),s(X9,X11)))) | s(X9,X11) = s(X9,sK9(X11,happ(s(fun('type.list.list'(fun(X9,bool)),fun(X9,bool)),'const.list.HD_0'),s('type.list.list'(fun(X9,bool)),X10)),X9))) )),
  inference(superposition,[],[f354,f338])).
fof(f2110,plain,(
  ( ! [X45,X46] : (~p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X45)))) | s(bool,'F_0') = s(bool,sK9('F_0',X46,bool)) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f354,f276])).
fof(f2150,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,sK9('const.marker.Cong_1'(s(bool,X60)),X61,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X61),s(bool,X60))))) )),
  inference(forward_demodulation,[],[f2118,f271])).
fof(f2118,plain,(
  ( ! [X61,X60] : (~p(s(bool,happ(s(fun(bool,bool),X61),s(bool,X60)))) | s(bool,'const.marker.Cong_1'(s(bool,X60))) = s(bool,sK9('const.marker.Cong_1'(s(bool,X60)),X61,bool))) )),
  inference(superposition,[],[f354,f271])).
fof(f5253,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,'T_0') = s(bool,sK9(X0,X1,bool))) )),
  inference(superposition,[],[f355,f276])).
fof(f10688,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X27)))) | s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f353,f276])).
fof(f10707,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,happ(s(fun(bool,bool),X95),s(bool,X94)))) | p(s(bool,happ(s(fun(bool,bool),X95),s(bool,X96)))) | s(bool,'F_0') != s(bool,X96) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f353,f276])).
fof(f2146,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X53,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X53),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f2114,f479])).
fof(f2114,plain,(
  ( ! [X52,X53] : (~p(s(bool,happ(s(fun(bool,bool),X53),s(bool,'T_0')))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52))))) = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X53,bool))) )),
  inference(superposition,[],[f354,f479])).
fof(f2149,plain,(
  ( ! [X59,X58] : (s(bool,X58) = s(bool,sK9('const.marker.Cong_1'(s(bool,X58)),X59,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X58)) )),
  inference(forward_demodulation,[],[f2117,f271])).
fof(f2117,plain,(
  ( ! [X59,X58] : (~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | s(bool,'const.marker.Cong_1'(s(bool,X58))) = s(bool,sK9('const.marker.Cong_1'(s(bool,X58)),X59,bool)) | s(bool,'F_0') = s(bool,X58)) )),
  inference(superposition,[],[f354,f489])).
fof(f2151,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') = s(bool,sK9('const.bool.BOUNDED_1'(s(bool,X62)),X63,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f2119,f272])).
fof(f2119,plain,(
  ( ! [X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0')))) | s(bool,'const.bool.BOUNDED_1'(s(bool,X62))) = s(bool,sK9('const.bool.BOUNDED_1'(s(bool,X62)),X63,bool))) )),
  inference(superposition,[],[f354,f272])).
fof(f2098,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,'F_0')))) | s(bool,X11) = s(bool,sK9(X11,X12,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f354,f276])).
fof(f2145,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num',X50),s('type.num.num','const.num.|48|_0')),X51,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X51),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f2113,f462])).
fof(f2113,plain,(
  ( ! [X50,X51] : (~p(s(bool,happ(s(fun(bool,bool),X51),s(bool,'F_0')))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X50),s('type.num.num','const.num.|48|_0'))) = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num',X50),s('type.num.num','const.num.|48|_0')),X51,bool))) )),
  inference(superposition,[],[f354,f462])).
fof(f2152,plain,(
  ( ! [X64,X65] : (s(fun('type.list.list'(X64),X64),'const.list.HD_0') = s(fun('type.list.list'(X64),X64),sK9('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X65,fun('type.list.list'(X64),X64))) | ~p(s(bool,happ(s(fun(fun('type.list.list'(X64),X64),bool),X65),s(fun('type.list.list'(X64),X64),'const.list.HD_0'))))) )),
  inference(forward_demodulation,[],[f2120,f275])).
fof(f2120,plain,(
  ( ! [X64,X65] : (~p(s(bool,happ(s(fun(fun('type.list.list'(X64),X64),bool),X65),s(fun('type.list.list'(X64),X64),'const.list.HD_0')))) | s(fun('type.list.list'(X64),X64),'const.list.EL_1'(s('type.num.num','const.num.|48|_0'))) = s(fun('type.list.list'(X64),X64),sK9('const.list.EL_1'(s('type.num.num','const.num.|48|_0')),X65,fun('type.list.list'(X64),X64)))) )),
  inference(superposition,[],[f354,f275])).
fof(f2122,plain,(
  ( ! [X68,X69] : (~p(s(bool,happ(s(fun('type.num.num',bool),X69),s('type.num.num',X68)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X68)))) = s('type.num.num',sK9('const.num.SUC_1'(s('type.num.num',sK5(X68))),X69,'type.num.num')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X68)) )),
  inference(superposition,[],[f354,f277])).
fof(f2154,plain,(
  ( ! [X70,X71] : (s('type.num.num',X70) = s('type.num.num',sK9(sK5('const.num.SUC_1'(s('type.num.num',X70))),X71,'type.num.num')) | ~p(s(bool,happ(s(fun('type.num.num',bool),X71),s('type.num.num',X70))))) )),
  inference(forward_demodulation,[],[f2123,f839])).
fof(f2123,plain,(
  ( ! [X70,X71] : (~p(s(bool,happ(s(fun('type.num.num',bool),X71),s('type.num.num',X70)))) | s('type.num.num',sK5('const.num.SUC_1'(s('type.num.num',X70)))) = s('type.num.num',sK9(sK5('const.num.SUC_1'(s('type.num.num',X70))),X71,'type.num.num'))) )),
  inference(superposition,[],[f354,f839])).
fof(f5516,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f356,f276])).
fof(f5629,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),X61),s('type.num.num',X62))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X60),s(fun('type.num.num',bool),X61),s('type.num.num','const.num.SUC_1'(s('type.num.num',X62)))))) )),
  inference(forward_demodulation,[],[f5545,f271])).
fof(f5545,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun('type.num.num',bool),X61),s('type.num.num',X62))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X60),s(fun('type.num.num',bool),X61),s('type.num.num','const.num.SUC_1'(s('type.num.num',X62)))))) )),
  inference(superposition,[],[f621,f356])).
fof(f10737,plain,(
  ( ! [X196,X197,X195] : (~p(s(bool,happ(s(fun('type.num.num',bool),X196),s('type.num.num',X195)))) | p(s(bool,happ(s(fun('type.num.num',bool),X196),s('type.num.num',X197)))) | s('type.num.num',X197) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X195)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X195)) )),
  inference(superposition,[],[f353,f277])).
fof(f2153,plain,(
  ( ! [X66,X67] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(X66),'const.list.NIL_0')),X67,'type.num.num')) | ~p(s(bool,happ(s(fun('type.num.num',bool),X67),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f2121,f273])).
fof(f2121,plain,(
  ( ! [X66,X67] : (~p(s(bool,happ(s(fun('type.num.num',bool),X67),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X66),'const.list.NIL_0'))) = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(X66),'const.list.NIL_0')),X67,'type.num.num'))) )),
  inference(superposition,[],[f354,f273])).
fof(f11888,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK19(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK19(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f379,f276])).
fof(f381,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) )),
  inference(cnf_transformation,[],[f256])).
fof(f396,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X4,X3) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f261])).
fof(f4047,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X3,fun(X4,bool)),X5),s('type.list.list'(X3),X7),s('type.list.list'(X4),X8)))) | sP0(X3,X4,'const.list.HD_1'(s('type.list.list'(fun(X3,fun(X4,bool))),'const.list.CONS_2'(s(fun(X3,fun(X4,bool)),X5),s('type.list.list'(fun(X3,fun(X4,bool))),X6)))),X8,X7)) )),
  inference(superposition,[],[f395,f343])).
fof(f4048,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X9,fun(X10,bool)),X11),s('type.list.list'(X9),X13),s('type.list.list'(X10),X14)))) | sP0(X9,X10,'const.ind_type.FCONS_3'(s(fun(X9,fun(X10,bool)),X11),s(fun('type.num.num',fun(X9,fun(X10,bool))),X12),s('type.num.num','const.num.|48|_0')),X14,X13)) )),
  inference(superposition,[],[f395,f344])).
fof(f4050,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X3,bool)),X4),s('type.list.list'(X0),X1),s('type.list.list'(X3),X5)))) | sP0(X0,X3,X4,X5,'const.list.HD_1'(s('type.list.list'('type.list.list'(X0)),'const.list.CONS_2'(s('type.list.list'(X0),X1),s('type.list.list'('type.list.list'(X0)),X2)))))) )),
  inference(superposition,[],[f395,f343])).
fof(f4051,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X9,bool)),X10),s('type.list.list'(X6),X7),s('type.list.list'(X9),X11)))) | sP0(X6,X9,X10,X11,'const.ind_type.FCONS_3'(s('type.list.list'(X6),X7),s(fun('type.num.num','type.list.list'(X6)),X8),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f395,f344])).
fof(f4053,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X3,fun(X0,bool)),X4),s('type.list.list'(X3),X5),s('type.list.list'(X0),X1)))) | sP0(X3,X0,X4,'const.list.HD_1'(s('type.list.list'('type.list.list'(X0)),'const.list.CONS_2'(s('type.list.list'(X0),X1),s('type.list.list'('type.list.list'(X0)),X2)))),X5)) )),
  inference(superposition,[],[f395,f343])).
fof(f4054,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X9,fun(X6,bool)),X10),s('type.list.list'(X9),X11),s('type.list.list'(X6),X7)))) | sP0(X9,X6,X10,'const.ind_type.FCONS_3'(s('type.list.list'(X6),X7),s(fun('type.num.num','type.list.list'(X6)),X8),s('type.num.num','const.num.|48|_0')),X11)) )),
  inference(superposition,[],[f395,f344])).
fof(f360,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X47,X27] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X47),s('type.num.num',X27))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.prim_rec.LESS_MONO_EQ')).
fof(f3794,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f360,f276])).
fof(f3861,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) )),
  inference(forward_demodulation,[],[f3817,f271])).
fof(f3817,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) )),
  inference(superposition,[],[f621,f360])).
fof(f3785,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X0)),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f360,f277])).
fof(f3793,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK5(X0)))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f360,f277])).
fof(f5310,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.list.EL_2'(s('type.num.num',X48),s('type.list.list'(bool),X49)))) | p(s(bool,'const.list.EL_2'(s('type.num.num',X48),s('type.list.list'(bool),sK9(X49,'const.list.EL_1'(s('type.num.num',X48)),'type.list.list'(bool))))))) )),
  inference(forward_demodulation,[],[f5280,f346])).
fof(f5280,plain,(
  ( ! [X48,X49] : (p(s(bool,'const.list.EL_2'(s('type.num.num',X48),s('type.list.list'(bool),sK9(X49,'const.list.EL_1'(s('type.num.num',X48)),'type.list.list'(bool)))))) | ~p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X48))),s('type.list.list'(bool),X49))))) )),
  inference(superposition,[],[f355,f346])).
fof(f4480,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(bool),X4))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X3))),s('type.list.list'(bool),X4)))) )),
  inference(superposition,[],[f346,f276])).
fof(f4496,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.list.EL_2'(s('type.num.num',X16),s('type.list.list'(bool),X17)))) | s('type.list.list'(bool),X17) = s('type.list.list'(bool),sK9(X17,'const.list.EL_1'(s('type.num.num',X16)),'type.list.list'(bool)))) )),
  inference(superposition,[],[f354,f346])).
fof(f4579,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,'const.list.EL_2'(s('type.num.num',X48),s('type.list.list'(bool),X49))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X48))),s('type.list.list'(bool),X49)))) )),
  inference(forward_demodulation,[],[f4510,f271])).
fof(f4510,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X48),s('type.list.list'(bool),X49))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X48))),s('type.list.list'(bool),X49)))) )),
  inference(superposition,[],[f621,f346])).
fof(f10842,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,'const.list.EL_2'(s('type.num.num',X75),s('type.list.list'(bool),X77)))) | ~p(s(bool,'const.list.EL_2'(s('type.num.num',X75),s('type.list.list'(bool),X76)))) | s('type.list.list'(bool),X76) != s('type.list.list'(bool),X77)) )),
  inference(forward_demodulation,[],[f10756,f346])).
fof(f10756,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,'const.list.EL_2'(s('type.num.num',X75),s('type.list.list'(bool),X76)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X75))),s('type.list.list'(bool),X77)))) | s('type.list.list'(bool),X76) != s('type.list.list'(bool),X77)) )),
  inference(superposition,[],[f353,f346])).
fof(f11309,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(bool),X6))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X5),s('type.list.list'(bool),X6)))))) )),
  inference(superposition,[],[f357,f276])).
fof(f11471,plain,(
  ( ! [X80,X81,X79] : (s(bool,'F_0') != s(bool,'const.list.EL_2'(s('type.num.num',X79),s('type.list.list'(bool),X81))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X80),s('type.list.list'(bool),X81)))))) )),
  inference(forward_demodulation,[],[f11344,f271])).
fof(f11344,plain,(
  ( ! [X80,X81,X79] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X79),s('type.list.list'(bool),X81))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X80),s('type.list.list'(bool),X81)))))) )),
  inference(superposition,[],[f621,f357])).
fof(f1461,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X11))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X11)))) )),
  inference(superposition,[],[f338,f276])).
fof(f1506,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X22))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X22)))) )),
  inference(forward_demodulation,[],[f1476,f271])).
fof(f1476,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X22))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X22)))) )),
  inference(superposition,[],[f621,f338])).
fof(f2133,plain,(
  ( ! [X34] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X34)))) | s('type.list.list'(bool),X34) = s('type.list.list'(bool),sK9(X34,'const.list.HD_0','type.list.list'(bool)))) )),
  inference(superposition,[],[f354,f338])).
fof(f2134,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) | s('type.list.list'(bool),X0) = s('type.list.list'(bool),sK9(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))) )),
  inference(forward_demodulation,[],[f2091,f338])).
fof(f2091,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X0)))) | s('type.list.list'(bool),X0) = s('type.list.list'(bool),sK9(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))) )),
  inference(superposition,[],[f354,f275])).
fof(f5283,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) | p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK9(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool))))))) )),
  inference(forward_demodulation,[],[f5282,f4561])).
fof(f5282,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X0)))) | p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK9(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool))))))) )),
  inference(forward_demodulation,[],[f5281,f346])).
fof(f5281,plain,(
  ( ! [X0] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK9(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))))) | ~p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num','const.num.|48|_0'))),s('type.list.list'(bool),X0))))) )),
  inference(forward_demodulation,[],[f5248,f338])).
fof(f5248,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),sK9(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))))) | ~p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num','const.num.|48|_0'))),s('type.list.list'(bool),X0))))) )),
  inference(superposition,[],[f355,f275])).
fof(f5309,plain,(
  ( ! [X47] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X47)))) | p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK9(X47,'const.list.HD_0','type.list.list'(bool))))))) )),
  inference(forward_demodulation,[],[f5279,f338])).
fof(f5279,plain,(
  ( ! [X47] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),sK9(X47,'const.list.HD_0','type.list.list'(bool)))))) | ~p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X47))))) )),
  inference(superposition,[],[f355,f338])).
fof(f10759,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X1)))) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) | s('type.list.list'(bool),X0) != s('type.list.list'(bool),X1)) )),
  inference(forward_demodulation,[],[f10758,f4561])).
fof(f10758,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(bool),X1)))) | ~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) | s('type.list.list'(bool),X0) != s('type.list.list'(bool),X1)) )),
  inference(forward_demodulation,[],[f10757,f346])).
fof(f10757,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X0)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num','const.num.|48|_0'))),s('type.list.list'(bool),X1)))) | s('type.list.list'(bool),X0) != s('type.list.list'(bool),X1)) )),
  inference(forward_demodulation,[],[f10677,f338])).
fof(f10677,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X0)))) | p(s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num','const.num.|48|_0'))),s('type.list.list'(bool),X1)))) | s('type.list.list'(bool),X0) != s('type.list.list'(bool),X1)) )),
  inference(superposition,[],[f353,f275])).
fof(f11872,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK19(X0,X1,X2,bool))) )),
  inference(superposition,[],[f379,f276])).
fof(f11428,plain,(
  ( ! [X424,X426,X423,X425] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X426))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.EL_2'(s('type.num.num',X423),s('type.list.list'('type.num.num'),X425))))) | s('type.num.num',X426) = s('type.num.num','const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X423))),s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X424),s('type.list.list'('type.num.num'),X425)))))) )),
  inference(superposition,[],[f365,f357])).
fof(f11245,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,'const.list.EL_2'(s('type.num.num',sK5(X0)),s('type.list.list'(X1),X2))) = s(X1,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X2))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f357,f277])).
fof(f6100,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X68) = s(bool,'const.list.LIST_REL_3'(s(fun(X63,fun(X64,bool)),X65),s('type.list.list'(X63),'const.list.NIL_0'),s('type.list.list'(X64),'const.list.CONS_2'(s(X64,X66),s('type.list.list'(X64),X67))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f6016,f426])).
fof(f6016,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (p(s(bool,'F_0')) | s(bool,X68) = s(bool,'const.list.LIST_REL_3'(s(fun(X63,fun(X64,bool)),X65),s('type.list.list'(X63),'const.list.NIL_0'),s('type.list.list'(X64),'const.list.CONS_2'(s(X64,X66),s('type.list.list'(X64),X67))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f546,f380])).
fof(f6095,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,'const.list.LIST_REL_3'(s(fun(X20,fun(X21,bool)),X22),s('type.list.list'(X20),'const.list.NIL_0'),s('type.list.list'(X21),'const.list.CONS_2'(s(X21,X23),s('type.list.list'(X21),X24)))))) )),
  inference(subsumption_resolution,[],[f6008,f426])).
fof(f6008,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.list.LIST_REL_3'(s(fun(X20,fun(X21,bool)),X22),s('type.list.list'(X20),'const.list.NIL_0'),s('type.list.list'(X21),'const.list.CONS_2'(s(X21,X23),s('type.list.list'(X21),X24)))))) )),
  inference(superposition,[],[f281,f380])).
fof(f5965,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,X105) = s(bool,'const.list.LIST_REL_3'(s(fun(X106,fun(bool,bool)),X107),s('type.list.list'(X106),'const.list.NIL_0'),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X105),s('type.list.list'(bool),X108))))) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f380,f276])).
fof(f5596,plain,(
  ( ! [X265,X267,X268,X266] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X268))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X266),s('type.num.num',X267))))) | s('type.num.num',X268) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X265),s(fun('type.num.num','type.num.num'),X266),s('type.num.num','const.num.SUC_1'(s('type.num.num',X267)))))) )),
  inference(superposition,[],[f365,f356])).
fof(f5515,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun('type.num.num',X1),X2),s('type.num.num',sK5(X0)))) = s(X1,'const.ind_type.FCONS_3'(s(X1,X3),s(fun('type.num.num',X1),X2),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f356,f277])).
fof(f4553,plain,(
  ( ! [X180,X181,X179] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X181))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.EL_2'(s('type.num.num',X179),s('type.list.list'('type.num.num'),X180))))) | s('type.num.num',X181) = s('type.num.num',happ(s(fun('type.list.list'('type.num.num'),'type.num.num'),'const.list.EL_1'(s('type.num.num',X179))),s('type.list.list'('type.num.num'),X180)))) )),
  inference(superposition,[],[f365,f346])).
fof(f4075,plain,(
  ( ! [X2,X3,X1] : (sP0(X1,X2,X3,'const.list.NIL_0','const.list.NIL_0')) )),
  inference(subsumption_resolution,[],[f4058,f264])).
fof(f4058,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | sP0(X1,X2,X3,'const.list.NIL_0','const.list.NIL_0')) )),
  inference(superposition,[],[f395,f368])).
fof(f3258,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X5))))))) | s('type.num.num',X6) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5)))))) )),
  inference(superposition,[],[f365,f345])).
fof(f3064,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,'const.list.HD_1'(s('type.list.list'('type.list.list'(X1)),'const.list.CONS_2'(s('type.list.list'(X1),'const.list.NIL_0'),s('type.list.list'('type.list.list'(X1)),X3)))),'const.list.HD_1'(s('type.list.list'('type.list.list'(X0)),'const.list.CONS_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'('type.list.list'(X0)),X4)))))) )),
  inference(unit_resulting_resolution,[],[f343,f343,f393])).
fof(f3065,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,'const.list.HD_1'(s('type.list.list'('type.list.list'(X1)),'const.list.CONS_2'(s('type.list.list'(X1),'const.list.NIL_0'),s('type.list.list'('type.list.list'(X1)),X3)))),'const.ind_type.FCONS_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(fun('type.num.num','type.list.list'(X0)),X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f344,f343,f393])).
fof(f3066,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,'const.ind_type.FCONS_3'(s('type.list.list'(X1),'const.list.NIL_0'),s(fun('type.num.num','type.list.list'(X1)),X3),s('type.num.num','const.num.|48|_0')),'const.list.HD_1'(s('type.list.list'('type.list.list'(X0)),'const.list.CONS_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'('type.list.list'(X0)),X4)))))) )),
  inference(unit_resulting_resolution,[],[f343,f344,f393])).
fof(f3067,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,'const.ind_type.FCONS_3'(s('type.list.list'(X1),'const.list.NIL_0'),s(fun('type.num.num','type.list.list'(X1)),X3),s('type.num.num','const.num.|48|_0')),'const.ind_type.FCONS_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(fun('type.num.num','type.list.list'(X0)),X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f344,f344,f393])).
fof(f1488,plain,(
  ( ! [X47,X46] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.HD_1'(s('type.list.list'('type.num.num'),X46))))) | s('type.num.num',X47) = s('type.num.num',happ(s(fun('type.list.list'('type.num.num'),'type.num.num'),'const.list.HD_0'),s('type.list.list'('type.num.num'),X46)))) )),
  inference(superposition,[],[f365,f338])).
fof(f966,plain,(
  ( ! [X2,X3] : (s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f366,f277])).
fof(f846,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num',sK5('const.num.SUC_1'(s('type.num.num',X2))))) )),
  inference(superposition,[],[f365,f573])).
fof(f845,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num',sK5(X0)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f365,f277])).
fof(f844,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X9),s(fun('type.num.num','type.num.num'),X10),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f365,f344])).
fof(f843,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) | s('type.num.num',X8) = s('type.num.num','const.list.HD_1'(s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X6),s('type.list.list'('type.num.num'),X7)))))) )),
  inference(superposition,[],[f365,f343])).
fof(f841,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X2)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f365,f277])).
fof(f840,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f365,f273])).
fof(f837,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f336,f365])).
fof(f839,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num',sK5('const.num.SUC_1'(s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f573,f365])).
fof(f799,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK6(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f795,f426])).
fof(f795,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK6(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f347,f276])).
fof(f789,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,'const.bool.BOUNDED_1'(s(bool,X13))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f780,f264])).
fof(f780,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,X14) = s(bool,'const.bool.BOUNDED_1'(s(bool,X13))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f623,f272])).
fof(f787,plain,(
  ( ! [X10,X9] : (s(bool,X10) = s(bool,'const.marker.Cong_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f778,f264])).
fof(f778,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s(bool,X10) = s(bool,'const.marker.Cong_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f623,f489])).
fof(f786,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f777,f264])).
fof(f777,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s(bool,X8) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f623,f479])).
fof(f764,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f756,f426])).
fof(f756,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f546,f462])).
fof(f747,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f730,f271])).
fof(f730,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X29))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f621,f344])).
fof(f746,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f729,f271])).
fof(f729,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X27))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f489,f344])).
fof(f717,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f344,f276])).
fof(f698,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X7),s('type.list.list'(bool),X8)))))) )),
  inference(forward_demodulation,[],[f686,f271])).
fof(f686,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X7),s('type.list.list'(bool),X8)))))) )),
  inference(superposition,[],[f621,f343])).
fof(f664,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X23),s('type.list.list'(bool),X24)))))) )),
  inference(forward_demodulation,[],[f650,f271])).
fof(f650,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X23),s('type.list.list'(bool),X24)))))) )),
  inference(superposition,[],[f489,f343])).
fof(f639,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X3),s('type.list.list'(bool),X4)))))) )),
  inference(superposition,[],[f343,f276])).
fof(f621,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X6))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f616,f264])).
fof(f616,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X6))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f285,f489])).
fof(f582,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f479,f277])).
fof(f573,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5('const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f336,f277])).
fof(f486,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f276,f276])).
fof(f479,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f274,f282])).
fof(f462,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f337,f286])).
fof(f438,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f337,f281])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X38,X39,X40] : ! [X41,X42,X43] : s(X38,happ(s(fun(X39,X38),X41),s(X39,happ(s(fun(X40,X39),X42),s(X40,X43))))) = s(X38,happ(s(fun(X40,X38),'_dst_x0x1_2'(s(fun(X39,X38),X41),s(fun(X40,X39),X42))),s(X40,X43)))),
  file('thm.list.LIST_REL_NIL|split|0','_dst_x0x1_ax')).
fof(f385,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) | s('type.list.list'(X0),X4) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK21(X4,X3,X2,X1,X0)),s('type.list.list'(X0),sK23(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f387,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) | s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK22(X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK24(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f388,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3) | s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK22(X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK24(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f393,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X4)) )),
  inference(cnf_transformation,[],[f260])).
fof(f379,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.list.LIST_REL_NIL|split|0',aHL_EXT)).
fof(f369,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK10(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK10(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK10(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK10(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X7,X6] : ! [X30] : ? [X29] : ! [X31,X32,X33] : s(X7,happ(s(fun('type.ind_type.recspace'(X6),X7),X29),s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X31),s(X6,X32),s(fun('type.num.num','type.ind_type.recspace'(X6)),X33))))) = s(X7,happ(s(fun(fun('type.num.num',X7),X7),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X7),X7)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X7),X7))),happ(s(fun('type.num.num',fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X7),X7)))),X30),s('type.num.num',X31))),s(X6,X32))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X33))),s(fun('type.num.num',X7),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X6),X7),X29),s(fun('type.num.num','type.ind_type.recspace'(X6)),X33)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.ind_type.CONSTR_REC')).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0'))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ! [X2] : p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0'))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X9] : p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X7),'const.list.NIL_0'))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LIST_REL_rules|split|0')).
fof(f365,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X47,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))) <=> s('type.num.num',X27) = s('type.num.num',X47))),
  file('thm.list.LIST_REL_NIL|split|0','thm.prim_rec.INV_SUC_EQ')).
fof(f366,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f246])).
fof(f359,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ! [X34,X35,X36] : s(X6,'const.ind_type.FCONS_3'(s(X6,X34),s(fun('type.num.num',X6),X35),s('type.num.num',X36))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X34))),s(fun('type.num.num',X6),X35))),s('type.num.num',X36)))),
  file('thm.list.LIST_REL_NIL|split|0','ARITY_0_3_const.ind_type.FCONS')).
fof(f358,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X34,X35,X36] : s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X34),s(X6,X35),s(fun('type.num.num','type.ind_type.recspace'(X6)),X36))) = s('type.ind_type.recspace'(X6),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6))),'const.ind_type.CONSTR_1'(s('type.num.num',X34))),s(X6,X35))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X36)))),
  file('thm.list.LIST_REL_NIL|split|0','ARITY_1_3_const.ind_type.CONSTR')).
fof(f357,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X7] : ! [X27,X45,X46] : s(X7,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(X7),X45))) = s(X7,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X46),s('type.list.list'(X7),X45)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.EL_restricted|split|1')).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X28,X29,X27] : s(X6,happ(s(fun('type.num.num',X6),X29),s('type.num.num',X27))) = s(X6,'const.ind_type.FCONS_3'(s(X6,X28),s(fun('type.num.num',X6),X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.ind_type.FCONS|split|1')).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ! [X28,X29] : s(X6,X28) = s(X6,'const.ind_type.FCONS_3'(s(X6,X28),s(fun('type.num.num',X6),X29),s('type.num.num','const.num.|48|_0')))),
  file('thm.list.LIST_REL_NIL|split|0','thm.ind_type.FCONS|split|0')).
fof(f337,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X27] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.num.|48|_0'))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.prim_rec.NOT_LESS_|48|')).
fof(f286,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f277,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f81])).
fof(f81,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X47] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X47) | ? [X27] : s('type.num.num',X47) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.arithmetic.num_CASES')).
fof(f274,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X27] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))))),
  file('thm.list.LIST_REL_NIL|split|0','thm.prim_rec.LESS_|48|')).
fof(f272,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X37] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X37)))),
  file('thm.list.LIST_REL_NIL|split|0','thm.bool.BOUNDED_THM')).
fof(f34597,plain,(
  ( ! [X125,X126] : (s(bool,X126) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X125)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f34070,f17030])).
fof(f34070,plain,(
  ( ! [X125,X126] : (s(bool,'T_0') != s(bool,X126) | s(bool,X126) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X125)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14699])).
fof(f14699,plain,(
  ( ! [X125,X126] : (s(bool,X126) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X125))) | s(bool,X126) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X125)))) ) | $spl0),
  inference(forward_demodulation,[],[f14618,f271])).
fof(f14618,plain,(
  ( ! [X125,X126] : (s(bool,X126) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X125))))) | s(bool,X126) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X125)))) ) | $spl0),
  inference(superposition,[],[f13848,f338])).
fof(f13848,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | $spl0),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f422,f438,f437,f447,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f785,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f875,f877,f878,f879,f966,f1004,f1006,f1007,f1027,f1029,f1045,f1046,f1136,f1396,f1397,f1450,f1429,f1431,f1432,f1488,f1930,f1923,f1560,f1561,f1563,f1564,f1566,f1942,f1944,f1945,f1587,f1588,f1985,f1992,f1710,f2007,f2008,f2013,f1810,f1798,f2014,f2074,f2044,f2045,f2047,f2156,f2175,f2178,f2308,f2359,f2360,f2371,f2374,f2499,f2500,f2944,f2951,f2992,f2993,f2995,f3067,f3066,f3065,f3064,f3103,f3201,f3202,f3210,f3258,f3336,f3337,f4070,f4074,f4075,f4154,f4155,f4160,f4389,f4561,f4553,f5002,f5009,f5187,f5377,f5378,f5450,f5426,f5437,f5438,f5515,f5596,f5693,f5704,f5705,f5814,f5815,f5962,f5963,f5965,f6095,f6100,f6166,f6167,f6197,f6293,f6294,f6213,f6330,f6441,f6478,f6434,f6481,f6428,f6429,f6431,f6296,f5817,f5439,f5309,f5283,f2501,f2134,f2133,f1589,f1562,f1506,f1496,f1483,f1461,f1458,f1456,f6169,f5707,f1485,f4542,f5583,f794,f793,f347,f4497,f6432,f6297,f6170,f5818,f5708,f5440,f5310,f4598,f4579,f4562,f4521,f4496,f4488,f4486,f4483,f4482,f4480,f3793,f3785,f6414,f6287,f6152,f5811,f5690,f5423,f3879,f3861,f3846,f3828,f3802,f3800,f3797,f3796,f3794,f360,f3277,f6430,f6295,f6168,f5816,f5706,f5648,f5629,f5610,f5558,f5556,f5524,f5522,f5519,f5518,f5516,f2154,f2122,f2153,f2152,f3884,f6116,f2515,f2513,f2312,f2148,f2145,f2100,f2098,f2538,f2366,f2151,f2149,f2146,f2144,f2138,f5288,f5256,f5255,f5253,f4463,f2943,f2516,f2150,f2143,f2142,f2110,f2101,f2537,f2094,f5286,f4541,f5287,f4607,f5665,f5664,f5582,f390,f389,f2536,f2535,f348,f2543,f2137,f4569,f5619,f5301,f5299,f5285,f5284,f2895,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4046,f4049,f4543,f5584,f392,f391,f4052,f4546,f5588,f4055,f4545,f5587,f4077,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f394,f384,f383,f382,f381,f378,f377,f376,f375,f374,f373,f372,f6980,f7074,f7079,f7138,f7277,f7776,f7777,f7778,f7863,f7873,f7874,f7849,f7890,f8185,f8186,f8187,f8652,f8296,f8299,f8311,f8312,f8313,f8314,f8315,f8650,f8684,f8687,f8692,f8693,f8694,f8695,f8696,f8808,f8809,f8877,f8912,f9308,f9360,f9361,f9372,f9373,f9876,f9883,f9884,f9980,f10054,f10186,f10142,f10310,f10337,f10338,f10340,f10510,f10528,f10529,f10530,f10531,f10532,f10759,f10688,f10770,f10690,f10691,f10778,f10785,f10707,f10786,f10713,f10737,f10824,f10842,f10961,f10974,f10879,f10897,f10898,f10899,f10900,f10901,f11245,f11309,f11448,f11311,f11312,f11315,f11317,f11320,f11461,f11471,f11355,f11357,f11358,f11477,f11369,f11490,f11493,f11494,f11376,f11379,f11390,f11532,f11533,f11411,f11412,f11413,f11421,f11422,f11428,f11437,f11439,f11720,f11738,f11739,f11740,f11741,f11742,f11743,f11872,f11904,f11874,f11875,f11878,f11879,f11880,f11907,f11911,f11888,f11912,f11890,f11891,f11894,f11895,f11896,f11899,f11918,f13363,f13364,f12190,f12209,f12210,f12211,f12212,f12213,f12214,f13441,f13444,f13445,f13446,f13447,f13448,f13449,f13451,f13452,f13453,f13454,f12358,f13470,f13471,f13472,f13473,f12393,f13475,f13476,f12407,f12408,f12409,f12410,f12411,f12412,f13543,f12581,f12590,f12759,f12768,f13730,f13732,f13782,f13783,f12883,f12884,f12885,f12886,f12887,f12889,f12890,f12891,f12895,f13797,f13798,f13799,f13811,f13817,f13835,f13837,f13838,f13840,f13845,f13119,f13846,f13847,f13128,f13139])).
fof(f13139,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(equality_factoring,[],[f9373])).
fof(f13128,plain,(
  ( ! [X472,X469,X471,X470] : (s(X471,f(s(bool,'T_0'))) != s(X471,g(s(bool,'T_0'))) | s(fun(bool,X471),X469) = s(fun(bool,X471),X470) | s(bool,X472) = s(bool,sK19(X469,X470,X471,bool)) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X472)))) ) | $spl0),
  inference(superposition,[],[f379,f9373])).
fof(f13847,plain,(
  ( ! [X466,X467,X468] : (s(bool,'T_0') = s(bool,X468) | p(s(bool,happ(s(fun(bool,bool),X467),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X467),s(bool,X466)))) | s(bool,X468) = s(bool,sK9(X466,X467,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f13127,f271])).
fof(f13127,plain,(
  ( ! [X466,X467,X468] : (p(s(bool,happ(s(fun(bool,bool),X467),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X467),s(bool,X466)))) | s(bool,X468) = s(bool,sK9(X466,X467,bool)) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X468)))) ) | $spl0),
  inference(superposition,[],[f355,f9373])).
fof(f13846,plain,(
  ( ! [X463,X465,X464] : (s(bool,'T_0') = s(bool,X465) | ~p(s(bool,happ(s(fun(bool,bool),X464),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X464)))) | s(bool,X465) = s(bool,sK6(X463,X464,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f13126,f271])).
fof(f13126,plain,(
  ( ! [X463,X465,X464] : (~p(s(bool,happ(s(fun(bool,bool),X464),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X464)))) | s(bool,X465) = s(bool,sK6(X463,X464,bool)) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X465)))) ) | $spl0),
  inference(superposition,[],[f348,f9373])).
fof(f13119,plain,(
  ( ! [X441,X443,X440,X442] : (s(bool,'T_0') != s(bool,f(s(X440,sK19(X441,X442,bool,X440)))) | s(fun(X440,bool),X441) = s(fun(X440,bool),X442) | s(bool,X443) = s(bool,g(s(X440,sK19(X441,X442,bool,X440)))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X443)))) ) | $spl0),
  inference(superposition,[],[f379,f9373])).
fof(f13845,plain,(
  ( ! [X422,X421,X423] : (s(bool,'T_0') = s(bool,X423) | ~p(s(bool,happ(s(fun(bool,bool),X422),s(bool,'T_0')))) | s(bool,X421) = s(bool,sK9(X421,X422,bool)) | s(bool,X421) = s(bool,X423)) ) | $spl0),
  inference(forward_demodulation,[],[f13115,f271])).
fof(f13115,plain,(
  ( ! [X422,X421,X423] : (~p(s(bool,happ(s(fun(bool,bool),X422),s(bool,'T_0')))) | s(bool,X421) = s(bool,sK9(X421,X422,bool)) | s(bool,X421) = s(bool,X423) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X423)))) ) | $spl0),
  inference(superposition,[],[f354,f9373])).
fof(f13840,plain,(
  ( ! [X296,X298,X297,X299] : (s(bool,'T_0') = s(bool,X299) | p(s(bool,'|dollar|exists_1'(s(fun(X296,bool),X297)))) | s(bool,X299) = s(bool,happ(s(fun(X296,bool),X297),s(X296,sK6(X298,X297,X296))))) ) | $spl0),
  inference(forward_demodulation,[],[f13839,f271])).
fof(f13839,plain,(
  ( ! [X296,X298,X297,X299] : (p(s(bool,'|dollar|exists_1'(s(fun(X296,bool),X297)))) | s(bool,X299) = s(bool,happ(s(fun(X296,bool),X297),s(X296,sK6(X298,X297,X296)))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X299)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13067,f264])).
fof(f13067,plain,(
  ( ! [X296,X298,X297,X299] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X296,bool),X297)))) | s(bool,X299) = s(bool,happ(s(fun(X296,bool),X297),s(X296,sK6(X298,X297,X296)))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X299)))) ) | $spl0),
  inference(superposition,[],[f348,f9373])).
fof(f13838,plain,(
  ( ! [X292,X294,X293,X295] : (s(X292,X294) = s(X292,sK9(X294,X293,X292)) | s(bool,X295) = s(bool,happ(s(fun(X292,bool),X293),s(X292,X294))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X295)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13066,f264])).
fof(f13066,plain,(
  ( ! [X292,X294,X293,X295] : (~p(s(bool,'T_0')) | s(X292,X294) = s(X292,sK9(X294,X293,X292)) | s(bool,X295) = s(bool,happ(s(fun(X292,bool),X293),s(X292,X294))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X295)))) ) | $spl0),
  inference(superposition,[],[f354,f9373])).
fof(f13837,plain,(
  ( ! [X288,X290,X287,X289,X291] : (s(bool,'T_0') = s(bool,X291) | p(s(bool,happ(s(fun(X287,bool),X288),s(X287,X290)))) | s(X287,X289) != s(X287,X290) | s(bool,X291) = s(bool,happ(s(fun(X287,bool),X288),s(X287,X289)))) ) | $spl0),
  inference(forward_demodulation,[],[f13836,f271])).
fof(f13836,plain,(
  ( ! [X288,X290,X287,X289,X291] : (p(s(bool,happ(s(fun(X287,bool),X288),s(X287,X290)))) | s(X287,X289) != s(X287,X290) | s(bool,X291) = s(bool,happ(s(fun(X287,bool),X288),s(X287,X289))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X291)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13065,f264])).
fof(f13065,plain,(
  ( ! [X288,X290,X287,X289,X291] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X287,bool),X288),s(X287,X290)))) | s(X287,X289) != s(X287,X290) | s(bool,X291) = s(bool,happ(s(fun(X287,bool),X288),s(X287,X289))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X291)))) ) | $spl0),
  inference(superposition,[],[f353,f9373])).
fof(f13835,plain,(
  ( ! [X282,X285,X281,X283,X284,X286] : (s(bool,'T_0') = s(bool,X286) | sP0(X281,X282,X283,X285,X284) | s(bool,X286) = s(bool,'const.list.LIST_REL_3'(s(fun(X281,fun(X282,bool)),X283),s('type.list.list'(X281),X284),s('type.list.list'(X282),X285)))) ) | $spl0),
  inference(forward_demodulation,[],[f13834,f271])).
fof(f13834,plain,(
  ( ! [X282,X285,X281,X283,X284,X286] : (sP0(X281,X282,X283,X285,X284) | s(bool,X286) = s(bool,'const.list.LIST_REL_3'(s(fun(X281,fun(X282,bool)),X283),s('type.list.list'(X281),X284),s('type.list.list'(X282),X285))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X286)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13064,f264])).
fof(f13064,plain,(
  ( ! [X282,X285,X281,X283,X284,X286] : (~p(s(bool,'T_0')) | sP0(X281,X282,X283,X285,X284) | s(bool,X286) = s(bool,'const.list.LIST_REL_3'(s(fun(X281,fun(X282,bool)),X283),s('type.list.list'(X281),X284),s('type.list.list'(X282),X285))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X286)))) ) | $spl0),
  inference(superposition,[],[f395,f9373])).
fof(f13817,plain,(
  ( ! [X151,X150] : (s(bool,X150) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X150) = s(bool,X151) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X151)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13015,f271])).
fof(f13015,plain,(
  ( ! [X151,X150] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X150) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X150) = s(bool,X151) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X151)))) ) | $spl0),
  inference(superposition,[],[f5002,f9373])).
fof(f13811,plain,(
  ( ! [X90,X88,X89] : (s(bool,X88) = s(bool,'const.bool.BOUNDED_1'(s(bool,X89))) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X90)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13810,f437])).
fof(f13810,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X88) = s(bool,'const.bool.BOUNDED_1'(s(bool,X89))) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X90)))) ) | $spl0),
  inference(forward_demodulation,[],[f12992,f271])).
fof(f12992,plain,(
  ( ! [X90,X88,X89] : (s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X88) = s(bool,'const.bool.BOUNDED_1'(s(bool,X89))) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X90)))) ) | $spl0),
  inference(superposition,[],[f1945,f9373])).
fof(f13799,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X11)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12963,f264])).
fof(f12963,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X11)))) ) | $spl0),
  inference(superposition,[],[f285,f9373])).
fof(f13798,plain,(
  ( ! [X243,X241,X242] : (s(bool,'T_0') = s(bool,X243) | ~p(s(bool,happ(s(fun(bool,bool),X242),s(bool,X241)))) | s(bool,'T_0') = s(bool,sK9('T_0',X242,bool)) | s(bool,X241) = s(bool,X243)) ) | $spl0),
  inference(forward_demodulation,[],[f12954,f271])).
fof(f12954,plain,(
  ( ! [X243,X241,X242] : (~p(s(bool,happ(s(fun(bool,bool),X242),s(bool,X241)))) | s(bool,'T_0') = s(bool,sK9('T_0',X242,bool)) | s(bool,X241) = s(bool,X243) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X243)))) ) | $spl0),
  inference(superposition,[],[f354,f9373])).
fof(f13797,plain,(
  ( ! [X239,X237,X240,X238] : (s(bool,'T_0') = s(bool,X240) | ~p(s(bool,happ(s(fun(bool,bool),X238),s(bool,X237)))) | p(s(bool,happ(s(fun(bool,bool),X238),s(bool,X239)))) | s(bool,'T_0') != s(bool,X239) | s(bool,X237) = s(bool,X240)) ) | $spl0),
  inference(forward_demodulation,[],[f12953,f271])).
fof(f12953,plain,(
  ( ! [X239,X237,X240,X238] : (~p(s(bool,happ(s(fun(bool,bool),X238),s(bool,X237)))) | p(s(bool,happ(s(fun(bool,bool),X238),s(bool,X239)))) | s(bool,'T_0') != s(bool,X239) | s(bool,X237) = s(bool,X240) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X240)))) ) | $spl0),
  inference(superposition,[],[f353,f9373])).
fof(f12895,plain,(
  ( ! [X59,X60] : (s(bool,X59) = s(bool,'const.bool.BOUNDED_1'(s(bool,X59))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X60)))) ) | $spl0),
  inference(superposition,[],[f272,f9373])).
fof(f12891,plain,(
  ( ! [X45,X46,X44] : (s(bool,X44) != s(bool,X45) | s(bool,X45) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X44) = s(bool,X46) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X46)))) ) | $spl0),
  inference(superposition,[],[f5187,f9373])).
fof(f12890,plain,(
  ( ! [X39,X41,X40] : (s(bool,X40) = s(bool,'const.marker.Cong_1'(s(bool,X39))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X40))) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X41)))) ) | $spl0),
  inference(superposition,[],[f4389,f9373])).
fof(f12889,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) = s(bool,'const.marker.Cong_1'(s(bool,X36))) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X38)))) ) | $spl0),
  inference(superposition,[],[f4160,f9373])).
fof(f12887,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X29))) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X30)))) ) | $spl0),
  inference(superposition,[],[f3210,f9373])).
fof(f12886,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X25) = s(bool,'const.bool.BOUNDED_1'(s(bool,X26))) | s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X27)))) ) | $spl0),
  inference(superposition,[],[f3103,f9373])).
fof(f12885,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X13) != s(bool,X14) | s(bool,X14) = s(bool,X15) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X15))) | s(bool,X13) = s(bool,X16) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X16)))) ) | $spl0),
  inference(superposition,[],[f1710,f9373])).
fof(f12884,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X11) = s(bool,'const.marker.Cong_1'(s(bool,X10))) | s(bool,'F_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X12)))) ) | $spl0),
  inference(superposition,[],[f1429,f9373])).
fof(f12883,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X5) != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X5) = s(bool,X8) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X8)))) ) | $spl0),
  inference(superposition,[],[f1136,f9373])).
fof(f13783,plain,(
  ( ! [X282,X281,X283] : (s(bool,X281) = s(bool,sK9('const.marker.Cong_1'(s(bool,X281)),X282,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X282),s(bool,'T_0')))) | s(bool,X281) = s(bool,X283) | s(bool,'T_0') = s(bool,X283)) ) | $spl0),
  inference(forward_demodulation,[],[f12878,f271])).
fof(f12878,plain,(
  ( ! [X282,X281,X283] : (~p(s(bool,happ(s(fun(bool,bool),X282),s(bool,'T_0')))) | s(bool,'const.marker.Cong_1'(s(bool,X281))) = s(bool,sK9('const.marker.Cong_1'(s(bool,X281)),X282,bool)) | s(bool,X281) = s(bool,X283) | s(bool,'T_0') = s(bool,X283)) ) | $spl0),
  inference(superposition,[],[f354,f9373])).
fof(f13782,plain,(
  ( ! [X280,X277,X279,X278] : (s(bool,X277) != s(bool,X279) | ~p(s(bool,happ(s(fun(bool,bool),X278),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X278),s(bool,X279)))) | s(bool,X277) = s(bool,X280) | s(bool,'T_0') = s(bool,X280)) ) | $spl0),
  inference(forward_demodulation,[],[f12877,f271])).
fof(f12877,plain,(
  ( ! [X280,X277,X279,X278] : (~p(s(bool,happ(s(fun(bool,bool),X278),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X278),s(bool,X279)))) | s(bool,X279) != s(bool,'const.marker.Cong_1'(s(bool,X277))) | s(bool,X277) = s(bool,X280) | s(bool,'T_0') = s(bool,X280)) ) | $spl0),
  inference(superposition,[],[f353,f9373])).
fof(f13732,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,X82)) | s(bool,X81) = s(bool,X82) | s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) ) | $spl0),
  inference(forward_demodulation,[],[f13731,f271])).
fof(f13731,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,X82)) | s(bool,X82) = s(bool,'const.marker.Cong_1'(s(bool,X81))) | s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) ) | $spl0),
  inference(subsumption_resolution,[],[f12804,f264])).
fof(f12804,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X82)) | s(bool,X82) = s(bool,'const.marker.Cong_1'(s(bool,X81))) | s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) ) | $spl0),
  inference(superposition,[],[f364,f9373])).
fof(f13730,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X71))) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) ) | $spl0),
  inference(subsumption_resolution,[],[f12800,f264])).
fof(f12800,plain,(
  ( ! [X72,X71] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X71))) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) ) | $spl0),
  inference(superposition,[],[f285,f9373])).
fof(f12768,plain,(
  ( ! [X492,X490,X489,X491] : (s(X491,f(s(bool,X492))) != s(X491,g(s(bool,X492))) | s(fun(bool,X491),X489) = s(fun(bool,X491),X490) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,sK19(X489,X490,X491,bool)))) | s(bool,'T_0') = s(bool,X492)) ) | $spl0),
  inference(superposition,[],[f379,f9373])).
fof(f12759,plain,(
  ( ! [X461,X463,X460,X462] : (s(bool,X463) != s(bool,f(s(X460,sK19(X461,X462,bool,X460)))) | s(fun(X460,bool),X461) = s(fun(X460,bool),X462) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,g(s(X460,sK19(X461,X462,bool,X460)))))) | s(bool,'T_0') = s(bool,X463)) ) | $spl0),
  inference(superposition,[],[f379,f9373])).
fof(f12590,plain,(
  ( ! [X492,X490,X489,X491] : (s(X491,f(s(bool,X492))) != s(X491,g(s(bool,X492))) | s(fun(bool,X491),X489) = s(fun(bool,X491),X490) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X492))) | s(bool,'T_0') = s(bool,sK19(X489,X490,X491,bool))) ) | $spl0),
  inference(superposition,[],[f379,f9373])).
fof(f12581,plain,(
  ( ! [X461,X463,X460,X462] : (s(bool,X463) != s(bool,f(s(X460,sK19(X461,X462,bool,X460)))) | s(fun(X460,bool),X461) = s(fun(X460,bool),X462) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X463))) | s(bool,'T_0') = s(bool,g(s(X460,sK19(X461,X462,bool,X460))))) ) | $spl0),
  inference(superposition,[],[f379,f9373])).
fof(f13543,plain,(
  ( ! [X280,X281] : (s(bool,X280) != s(bool,X281) | s(bool,X281) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X280)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12516,f437])).
fof(f12516,plain,(
  ( ! [X280,X281] : (s(bool,X280) != s(bool,X281) | s(bool,X281) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X280))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f10142,f9373])).
fof(f12412,plain,(
  ( ! [X132,X133,X131] : (s(bool,'T_0') = s(bool,X131) | s(bool,X133) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X131),s('type.list.list'(bool),X132))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X133)))) ) | $spl0),
  inference(superposition,[],[f9373,f343])).
fof(f12411,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X128),s('type.num.num',X129))) | s(bool,X130) = s(bool,'const.ind_type.FCONS_3'(s(bool,X127),s(fun('type.num.num',bool),X128),s('type.num.num','const.num.SUC_1'(s('type.num.num',X129))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X130)))) ) | $spl0),
  inference(superposition,[],[f9373,f356])).
fof(f12410,plain,(
  ( ! [X125,X126,X124] : (s(bool,'T_0') = s(bool,X124) | s(bool,X126) = s(bool,'const.ind_type.FCONS_3'(s(bool,X124),s(fun('type.num.num',bool),X125),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X126)))) ) | $spl0),
  inference(superposition,[],[f9373,f344])).
fof(f12409,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',X120),s('type.list.list'(bool),X122))) | s(bool,X123) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X120))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X121),s('type.list.list'(bool),X122))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X123)))) ) | $spl0),
  inference(superposition,[],[f9373,f357])).
fof(f12408,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',X117),s('type.list.list'(bool),X118))) | s(bool,X119) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X117))),s('type.list.list'(bool),X118))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X119)))) ) | $spl0),
  inference(superposition,[],[f9373,f346])).
fof(f12407,plain,(
  ( ! [X116,X115] : (s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X115))) | s(bool,X116) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X115))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X116)))) ) | $spl0),
  inference(superposition,[],[f9373,f338])).
fof(f13476,plain,(
  ( ! [X99,X100] : (s(bool,X100) = s(bool,'const.marker.Cong_1'(s(bool,X99))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X100))) | s(bool,X99) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f12401,f437])).
fof(f12401,plain,(
  ( ! [X99,X100] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X100) = s(bool,'const.marker.Cong_1'(s(bool,X99))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X100))) | s(bool,X99) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f9373,f4389])).
fof(f13475,plain,(
  ( ! [X94,X95,X93] : (s(bool,X94) = s(bool,'const.marker.Cong_1'(s(bool,X93))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X94))) | s(bool,X93) = s(bool,'const.bool.BOUNDED_1'(s(bool,X95)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12399,f437])).
fof(f12399,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X94) = s(bool,'const.marker.Cong_1'(s(bool,X93))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X94))) | s(bool,X93) = s(bool,'const.bool.BOUNDED_1'(s(bool,X95)))) ) | $spl0),
  inference(superposition,[],[f9373,f1945])).
fof(f12393,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X78),s('type.num.num',X79))) | s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X78))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X80)))) ) | $spl0),
  inference(superposition,[],[f9373,f360])).
fof(f13473,plain,(
  ( ! [X76,X77] : (s(bool,X77) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X77)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12392,f437])).
fof(f12392,plain,(
  ( ! [X76,X77] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X77) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X77)))) ) | $spl0),
  inference(superposition,[],[f9373,f6330])).
fof(f13472,plain,(
  ( ! [X74,X73] : (s(bool,X74) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X74))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X73)) ) | $spl0),
  inference(subsumption_resolution,[],[f12391,f437])).
fof(f12391,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X74) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X74))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X73)) ) | $spl0),
  inference(superposition,[],[f9373,f6980])).
fof(f13471,plain,(
  ( ! [X72,X71] : (s(bool,X72) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X71),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X72)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12390,f437])).
fof(f12390,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X72) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X71),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X72)))) ) | $spl0),
  inference(superposition,[],[f9373,f462])).
fof(f13470,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,X70) = s(bool,'const.list.LIST_REL_3'(s(fun(X65,fun(X66,bool)),X67),s('type.list.list'(X65),'const.list.NIL_0'),s('type.list.list'(X66),'const.list.CONS_2'(s(X66,X68),s('type.list.list'(X66),X69))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X70)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12389,f437])).
fof(f12389,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X70) = s(bool,'const.list.LIST_REL_3'(s(fun(X65,fun(X66,bool)),X67),s('type.list.list'(X65),'const.list.NIL_0'),s('type.list.list'(X66),'const.list.CONS_2'(s(X66,X68),s('type.list.list'(X66),X69))))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X70)))) ) | $spl0),
  inference(superposition,[],[f9373,f380])).
fof(f12358,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X2))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X3)))) ) | $spl0),
  inference(superposition,[],[f9373,f9373])).
fof(f13454,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.marker.Cong_1'(s(bool,X50))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X52)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12348,f437])).
fof(f12348,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.marker.Cong_1'(s(bool,X50))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X52)))) ) | $spl0),
  inference(superposition,[],[f9373,f1004])).
fof(f13453,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,'const.marker.Cong_1'(s(bool,X32))) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(subsumption_resolution,[],[f12342,f437])).
fof(f12342,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,'const.marker.Cong_1'(s(bool,X32))) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(superposition,[],[f9373,f879])).
fof(f13452,plain,(
  ( ! [X17,X16] : (s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X16) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f12335,f437])).
fof(f12335,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X16) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f9373,f4389])).
fof(f13451,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,'const.bool.BOUNDED_1'(s(bool,X12)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12333,f437])).
fof(f12333,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,'const.bool.BOUNDED_1'(s(bool,X12)))) ) | $spl0),
  inference(superposition,[],[f9373,f1945])).
fof(f13449,plain,(
  ( ! [X146,X144,X145] : (s(bool,'T_0') = s(bool,X144) | s(bool,X146) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X144),s('type.list.list'(bool),X145))))) | s(bool,'T_0') = s(bool,X146)) ) | $spl0),
  inference(forward_demodulation,[],[f12328,f271])).
fof(f12328,plain,(
  ( ! [X146,X144,X145] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X144))) | s(bool,X146) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X144),s('type.list.list'(bool),X145))))) | s(bool,'T_0') = s(bool,X146)) ) | $spl0),
  inference(superposition,[],[f9373,f343])).
fof(f13448,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X141),s('type.num.num',X142))) | s(bool,X143) = s(bool,'const.ind_type.FCONS_3'(s(bool,X140),s(fun('type.num.num',bool),X141),s('type.num.num','const.num.SUC_1'(s('type.num.num',X142))))) | s(bool,'T_0') = s(bool,X143)) ) | $spl0),
  inference(forward_demodulation,[],[f12327,f271])).
fof(f12327,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun('type.num.num',bool),X141),s('type.num.num',X142))))) | s(bool,X143) = s(bool,'const.ind_type.FCONS_3'(s(bool,X140),s(fun('type.num.num',bool),X141),s('type.num.num','const.num.SUC_1'(s('type.num.num',X142))))) | s(bool,'T_0') = s(bool,X143)) ) | $spl0),
  inference(superposition,[],[f9373,f356])).
fof(f13447,plain,(
  ( ! [X138,X139,X137] : (s(bool,'T_0') = s(bool,X137) | s(bool,X139) = s(bool,'const.ind_type.FCONS_3'(s(bool,X137),s(fun('type.num.num',bool),X138),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X139)) ) | $spl0),
  inference(forward_demodulation,[],[f12326,f271])).
fof(f12326,plain,(
  ( ! [X138,X139,X137] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X137))) | s(bool,X139) = s(bool,'const.ind_type.FCONS_3'(s(bool,X137),s(fun('type.num.num',bool),X138),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X139)) ) | $spl0),
  inference(superposition,[],[f9373,f344])).
fof(f13446,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',X133),s('type.list.list'(bool),X135))) | s(bool,X136) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X133))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X134),s('type.list.list'(bool),X135))))) | s(bool,'T_0') = s(bool,X136)) ) | $spl0),
  inference(forward_demodulation,[],[f12325,f271])).
fof(f12325,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X133),s('type.list.list'(bool),X135))))) | s(bool,X136) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X133))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X134),s('type.list.list'(bool),X135))))) | s(bool,'T_0') = s(bool,X136)) ) | $spl0),
  inference(superposition,[],[f9373,f357])).
fof(f13445,plain,(
  ( ! [X132,X130,X131] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',X130),s('type.list.list'(bool),X131))) | s(bool,X132) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X130))),s('type.list.list'(bool),X131))) | s(bool,'T_0') = s(bool,X132)) ) | $spl0),
  inference(forward_demodulation,[],[f12324,f271])).
fof(f12324,plain,(
  ( ! [X132,X130,X131] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X130),s('type.list.list'(bool),X131))))) | s(bool,X132) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X130))),s('type.list.list'(bool),X131))) | s(bool,'T_0') = s(bool,X132)) ) | $spl0),
  inference(superposition,[],[f9373,f346])).
fof(f13444,plain,(
  ( ! [X128,X129] : (s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X128))) | s(bool,X129) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X128))) | s(bool,'T_0') = s(bool,X129)) ) | $spl0),
  inference(forward_demodulation,[],[f12323,f271])).
fof(f12323,plain,(
  ( ! [X128,X129] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X128))))) | s(bool,X129) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X128))) | s(bool,'T_0') = s(bool,X129)) ) | $spl0),
  inference(superposition,[],[f9373,f338])).
fof(f13441,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X85),s('type.num.num',X86))) | s(bool,X87) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X85))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X86))))) | s(bool,'T_0') = s(bool,X87)) ) | $spl0),
  inference(forward_demodulation,[],[f12304,f271])).
fof(f12304,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X85),s('type.num.num',X86))))) | s(bool,X87) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X85))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X86))))) | s(bool,'T_0') = s(bool,X87)) ) | $spl0),
  inference(superposition,[],[f9373,f360])).
fof(f12214,plain,(
  ( ! [X144,X145,X143] : (s(bool,X143) = s(bool,X145) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X145))) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X143),s('type.list.list'(bool),X144)))))) ) | $spl0),
  inference(superposition,[],[f9373,f343])).
fof(f12213,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,X142) = s(bool,happ(s(fun('type.num.num',bool),X140),s('type.num.num',X141))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X142))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X139),s(fun('type.num.num',bool),X140),s('type.num.num','const.num.SUC_1'(s('type.num.num',X141)))))) ) | $spl0),
  inference(superposition,[],[f9373,f356])).
fof(f12212,plain,(
  ( ! [X138,X136,X137] : (s(bool,X136) = s(bool,X138) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X138))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X136),s(fun('type.num.num',bool),X137),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f9373,f344])).
fof(f12211,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,X135) = s(bool,'const.list.EL_2'(s('type.num.num',X132),s('type.list.list'(bool),X134))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X135))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X132))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X133),s('type.list.list'(bool),X134)))))) ) | $spl0),
  inference(superposition,[],[f9373,f357])).
fof(f12210,plain,(
  ( ! [X130,X131,X129] : (s(bool,X131) = s(bool,'const.list.EL_2'(s('type.num.num',X129),s('type.list.list'(bool),X130))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X131))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X129))),s('type.list.list'(bool),X130)))) ) | $spl0),
  inference(superposition,[],[f9373,f346])).
fof(f12209,plain,(
  ( ! [X127,X128] : (s(bool,X128) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X127))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X128))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X127)))) ) | $spl0),
  inference(superposition,[],[f9373,f338])).
fof(f12190,plain,(
  ( ! [X85,X86,X84] : (s(bool,X86) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X84),s('type.num.num',X85))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X86))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X84))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X85)))))) ) | $spl0),
  inference(superposition,[],[f9373,f360])).
fof(f13364,plain,(
  ( ! [X57,X56,X55] : (s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X56))) | s(bool,X55) = s(bool,X57) | s(bool,'F_0') = s(bool,X57)) ) | $spl0),
  inference(subsumption_resolution,[],[f12179,f437])).
fof(f12179,plain,(
  ( ! [X57,X56,X55] : (s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X56))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X55) = s(bool,X57) | s(bool,'F_0') = s(bool,X57)) ) | $spl0),
  inference(superposition,[],[f9373,f879])).
fof(f13363,plain,(
  ( ! [X54,X52,X53] : (s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X53))) | s(bool,X52) = s(bool,X54) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X54)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12178,f437])).
fof(f12178,plain,(
  ( ! [X54,X52,X53] : (s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X53))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X52) = s(bool,X54) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X54)))) ) | $spl0),
  inference(superposition,[],[f9373,f1004])).
fof(f11918,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') != s(bool,f(s(X54,sK19(X55,X56,bool,X54)))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,g(s(X54,sK19(X55,X56,bool,X54))))))) ) | $spl0),
  inference(forward_demodulation,[],[f11903,f271])).
fof(f11903,plain,(
  ( ! [X54,X56,X55] : (s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) != s(bool,f(s(X54,sK19(X55,X56,bool,X54)))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,g(s(X54,sK19(X55,X56,bool,X54))))))) ) | $spl0),
  inference(superposition,[],[f379,f8912])).
fof(f11899,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'const.marker.Cong_1'(s(bool,X44))) != s(bool,f(s(X41,sK19(X42,X43,bool,X41)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,'T_0') = s(bool,g(s(X41,sK19(X42,X43,bool,X41)))) | s(bool,'T_0') = s(bool,X44)) ) | $spl0),
  inference(superposition,[],[f379,f2308])).
fof(f11896,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,f(s(X30,sK19(X31,X32,bool,X30)))) | s(fun(X30,bool),X31) = s(fun(X30,bool),X32) | s(bool,X33) = s(bool,g(s(X30,sK19(X31,X32,bool,X30)))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X33)))) ) | $spl0),
  inference(superposition,[],[f379,f1004])).
fof(f11895,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X29) != s(bool,f(s(X26,sK19(X27,X28,bool,X26)))) | s(fun(X26,bool),X27) = s(fun(X26,bool),X28) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,g(s(X26,sK19(X27,X28,bool,X26))))))) ) | $spl0),
  inference(superposition,[],[f379,f1004])).
fof(f11894,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X25) != s(bool,f(s(X22,sK19(X23,X24,bool,X22)))) | s(fun(X22,bool),X23) = s(fun(X22,bool),X24) | s(bool,'F_0') = s(bool,g(s(X22,sK19(X23,X24,bool,X22)))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X25)))) ) | $spl0),
  inference(superposition,[],[f379,f1004])).
fof(f11891,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X13) != s(bool,f(s(X10,sK19(X11,X12,bool,X10)))) | s(fun(X10,bool),X11) = s(fun(X10,bool),X12) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,g(s(X10,sK19(X11,X12,bool,X10))))) ) | $spl0),
  inference(superposition,[],[f379,f879])).
fof(f11890,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(X6,sK19(X7,X8,bool,X6)))) | s(fun(X6,bool),X7) = s(fun(X6,bool),X8) | s(bool,X9) = s(bool,g(s(X6,sK19(X7,X8,bool,X6)))) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f379,f879])).
fof(f11912,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK19(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'F_0') = s(bool,g(s(X3,sK19(X4,X5,bool,X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f11889,f447])).
fof(f11889,plain,(
  ( ! [X4,X5,X3] : (s(bool,f(s(X3,sK19(X4,X5,bool,X3)))) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'F_0') = s(bool,g(s(X3,sK19(X4,X5,bool,X3))))) ) | $spl0),
  inference(superposition,[],[f379,f785])).
fof(f11911,plain,(
  ( ! [X54,X56,X55] : (s(X56,f(s(bool,'F_0'))) != s(X56,g(s(bool,'F_0'))) | s(fun(bool,X56),X54) = s(fun(bool,X56),X55) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,sK19(X54,X55,X56,bool))))) ) | $spl0),
  inference(forward_demodulation,[],[f11887,f271])).
fof(f11887,plain,(
  ( ! [X54,X56,X55] : (s(X56,f(s(bool,'const.marker.Cong_1'(s(bool,'F_0'))))) != s(X56,g(s(bool,'const.marker.Cong_1'(s(bool,'F_0'))))) | s(fun(bool,X56),X54) = s(fun(bool,X56),X55) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,sK19(X54,X55,X56,bool))))) ) | $spl0),
  inference(superposition,[],[f379,f8912])).
fof(f11907,plain,(
  ( ! [X43,X41,X44,X42] : (s(X43,f(s(bool,X44))) != s(X43,g(s(bool,X44))) | s(fun(bool,X43),X41) = s(fun(bool,X43),X42) | s(bool,'T_0') = s(bool,sK19(X41,X42,X43,bool)) | s(bool,'T_0') = s(bool,X44)) ) | $spl0),
  inference(forward_demodulation,[],[f11883,f271])).
fof(f11883,plain,(
  ( ! [X43,X41,X44,X42] : (s(X43,f(s(bool,'const.marker.Cong_1'(s(bool,X44))))) != s(X43,g(s(bool,'const.marker.Cong_1'(s(bool,X44))))) | s(fun(bool,X43),X41) = s(fun(bool,X43),X42) | s(bool,'T_0') = s(bool,sK19(X41,X42,X43,bool)) | s(bool,'T_0') = s(bool,X44)) ) | $spl0),
  inference(superposition,[],[f379,f2308])).
fof(f11880,plain,(
  ( ! [X30,X33,X31,X32] : (s(X32,f(s(bool,'F_0'))) != s(X32,g(s(bool,'F_0'))) | s(fun(bool,X32),X30) = s(fun(bool,X32),X31) | s(bool,X33) = s(bool,sK19(X30,X31,X32,bool)) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X33)))) ) | $spl0),
  inference(superposition,[],[f379,f1004])).
fof(f11879,plain,(
  ( ! [X28,X26,X29,X27] : (s(X28,f(s(bool,X29))) != s(X28,g(s(bool,X29))) | s(fun(bool,X28),X26) = s(fun(bool,X28),X27) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,sK19(X26,X27,X28,bool))))) ) | $spl0),
  inference(superposition,[],[f379,f1004])).
fof(f11878,plain,(
  ( ! [X24,X23,X25,X22] : (s(X24,f(s(bool,X25))) != s(X24,g(s(bool,X25))) | s(fun(bool,X24),X22) = s(fun(bool,X24),X23) | s(bool,'F_0') = s(bool,sK19(X22,X23,X24,bool)) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X25)))) ) | $spl0),
  inference(superposition,[],[f379,f1004])).
fof(f11875,plain,(
  ( ! [X12,X10,X13,X11] : (s(X12,f(s(bool,X13))) != s(X12,g(s(bool,X13))) | s(fun(bool,X12),X10) = s(fun(bool,X12),X11) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,sK19(X10,X11,X12,bool))) ) | $spl0),
  inference(superposition,[],[f379,f879])).
fof(f11874,plain,(
  ( ! [X6,X8,X7,X9] : (s(X8,f(s(bool,'F_0'))) != s(X8,g(s(bool,'F_0'))) | s(fun(bool,X8),X6) = s(fun(bool,X8),X7) | s(bool,X9) = s(bool,sK19(X6,X7,X8,bool)) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f379,f879])).
fof(f11904,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'F_0') = s(bool,sK19(X3,X4,X5,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f11873,f447])).
fof(f11873,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))))) != s(X5,g(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'F_0') = s(bool,sK19(X3,X4,X5,bool))) ) | $spl0),
  inference(superposition,[],[f379,f785])).
fof(f11743,plain,(
  ( ! [X127,X128,X126] : (s(bool,X126) != s(bool,X128) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X128))) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X126),s('type.list.list'(bool),X127)))))) ) | $spl0),
  inference(superposition,[],[f3336,f343])).
fof(f11742,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,X125) != s(bool,happ(s(fun('type.num.num',bool),X123),s('type.num.num',X124))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X125))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X122),s(fun('type.num.num',bool),X123),s('type.num.num','const.num.SUC_1'(s('type.num.num',X124)))))) ) | $spl0),
  inference(superposition,[],[f3336,f356])).
fof(f11741,plain,(
  ( ! [X121,X120,X119] : (s(bool,X119) != s(bool,X121) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X121))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f3336,f344])).
fof(f11740,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,X118) != s(bool,'const.list.EL_2'(s('type.num.num',X115),s('type.list.list'(bool),X117))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X118))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X115))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X116),s('type.list.list'(bool),X117)))))) ) | $spl0),
  inference(superposition,[],[f3336,f357])).
fof(f11739,plain,(
  ( ! [X114,X112,X113] : (s(bool,X114) != s(bool,'const.list.EL_2'(s('type.num.num',X112),s('type.list.list'(bool),X113))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X114))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X112))),s('type.list.list'(bool),X113)))) ) | $spl0),
  inference(superposition,[],[f3336,f346])).
fof(f11738,plain,(
  ( ! [X111,X110] : (s(bool,X111) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X110))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X111))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X110)))) ) | $spl0),
  inference(superposition,[],[f3336,f338])).
fof(f11720,plain,(
  ( ! [X70,X72,X71] : (s(bool,X72) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X72))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X71)))))) ) | $spl0),
  inference(superposition,[],[f3336,f360])).
fof(f11439,plain,(
  ( ! [X461,X460,X462] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.list.EL_2'(s('type.num.num',X460),s('type.list.list'('type.num.num'),X462))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X460))),s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X461),s('type.list.list'('type.num.num'),X462)))))) ) | $spl0),
  inference(superposition,[],[f7138,f357])).
fof(f11437,plain,(
  ( ! [X455,X454,X456] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.list.EL_2'(s('type.num.num',X454),s('type.list.list'('type.num.num'),X456))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X454))),s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X455),s('type.list.list'('type.num.num'),X456)))))) ) | $spl0),
  inference(superposition,[],[f6980,f357])).
fof(f11390,plain,(
  ( ! [X251,X249,X250,X248] : (s(bool,X251) != s(bool,'const.list.EL_2'(s('type.num.num',X248),s('type.list.list'(bool),X250))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X248))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X249),s('type.list.list'(bool),X250))))) | s(bool,'T_0') = s(bool,X251)) ) | $spl0),
  inference(superposition,[],[f6213,f357])).
fof(f11379,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,X213) != s(bool,'const.list.EL_2'(s('type.num.num',X210),s('type.list.list'(bool),X212))) | s(bool,X213) = s(bool,'const.bool.BOUNDED_1'(s(bool,X214))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X210))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X211),s('type.list.list'(bool),X212)))))) ) | $spl0),
  inference(superposition,[],[f3201,f357])).
fof(f11376,plain,(
  ( ! [X198,X200,X199,X201] : (s(bool,X201) != s(bool,'const.list.EL_2'(s('type.num.num',X198),s('type.list.list'(bool),X200))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X201))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X198))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X199),s('type.list.list'(bool),X200)))))) ) | $spl0),
  inference(superposition,[],[f2374,f357])).
fof(f11494,plain,(
  ( ! [X189,X187,X188,X186] : (s(bool,X189) = s(bool,'const.list.EL_2'(s('type.num.num',X186),s('type.list.list'(bool),X188))) | s(bool,'T_0') = s(bool,X189) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X186))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X187),s('type.list.list'(bool),X188)))))) ) | $spl0),
  inference(forward_demodulation,[],[f11373,f271])).
fof(f11373,plain,(
  ( ! [X189,X187,X188,X186] : (s(bool,X189) = s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X186),s('type.list.list'(bool),X188))))) | s(bool,'T_0') = s(bool,X189) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X186))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X187),s('type.list.list'(bool),X188)))))) ) | $spl0),
  inference(superposition,[],[f2308,f357])).
fof(f11493,plain,(
  ( ! [X185,X184,X183] : (s(bool,'T_0') != s(bool,'const.list.EL_2'(s('type.num.num',X183),s('type.list.list'(bool),X185))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X183))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X184),s('type.list.list'(bool),X185)))))) ) | $spl0),
  inference(forward_demodulation,[],[f11372,f271])).
fof(f11372,plain,(
  ( ! [X185,X184,X183] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X183),s('type.list.list'(bool),X185))))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X183))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X184),s('type.list.list'(bool),X185)))))) ) | $spl0),
  inference(superposition,[],[f2178,f357])).
fof(f11490,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,X179) != s(bool,'const.list.EL_2'(s('type.num.num',X176),s('type.list.list'(bool),X178))) | s(bool,'F_0') = s(bool,X179) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X176))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X177),s('type.list.list'(bool),X178)))))) ) | $spl0),
  inference(forward_demodulation,[],[f11370,f271])).
fof(f11370,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,X179) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X176),s('type.list.list'(bool),X178))))) | s(bool,'F_0') = s(bool,X179) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X176))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X177),s('type.list.list'(bool),X178)))))) ) | $spl0),
  inference(superposition,[],[f2047,f357])).
fof(f11369,plain,(
  ( ! [X175,X173,X174,X172] : (s(bool,'const.marker.Cong_1'(s(bool,X175))) != s(bool,'const.list.EL_2'(s('type.num.num',X172),s('type.list.list'(bool),X174))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X172))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X173),s('type.list.list'(bool),X174))))) | s(bool,'T_0') = s(bool,X175)) ) | $spl0),
  inference(superposition,[],[f2047,f357])).
fof(f11477,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X140) != s(bool,'const.list.EL_2'(s('type.num.num',X137),s('type.list.list'(bool),X139))) | s(bool,X140) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X137))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X138),s('type.list.list'(bool),X139))))) | s(bool,'F_0') = s(bool,X140)) ) | $spl0),
  inference(forward_demodulation,[],[f11359,f271])).
fof(f11359,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X140) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X137),s('type.list.list'(bool),X139))))) | s(bool,X140) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X137))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X138),s('type.list.list'(bool),X139))))) | s(bool,'F_0') = s(bool,X140)) ) | $spl0),
  inference(superposition,[],[f1798,f357])).
fof(f11358,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'const.marker.Cong_1'(s(bool,X136))) != s(bool,'const.list.EL_2'(s('type.num.num',X133),s('type.list.list'(bool),X135))) | s(bool,X136) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X133))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X134),s('type.list.list'(bool),X135))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X133))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X134),s('type.list.list'(bool),X135)))))) ) | $spl0),
  inference(superposition,[],[f1798,f357])).
fof(f11357,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.list.EL_2'(s('type.num.num',X129),s('type.list.list'(bool),X131))) | s(bool,X132) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X129))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X130),s('type.list.list'(bool),X131))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X132)))) ) | $spl0),
  inference(superposition,[],[f1710,f357])).
fof(f11355,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,'T_0') != s(bool,'const.list.EL_2'(s('type.num.num',X121),s('type.list.list'(bool),X123))) | s(bool,X124) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X121))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X122),s('type.list.list'(bool),X123))))) | s(bool,'F_0') = s(bool,X124)) ) | $spl0),
  inference(superposition,[],[f1136,f357])).
fof(f11320,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'const.marker.Cong_1'(s(bool,X48))) = s(bool,'const.list.EL_2'(s('type.num.num',X45),s('type.list.list'(bool),X47))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X46),s('type.list.list'(bool),X47))))) | s(bool,'T_0') = s(bool,X48)) ) | $spl0),
  inference(superposition,[],[f357,f2308])).
fof(f11317,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',X34),s('type.list.list'(bool),X36))) | s(bool,X37) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X35),s('type.list.list'(bool),X36))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X37)))) ) | $spl0),
  inference(superposition,[],[f357,f1004])).
fof(f11315,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X29) = s(bool,'const.list.EL_2'(s('type.num.num',X26),s('type.list.list'(bool),X28))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X27),s('type.list.list'(bool),X28))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X29)))) ) | $spl0),
  inference(superposition,[],[f357,f1004])).
fof(f11312,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) = s(bool,'const.list.EL_2'(s('type.num.num',X14),s('type.list.list'(bool),X16))) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X15),s('type.list.list'(bool),X16)))))) ) | $spl0),
  inference(superposition,[],[f357,f879])).
fof(f11311,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(bool),X12))) | s(bool,X13) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X11),s('type.list.list'(bool),X12))))) | s(bool,'F_0') = s(bool,X13)) ) | $spl0),
  inference(superposition,[],[f357,f879])).
fof(f11448,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',X7),s('type.list.list'(bool),X9))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X8),s('type.list.list'(bool),X9)))))) ) | $spl0),
  inference(forward_demodulation,[],[f11310,f447])).
fof(f11310,plain,(
  ( ! [X8,X7,X9] : (s(bool,'const.list.EL_2'(s('type.num.num',X7),s('type.list.list'(bool),X9))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X8),s('type.list.list'(bool),X9)))))) ) | $spl0),
  inference(superposition,[],[f357,f785])).
fof(f10901,plain,(
  ( ! [X175,X173,X174,X172] : (s(bool,X174) != s(bool,'const.list.EL_2'(s('type.num.num',X172),s('type.list.list'(bool),X173))) | s(bool,X174) = s(bool,'const.bool.BOUNDED_1'(s(bool,X175))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X172))),s('type.list.list'(bool),X173)))) ) | $spl0),
  inference(superposition,[],[f3201,f346])).
fof(f10900,plain,(
  ( ! [X171,X169,X170] : (s(bool,X170) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X169))) | s(bool,X170) = s(bool,'const.bool.BOUNDED_1'(s(bool,X171))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X169)))) ) | $spl0),
  inference(superposition,[],[f3201,f338])).
fof(f10899,plain,(
  ( ! [X167,X165,X166,X164,X168] : (s(bool,X167) != s(bool,happ(s(fun('type.num.num',bool),X165),s('type.num.num',X166))) | s(bool,X167) = s(bool,'const.bool.BOUNDED_1'(s(bool,X168))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X164),s(fun('type.num.num',bool),X165),s('type.num.num','const.num.SUC_1'(s('type.num.num',X166)))))) ) | $spl0),
  inference(superposition,[],[f3201,f356])).
fof(f10898,plain,(
  ( ! [X163,X161,X162,X160] : (s(bool,X160) != s(bool,X162) | s(bool,X162) = s(bool,'const.bool.BOUNDED_1'(s(bool,X163))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X160),s(fun('type.num.num',bool),X161),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f3201,f344])).
fof(f10897,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,X156) != s(bool,X158) | s(bool,X158) = s(bool,'const.bool.BOUNDED_1'(s(bool,X159))) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X156),s('type.list.list'(bool),X157)))))) ) | $spl0),
  inference(superposition,[],[f3201,f343])).
fof(f10879,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,X100) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X98),s('type.num.num',X99))) | s(bool,X100) = s(bool,'const.bool.BOUNDED_1'(s(bool,X101))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X98))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X99)))))) ) | $spl0),
  inference(superposition,[],[f3201,f360])).
fof(f10824,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X7)))) | s(X4,X6) != s(X4,X7) | s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) ) | $spl0),
  inference(subsumption_resolution,[],[f10823,f264])).
fof(f10823,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X7)))) | s(X4,X6) != s(X4,X7) | s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) ) | $spl0),
  inference(forward_demodulation,[],[f10740,f447])).
fof(f10740,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X7)))) | s(X4,X6) != s(X4,X7) | s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) ) | $spl0),
  inference(superposition,[],[f353,f785])).
fof(f10713,plain,(
  ( ! [X118,X119,X117] : (~p(s(bool,happ(s(fun(bool,bool),X118),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X118),s(bool,X119)))) | s(bool,X119) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X117),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X117)) ) | $spl0),
  inference(superposition,[],[f353,f6980])).
fof(f10786,plain,(
  ( ! [X99,X97,X98] : (s(bool,'T_0') != s(bool,X99) | ~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X97)))) | p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X99)))) | s(bool,'F_0') = s(bool,X97)) ) | $spl0),
  inference(forward_demodulation,[],[f10708,f447])).
fof(f10708,plain,(
  ( ! [X99,X97,X98] : (~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X97)))) | p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X99)))) | s(bool,X99) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,X97)) ) | $spl0),
  inference(superposition,[],[f353,f785])).
fof(f10785,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'F_0') = s(bool,X85) | ~p(s(bool,happ(s(fun(bool,bool),X83),s(bool,X82)))) | p(s(bool,happ(s(fun(bool,bool),X83),s(bool,X84)))) | s(bool,'F_0') != s(bool,X84) | s(bool,X82) = s(bool,X85)) ) | $spl0),
  inference(forward_demodulation,[],[f10704,f271])).
fof(f10704,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,happ(s(fun(bool,bool),X83),s(bool,X82)))) | p(s(bool,happ(s(fun(bool,bool),X83),s(bool,X84)))) | s(bool,'F_0') != s(bool,X84) | s(bool,X82) = s(bool,X85) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X85)))) ) | $spl0),
  inference(superposition,[],[f353,f1004])).
fof(f10778,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,happ(s(fun(bool,bool),X68),s(bool,X67)))) | p(s(bool,happ(s(fun(bool,bool),X68),s(bool,X69)))) | s(bool,X66) != s(bool,X69) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X67)) ) | $spl0),
  inference(forward_demodulation,[],[f10699,f271])).
fof(f10699,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,happ(s(fun(bool,bool),X68),s(bool,'const.marker.Cong_1'(s(bool,X67)))))) | p(s(bool,happ(s(fun(bool,bool),X68),s(bool,X69)))) | s(bool,X66) != s(bool,X69) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X67)) ) | $spl0),
  inference(superposition,[],[f353,f2308])).
fof(f10691,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X36)))) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X38)))) | s(bool,X35) != s(bool,X38) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X35)) ) | $spl0),
  inference(superposition,[],[f353,f879])).
fof(f10690,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X33)))) | s(bool,X31) != s(bool,X33) | s(bool,X31) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(superposition,[],[f353,f879])).
fof(f10770,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X28)) ) | $spl0),
  inference(forward_demodulation,[],[f10689,f447])).
fof(f10689,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))))) | p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X28)) ) | $spl0),
  inference(superposition,[],[f353,f785])).
fof(f10532,plain,(
  ( ! [X127,X125,X126] : (s(bool,X127) != s(bool,'const.list.EL_2'(s('type.num.num',X125),s('type.list.list'(bool),X126))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X127))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X125))),s('type.list.list'(bool),X126)))) ) | $spl0),
  inference(superposition,[],[f2374,f346])).
fof(f10531,plain,(
  ( ! [X123,X124] : (s(bool,X124) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X123))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X124))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X123)))) ) | $spl0),
  inference(superposition,[],[f2374,f338])).
fof(f10530,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,X122) != s(bool,happ(s(fun('type.num.num',bool),X120),s('type.num.num',X121))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X122))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.SUC_1'(s('type.num.num',X121)))))) ) | $spl0),
  inference(superposition,[],[f2374,f356])).
fof(f10529,plain,(
  ( ! [X118,X116,X117] : (s(bool,X116) != s(bool,X118) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X118))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X116),s(fun('type.num.num',bool),X117),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f2374,f344])).
fof(f10528,plain,(
  ( ! [X114,X115,X113] : (s(bool,X113) != s(bool,X115) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X115))) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X113),s('type.list.list'(bool),X114)))))) ) | $spl0),
  inference(superposition,[],[f2374,f343])).
fof(f10510,plain,(
  ( ! [X74,X75,X73] : (s(bool,X75) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num',X74))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X75))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X73))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74)))))) ) | $spl0),
  inference(superposition,[],[f2374,f360])).
fof(f10340,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X10) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f10142,f276])).
fof(f10338,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f10142,f879])).
fof(f10337,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f10142,f1004])).
fof(f10310,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X32)))) | ~p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X30)))) | s(bool,'T_0') = s(bool,sK9(X30,X31,bool)) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(forward_demodulation,[],[f10271,f271])).
fof(f10271,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,happ(s(fun(bool,bool),X31),s(bool,'const.marker.Cong_1'(s(bool,X32)))))) | ~p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X30)))) | s(bool,'T_0') = s(bool,sK9(X30,X31,bool)) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f355,f2308])).
fof(f10142,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f10054,f271])).
fof(f10186,plain,(
  ( ! [X14,X15] : (s(bool,X14) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X14) = s(bool,X15) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X15)))) ) | $spl0),
  inference(subsumption_resolution,[],[f10099,f271])).
fof(f10099,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X14) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X14) = s(bool,X15) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X15)))) ) | $spl0),
  inference(superposition,[],[f10054,f1004])).
fof(f10054,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X35))) | s(bool,X35) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f9703,f264])).
fof(f9703,plain,(
  ( ! [X35] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X35))) | s(bool,X35) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f285,f8912])).
fof(f9884,plain,(
  ( ! [X30,X31] : (s(bool,X30) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X31) = s(bool,'const.marker.Cong_1'(s(bool,X30))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X31)))) ) | $spl0),
  inference(subsumption_resolution,[],[f9466,f437])).
fof(f9466,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X30) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X31) = s(bool,'const.marker.Cong_1'(s(bool,X30))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X31)))) ) | $spl0),
  inference(superposition,[],[f8912,f1004])).
fof(f9883,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,X18))) | s(bool,'F_0') = s(bool,X19)) ) | $spl0),
  inference(subsumption_resolution,[],[f9460,f437])).
fof(f9460,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,X18))) | s(bool,'F_0') = s(bool,X19)) ) | $spl0),
  inference(superposition,[],[f8912,f879])).
fof(f9876,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(subsumption_resolution,[],[f9448,f437])).
fof(f9448,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X2) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f8912,f1004])).
fof(f9373,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f8912,f8650])).
fof(f9372,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X0)))) ) | $spl0),
  inference(superposition,[],[f8912,f8912])).
fof(f9361,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,'const.marker.Cong_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f8912,f879])).
fof(f9360,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f8912,f1004])).
fof(f9308,plain,(
  ( ! [X125,X126] : (s(bool,'F_0') = s(bool,sK9('const.marker.Cong_1'(s(bool,'F_0')),X126,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X126),s(bool,X125)))) | s(bool,'T_0') = s(bool,X125)) ) | $spl0),
  inference(forward_demodulation,[],[f9153,f271])).
fof(f9153,plain,(
  ( ! [X125,X126] : (~p(s(bool,happ(s(fun(bool,bool),X126),s(bool,X125)))) | s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) = s(bool,sK9('const.marker.Cong_1'(s(bool,'F_0')),X126,bool)) | s(bool,'T_0') = s(bool,X125)) ) | $spl0),
  inference(superposition,[],[f354,f8650])).
fof(f8877,plain,(
  ( ! [X204,X205,X203] : (s(bool,X203) = s(bool,sK9('const.marker.Cong_1'(s(bool,X203)),X205,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X205),s(bool,X204)))) | s(bool,'T_0') = s(bool,X204) | s(bool,'T_0') = s(bool,X203)) ) | $spl0),
  inference(forward_demodulation,[],[f8588,f271])).
fof(f8588,plain,(
  ( ! [X204,X205,X203] : (~p(s(bool,happ(s(fun(bool,bool),X205),s(bool,X204)))) | s(bool,'const.marker.Cong_1'(s(bool,X203))) = s(bool,sK9('const.marker.Cong_1'(s(bool,X203)),X205,bool)) | s(bool,'T_0') = s(bool,X204) | s(bool,'T_0') = s(bool,X203)) ) | $spl0),
  inference(superposition,[],[f354,f2308])).
fof(f8809,plain,(
  ( ! [X339,X340,X338] : (~p(s(bool,happ(s(fun(bool,bool),X339),s(bool,X340)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X339)))) | s(bool,'T_0') = s(bool,sK6(X338,X339,bool)) | s(bool,'T_0') = s(bool,X340)) ) | $spl0),
  inference(forward_demodulation,[],[f8504,f271])).
fof(f8504,plain,(
  ( ! [X339,X340,X338] : (~p(s(bool,happ(s(fun(bool,bool),X339),s(bool,'const.marker.Cong_1'(s(bool,X340)))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X339)))) | s(bool,'T_0') = s(bool,sK6(X338,X339,bool)) | s(bool,'T_0') = s(bool,X340)) ) | $spl0),
  inference(superposition,[],[f348,f2308])).
fof(f8808,plain,(
  ( ! [X302,X304,X303] : (~p(s(bool,happ(s(fun(bool,bool),X304),s(bool,X303)))) | s(bool,X302) = s(bool,sK9(X302,X304,bool)) | s(bool,'T_0') = s(bool,X302) | s(bool,'T_0') = s(bool,X303)) ) | $spl0),
  inference(forward_demodulation,[],[f8493,f271])).
fof(f8493,plain,(
  ( ! [X302,X304,X303] : (~p(s(bool,happ(s(fun(bool,bool),X304),s(bool,'const.marker.Cong_1'(s(bool,X303)))))) | s(bool,X302) = s(bool,sK9(X302,X304,bool)) | s(bool,'T_0') = s(bool,X302) | s(bool,'T_0') = s(bool,X303)) ) | $spl0),
  inference(superposition,[],[f354,f2308])).
fof(f8696,plain,(
  ( ! [X116,X117,X115] : (s(bool,X117) = s(bool,'const.list.EL_2'(s('type.num.num',X115),s('type.list.list'(bool),X116))) | s(bool,'T_0') = s(bool,X117) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X115))),s('type.list.list'(bool),X116)))) ) | $spl0),
  inference(forward_demodulation,[],[f8361,f271])).
fof(f8361,plain,(
  ( ! [X116,X117,X115] : (s(bool,X117) = s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X115),s('type.list.list'(bool),X116))))) | s(bool,'T_0') = s(bool,X117) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X115))),s('type.list.list'(bool),X116)))) ) | $spl0),
  inference(superposition,[],[f2308,f346])).
fof(f8695,plain,(
  ( ! [X114,X113] : (s(bool,X114) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X113))) | s(bool,'T_0') = s(bool,X114) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X113)))) ) | $spl0),
  inference(forward_demodulation,[],[f8360,f271])).
fof(f8360,plain,(
  ( ! [X114,X113] : (s(bool,X114) = s(bool,'const.marker.Cong_1'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X113))))) | s(bool,'T_0') = s(bool,X114) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X113)))) ) | $spl0),
  inference(superposition,[],[f2308,f338])).
fof(f8694,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X112) = s(bool,happ(s(fun('type.num.num',bool),X110),s('type.num.num',X111))) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X109),s(fun('type.num.num',bool),X110),s('type.num.num','const.num.SUC_1'(s('type.num.num',X111)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8359,f271])).
fof(f8359,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X112) = s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun('type.num.num',bool),X110),s('type.num.num',X111))))) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X109),s(fun('type.num.num',bool),X110),s('type.num.num','const.num.SUC_1'(s('type.num.num',X111)))))) ) | $spl0),
  inference(superposition,[],[f2308,f356])).
fof(f8693,plain,(
  ( ! [X107,X108,X106] : (s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X106),s(fun('type.num.num',bool),X107),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f8358,f271])).
fof(f8358,plain,(
  ( ! [X107,X108,X106] : (s(bool,X108) = s(bool,'const.marker.Cong_1'(s(bool,X106))) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X106),s(fun('type.num.num',bool),X107),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f2308,f344])).
fof(f8692,plain,(
  ( ! [X103,X105,X104] : (s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X103),s('type.list.list'(bool),X104)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8357,f271])).
fof(f8357,plain,(
  ( ! [X103,X105,X104] : (s(bool,X105) = s(bool,'const.marker.Cong_1'(s(bool,X103))) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X103),s('type.list.list'(bool),X104)))))) ) | $spl0),
  inference(superposition,[],[f2308,f343])).
fof(f8687,plain,(
  ( ! [X70,X71,X69] : (s(bool,X71) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X70))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X69))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))))) ) | $spl0),
  inference(forward_demodulation,[],[f8342,f271])).
fof(f8342,plain,(
  ( ! [X70,X71,X69] : (s(bool,X71) = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X70))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X69))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))))) ) | $spl0),
  inference(superposition,[],[f2308,f360])).
fof(f8684,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) = s(bool,'const.marker.Cong_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X39)))) ) | $spl0),
  inference(subsumption_resolution,[],[f8330,f437])).
fof(f8330,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) = s(bool,'const.marker.Cong_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X37) = s(bool,X39) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X39)))) ) | $spl0),
  inference(superposition,[],[f2308,f1004])).
fof(f8315,plain,(
  ( ! [X114,X112,X113] : (s(bool,'const.marker.Cong_1'(s(bool,X114))) = s(bool,'const.list.EL_2'(s('type.num.num',X112),s('type.list.list'(bool),X113))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X112))),s('type.list.list'(bool),X113))) | s(bool,'T_0') = s(bool,X114)) ) | $spl0),
  inference(superposition,[],[f2308,f346])).
fof(f8314,plain,(
  ( ! [X111,X110] : (s(bool,'const.marker.Cong_1'(s(bool,X111))) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X110))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X110))) | s(bool,'T_0') = s(bool,X111)) ) | $spl0),
  inference(superposition,[],[f2308,f338])).
fof(f8313,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'const.marker.Cong_1'(s(bool,X109))) = s(bool,happ(s(fun('type.num.num',bool),X107),s('type.num.num',X108))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X106),s(fun('type.num.num',bool),X107),s('type.num.num','const.num.SUC_1'(s('type.num.num',X108))))) | s(bool,'T_0') = s(bool,X109)) ) | $spl0),
  inference(superposition,[],[f2308,f356])).
fof(f8312,plain,(
  ( ! [X103,X105,X104] : (s(bool,X103) = s(bool,'const.marker.Cong_1'(s(bool,X105))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X103),s(fun('type.num.num',bool),X104),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X105)) ) | $spl0),
  inference(superposition,[],[f2308,f344])).
fof(f8311,plain,(
  ( ! [X101,X102,X100] : (s(bool,X100) = s(bool,'const.marker.Cong_1'(s(bool,X102))) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X100),s('type.list.list'(bool),X101))))) | s(bool,'T_0') = s(bool,X102)) ) | $spl0),
  inference(superposition,[],[f2308,f343])).
fof(f8299,plain,(
  ( ! [X72,X73] : (s(bool,X72) = s(bool,'const.marker.Cong_1'(s(bool,X73))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X73)) ) | $spl0),
  inference(superposition,[],[f2308,f271])).
fof(f8296,plain,(
  ( ! [X68,X66,X67] : (s(bool,'const.marker.Cong_1'(s(bool,X68))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X66),s('type.num.num',X67))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X66))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X67))))) | s(bool,'T_0') = s(bool,X68)) ) | $spl0),
  inference(superposition,[],[f2308,f360])).
fof(f8652,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.marker.Cong_1'(s(bool,X0))) = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(condensation,[],[f8651])).
fof(f8651,plain,(
  ( ! [X30,X31,X32] : (s(bool,'const.marker.Cong_1'(s(bool,X31))) = s(bool,'const.marker.Cong_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f8285])).
fof(f8285,plain,(
  ( ! [X30,X31,X32] : (s(bool,'const.marker.Cong_1'(s(bool,X31))) = s(bool,'const.marker.Cong_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(superposition,[],[f2308,f2308])).
fof(f8187,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X11),s('type.num.num',X12))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X10),s(fun('type.num.num','type.num.num'),X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12)))))) ) | $spl0),
  inference(superposition,[],[f7138,f356])).
fof(f8186,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X8),s(fun('type.num.num','type.num.num'),X9),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f7138,f344])).
fof(f8185,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.HD_1'(s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X6),s('type.list.list'('type.num.num'),X7)))))) ) | $spl0),
  inference(superposition,[],[f7138,f343])).
fof(f7890,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f7074,f360,f6213])).
fof(f7849,plain,(
  ( ! [X92,X91] : (~p(s(bool,happ(s(fun(bool,bool),X92),s(bool,'F_0')))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X91),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num',X91),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X92,bool)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X91)) ) | $spl0),
  inference(superposition,[],[f354,f6980])).
fof(f7874,plain,(
  ( ! [X17,X16] : (s(bool,X17) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X17) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) ) | $spl0),
  inference(subsumption_resolution,[],[f7803,f426])).
fof(f7803,plain,(
  ( ! [X17,X16] : (p(s(bool,'F_0')) | s(bool,X17) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X17) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) ) | $spl0),
  inference(superposition,[],[f546,f6980])).
fof(f7873,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s(bool,X10) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) ) | $spl0),
  inference(subsumption_resolution,[],[f7799,f426])).
fof(f7799,plain,(
  ( ! [X10,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X10) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) ) | $spl0),
  inference(superposition,[],[f361,f6980])).
fof(f7863,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl0),
  inference(condensation,[],[f7862])).
fof(f7862,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,X13) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f7786])).
fof(f7786,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,X13) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl0),
  inference(superposition,[],[f6980,f1004])).
fof(f7778,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X11),s('type.num.num',X12))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X10),s(fun('type.num.num','type.num.num'),X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12)))))) ) | $spl0),
  inference(superposition,[],[f6980,f356])).
fof(f7777,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X8),s(fun('type.num.num','type.num.num'),X9),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f6980,f344])).
fof(f7776,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.HD_1'(s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X6),s('type.list.list'('type.num.num'),X7)))))) ) | $spl0),
  inference(superposition,[],[f6980,f343])).
fof(f7277,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X82))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X83,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f7219,f6330])).
fof(f7219,plain,(
  ( ! [X83,X82] : (~p(s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0')))) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X82))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,sK9('const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X82))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X83,bool))) ) | $spl0),
  inference(superposition,[],[f354,f6330])).
fof(f7138,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl0),
  inference(superposition,[],[f7074,f277])).
fof(f7079,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X15)) ) | $spl0),
  inference(subsumption_resolution,[],[f7000,f426])).
fof(f7000,plain,(
  ( ! [X14,X15] : (p(s(bool,'F_0')) | s(bool,X15) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X15)) ) | $spl0),
  inference(superposition,[],[f546,f6330])).
fof(f7074,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(subsumption_resolution,[],[f6992,f426])).
fof(f6992,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(superposition,[],[f281,f6330])).
fof(f6980,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl0),
  inference(superposition,[],[f6330,f277])).
fof(f4077,plain,(
  ( ! [X12,X10,X13,X11,X9] : (sP0(X9,X10,X11,X13,X12) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X9,fun(X10,bool)),X11),s('type.list.list'(X9),X12),s('type.list.list'(X10),X13)))) ) | $spl0),
  inference(subsumption_resolution,[],[f4076,f264])).
fof(f4076,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,'T_0')) | sP0(X9,X10,X11,X13,X12) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X9,fun(X10,bool)),X11),s('type.list.list'(X9),X12),s('type.list.list'(X10),X13)))) ) | $spl0),
  inference(forward_demodulation,[],[f4060,f447])).
fof(f4060,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) | sP0(X9,X10,X11,X13,X12) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X9,fun(X10,bool)),X11),s('type.list.list'(X9),X12),s('type.list.list'(X10),X13)))) ) | $spl0),
  inference(superposition,[],[f395,f785])).
fof(f2895,plain,(
  ( ! [X94,X95,X93] : (s(X93,X95) = s(X93,sK9(X95,X94,X93)) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun(X93,bool),X94),s(X93,X95)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2894,f264])).
fof(f2894,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'T_0')) | s(X93,X95) = s(X93,sK9(X95,X94,X93)) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun(X93,bool),X94),s(X93,X95)))))) ) | $spl0),
  inference(forward_demodulation,[],[f2686,f272])).
fof(f2686,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'const.bool.BOUNDED_1'(s(bool,X96)))) | s(X93,X95) = s(X93,sK9(X95,X94,X93)) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun(X93,bool),X94),s(X93,X95)))))) ) | $spl0),
  inference(superposition,[],[f354,f1945])).
fof(f5301,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | s(bool,X9) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK9(X8,X7,X6)))) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(subsumption_resolution,[],[f5268,f426])).
fof(f5268,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | s(bool,X9) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK9(X8,X7,X6)))) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f355,f879])).
fof(f2543,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(bool,'F_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK6(X5,X4,X3))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2542,f264])).
fof(f2542,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(bool,'F_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK6(X5,X4,X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f2523,f447])).
fof(f2523,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(bool,'F_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK6(X5,X4,X3))))) ) | $spl0),
  inference(superposition,[],[f348,f785])).
fof(f2101,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X19)))) | s(bool,X18) = s(bool,sK9(X18,X20,bool)) | s(bool,'F_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X18)) ) | $spl0),
  inference(superposition,[],[f354,f879])).
fof(f2142,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') = s(bool,X38) | ~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X36)))) | s(bool,'F_0') = s(bool,sK9('F_0',X37,bool)) | s(bool,X36) = s(bool,X38)) ) | $spl0),
  inference(forward_demodulation,[],[f2107,f271])).
fof(f2107,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X36)))) | s(bool,'F_0') = s(bool,sK9('F_0',X37,bool)) | s(bool,X36) = s(bool,X38) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X38)))) ) | $spl0),
  inference(superposition,[],[f354,f1004])).
fof(f2143,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') = s(bool,sK9('const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)),X48,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X48),s(bool,X47)))) | s(bool,'F_0') = s(bool,X47)) ) | $spl0),
  inference(forward_demodulation,[],[f2111,f447])).
fof(f2111,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(bool,bool),X48),s(bool,X47)))) | s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) = s(bool,sK9('const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)),X48,bool)) | s(bool,'F_0') = s(bool,X47)) ) | $spl0),
  inference(superposition,[],[f354,f785])).
fof(f2516,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'F_0') = s(bool,X9) | s(bool,'F_0') = s(bool,sK6(X7,X8,bool))) ) | $spl0),
  inference(superposition,[],[f348,f879])).
fof(f2943,plain,(
  ( ! [X74,X72,X73] : (s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,sK9('const.bool.BOUNDED_1'(s(bool,X72)),X74,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X74),s(bool,X73))))) ) | $spl0),
  inference(forward_demodulation,[],[f2942,f271])).
fof(f2942,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,sK9('const.bool.BOUNDED_1'(s(bool,X72)),X74,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X74),s(bool,X73)))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X73)))) ) | $spl0),
  inference(forward_demodulation,[],[f2736,f272])).
fof(f2736,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,happ(s(fun(bool,bool),X74),s(bool,X73)))) | s(bool,'const.bool.BOUNDED_1'(s(bool,X72))) = s(bool,sK9('const.bool.BOUNDED_1'(s(bool,X72)),X74,bool)) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X73)))) ) | $spl0),
  inference(superposition,[],[f354,f1945])).
fof(f4463,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,sK9('const.marker.Cong_1'(s(bool,'T_0')),X76,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X76),s(bool,X75)))) | s(bool,'F_0') = s(bool,X75)) ) | $spl0),
  inference(forward_demodulation,[],[f4349,f271])).
fof(f4349,plain,(
  ( ! [X76,X75] : (~p(s(bool,happ(s(fun(bool,bool),X76),s(bool,X75)))) | s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) = s(bool,sK9('const.marker.Cong_1'(s(bool,'T_0')),X76,bool)) | s(bool,'F_0') = s(bool,X75)) ) | $spl0),
  inference(superposition,[],[f354,f4160])).
fof(f5255,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X4)))) | s(bool,X6) = s(bool,sK9(X4,X5,bool)) | s(bool,'F_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f355,f879])).
fof(f5256,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X7)))) | s(bool,'F_0') = s(bool,X9) | s(bool,'F_0') = s(bool,sK9(X7,X8,bool))) ) | $spl0),
  inference(superposition,[],[f355,f879])).
fof(f5288,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) | s(bool,'F_0') = s(bool,sK9(X2,X3,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f5254,f447])).
fof(f5254,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))))) | ~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) | s(bool,'F_0') = s(bool,sK9(X2,X3,bool))) ) | $spl0),
  inference(superposition,[],[f355,f785])).
fof(f2138,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0')))) | s(bool,X13) = s(bool,sK9(X13,X14,bool)) | s(bool,'F_0') = s(bool,X13)) ) | $spl0),
  inference(forward_demodulation,[],[f2099,f447])).
fof(f2099,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))))) | s(bool,X13) = s(bool,sK9(X13,X14,bool)) | s(bool,'F_0') = s(bool,X13)) ) | $spl0),
  inference(superposition,[],[f354,f785])).
fof(f2144,plain,(
  ( ! [X49] : (s(bool,'T_0') = s(bool,sK9('const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)),X49,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,'T_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f2112,f447])).
fof(f2112,plain,(
  ( ! [X49] : (~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,'T_0')))) | s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) = s(bool,sK9('const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)),X49,bool))) ) | $spl0),
  inference(superposition,[],[f354,f447])).
fof(f2538,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3)))) | s(bool,'F_0') = s(bool,sK6(X2,X3,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f2514,f447])).
fof(f2514,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3)))) | s(bool,'F_0') = s(bool,sK6(X2,X3,bool))) ) | $spl0),
  inference(superposition,[],[f348,f785])).
fof(f2100,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,X15) = s(bool,sK9(X15,X16,bool)) | s(bool,X15) = s(bool,X17) | s(bool,'F_0') = s(bool,X17)) ) | $spl0),
  inference(superposition,[],[f354,f879])).
fof(f2148,plain,(
  ( ! [X57,X56,X55] : (s(bool,X55) = s(bool,sK9('const.marker.Cong_1'(s(bool,X55)),X56,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X56),s(bool,'F_0')))) | s(bool,X55) = s(bool,X57) | s(bool,'F_0') = s(bool,X57)) ) | $spl0),
  inference(forward_demodulation,[],[f2116,f271])).
fof(f2116,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X56),s(bool,'F_0')))) | s(bool,'const.marker.Cong_1'(s(bool,X55))) = s(bool,sK9('const.marker.Cong_1'(s(bool,X55)),X56,bool)) | s(bool,X55) = s(bool,X57) | s(bool,'F_0') = s(bool,X57)) ) | $spl0),
  inference(superposition,[],[f354,f1004])).
fof(f2312,plain,(
  ( ! [X41,X42] : (s(bool,X41) = s(bool,sK9('const.marker.Cong_1'(s(bool,X41)),X42,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X41)) ) | $spl0),
  inference(forward_demodulation,[],[f2271,f271])).
fof(f2271,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,'F_0')))) | s(bool,'const.marker.Cong_1'(s(bool,X41))) = s(bool,sK9('const.marker.Cong_1'(s(bool,X41)),X42,bool)) | s(bool,'T_0') = s(bool,X41)) ) | $spl0),
  inference(superposition,[],[f354,f1930])).
fof(f2515,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5)))) | s(bool,X6) = s(bool,sK6(X4,X5,bool)) | s(bool,'F_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f348,f879])).
fof(f5518,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X11),s('type.num.num',X12))) | s(bool,X13) = s(bool,'const.ind_type.FCONS_3'(s(bool,X10),s(fun('type.num.num',bool),X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))))) | s(bool,'F_0') = s(bool,X13)) ) | $spl0),
  inference(superposition,[],[f356,f879])).
fof(f5519,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) = s(bool,happ(s(fun('type.num.num',bool),X15),s('type.num.num',X16))) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X14),s(fun('type.num.num',bool),X15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16)))))) ) | $spl0),
  inference(superposition,[],[f356,f879])).
fof(f5522,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X29) = s(bool,happ(s(fun('type.num.num',bool),X27),s('type.num.num',X28))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X26),s(fun('type.num.num',bool),X27),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X29)))) ) | $spl0),
  inference(superposition,[],[f356,f1004])).
fof(f5524,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',X36))) | s(bool,X37) = s(bool,'const.ind_type.FCONS_3'(s(bool,X34),s(fun('type.num.num',bool),X35),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X37)))) ) | $spl0),
  inference(superposition,[],[f356,f1004])).
fof(f5556,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))) | s(bool,X105) = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104))))) | s(bool,'F_0') = s(bool,X105)) ) | $spl0),
  inference(superposition,[],[f1136,f356])).
fof(f5558,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X111),s('type.num.num',X112))) | s(bool,X113) = s(bool,'const.ind_type.FCONS_3'(s(bool,X110),s(fun('type.num.num',bool),X111),s('type.num.num','const.num.SUC_1'(s('type.num.num',X112))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X113)))) ) | $spl0),
  inference(superposition,[],[f1710,f356])).
fof(f5610,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',X9))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X7),s(fun('type.num.num',bool),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))) ) | $spl0),
  inference(forward_demodulation,[],[f5517,f447])).
fof(f5517,plain,(
  ( ! [X8,X7,X9] : (s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',X9))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X7),s(fun('type.num.num',bool),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))) ) | $spl0),
  inference(superposition,[],[f356,f785])).
fof(f5648,plain,(
  ( ! [X146,X144,X145] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X145),s('type.num.num',X146))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X144),s(fun('type.num.num',bool),X145),s('type.num.num','const.num.SUC_1'(s('type.num.num',X146)))))) ) | $spl0),
  inference(forward_demodulation,[],[f5568,f271])).
fof(f5568,plain,(
  ( ! [X146,X144,X145] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun('type.num.num',bool),X145),s('type.num.num',X146))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X144),s(fun('type.num.num',bool),X145),s('type.num.num','const.num.SUC_1'(s('type.num.num',X146)))))) ) | $spl0),
  inference(superposition,[],[f2178,f356])).
fof(f5706,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,'const.marker.Cong_1'(s(bool,X93))) != s(bool,happ(s(fun('type.num.num',bool),X91),s('type.num.num',X92))) | s(bool,X93) = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92)))))) ) | $spl0),
  inference(superposition,[],[f1798,f356])).
fof(f5816,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,X96) != s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95))) | s(bool,X96) = s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95))))) | s(bool,'F_0') = s(bool,X96)) ) | $spl0),
  inference(forward_demodulation,[],[f5747,f271])).
fof(f5747,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,X96) != s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95))))) | s(bool,X96) = s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95))))) | s(bool,'F_0') = s(bool,X96)) ) | $spl0),
  inference(superposition,[],[f1798,f356])).
fof(f6168,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'const.marker.Cong_1'(s(bool,X99))) != s(bool,happ(s(fun('type.num.num',bool),X97),s('type.num.num',X98))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.SUC_1'(s('type.num.num',X98))))) | s(bool,'T_0') = s(bool,X99)) ) | $spl0),
  inference(superposition,[],[f2047,f356])).
fof(f6295,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,X102) != s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))) | s(bool,'F_0') = s(bool,X102) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) ) | $spl0),
  inference(forward_demodulation,[],[f6210,f271])).
fof(f6210,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,X102) != s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,X102) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) ) | $spl0),
  inference(superposition,[],[f2047,f356])).
fof(f6430,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X99) != s(bool,happ(s(fun('type.num.num',bool),X97),s('type.num.num',X98))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.SUC_1'(s('type.num.num',X98))))) | s(bool,'T_0') = s(bool,X99)) ) | $spl0),
  inference(superposition,[],[f6213,f356])).
fof(f3796,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))) | s(bool,X8) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f360,f879])).
fof(f3797,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10))) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))) ) | $spl0),
  inference(superposition,[],[f360,f879])).
fof(f3800,plain,(
  ( ! [X19,X20,X18] : (s(bool,X20) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X19))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X20)))) ) | $spl0),
  inference(superposition,[],[f360,f1004])).
fof(f3802,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X25))) | s(bool,X26) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X26)))) ) | $spl0),
  inference(superposition,[],[f360,f1004])).
fof(f3828,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X60))) | s(bool,X61) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X59))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60))))) | s(bool,'F_0') = s(bool,X61)) ) | $spl0),
  inference(superposition,[],[f1136,f360])).
fof(f3846,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))) ) | $spl0),
  inference(forward_demodulation,[],[f3795,f447])).
fof(f3795,plain,(
  ( ! [X4,X5] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))) ) | $spl0),
  inference(superposition,[],[f360,f785])).
fof(f3879,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X83),s('type.num.num',X84))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84)))))) ) | $spl0),
  inference(forward_demodulation,[],[f3838,f271])).
fof(f3838,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X83),s('type.num.num',X84))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84)))))) ) | $spl0),
  inference(superposition,[],[f2178,f360])).
fof(f5423,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X56),s('type.num.num',X57))) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X57))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X58)))) ) | $spl0),
  inference(superposition,[],[f1710,f360])).
fof(f5690,plain,(
  ( ! [X54,X55,X53] : (s(bool,'const.marker.Cong_1'(s(bool,X55))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X53),s('type.num.num',X54))) | s(bool,X55) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X53))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X54))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X53))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X54)))))) ) | $spl0),
  inference(superposition,[],[f1798,f360])).
fof(f5811,plain,(
  ( ! [X57,X58,X56] : (s(bool,X58) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X56),s('type.num.num',X57))) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X57))))) | s(bool,'F_0') = s(bool,X58)) ) | $spl0),
  inference(forward_demodulation,[],[f5731,f271])).
fof(f5731,plain,(
  ( ! [X57,X58,X56] : (s(bool,X58) != s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X56),s('type.num.num',X57))))) | s(bool,X58) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X57))))) | s(bool,'F_0') = s(bool,X58)) ) | $spl0),
  inference(superposition,[],[f1798,f360])).
fof(f6152,plain,(
  ( ! [X61,X59,X60] : (s(bool,'const.marker.Cong_1'(s(bool,X61))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X60))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X59))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60))))) | s(bool,'T_0') = s(bool,X61)) ) | $spl0),
  inference(superposition,[],[f2047,f360])).
fof(f6287,plain,(
  ( ! [X64,X62,X63] : (s(bool,X64) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X62),s('type.num.num',X63))) | s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X63)))))) ) | $spl0),
  inference(forward_demodulation,[],[f6194,f271])).
fof(f6194,plain,(
  ( ! [X64,X62,X63] : (s(bool,X64) != s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X62),s('type.num.num',X63))))) | s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X63)))))) ) | $spl0),
  inference(superposition,[],[f2047,f360])).
fof(f6414,plain,(
  ( ! [X61,X59,X60] : (s(bool,X61) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X60))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X59))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60))))) | s(bool,'T_0') = s(bool,X61)) ) | $spl0),
  inference(superposition,[],[f6213,f360])).
fof(f4482,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',X7),s('type.list.list'(bool),X8))) | s(bool,X9) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X7))),s('type.list.list'(bool),X8))) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f346,f879])).
fof(f4483,plain,(
  ( ! [X12,X10,X11] : (s(bool,X12) = s(bool,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(bool),X11))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X10))),s('type.list.list'(bool),X11)))) ) | $spl0),
  inference(superposition,[],[f346,f879])).
fof(f4486,plain,(
  ( ! [X21,X19,X20] : (s(bool,X21) = s(bool,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(bool),X20))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X19))),s('type.list.list'(bool),X20))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X21)))) ) | $spl0),
  inference(superposition,[],[f346,f1004])).
fof(f4488,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',X25),s('type.list.list'(bool),X26))) | s(bool,X27) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X25))),s('type.list.list'(bool),X26))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X27)))) ) | $spl0),
  inference(superposition,[],[f346,f1004])).
fof(f4521,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') != s(bool,'const.list.EL_2'(s('type.num.num',X79),s('type.list.list'(bool),X80))) | s(bool,X81) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X79))),s('type.list.list'(bool),X80))) | s(bool,'F_0') = s(bool,X81)) ) | $spl0),
  inference(superposition,[],[f1136,f346])).
fof(f4562,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(bool),X6))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X5))),s('type.list.list'(bool),X6)))) ) | $spl0),
  inference(forward_demodulation,[],[f4481,f447])).
fof(f4481,plain,(
  ( ! [X6,X5] : (s(bool,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(bool),X6))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X5))),s('type.list.list'(bool),X6)))) ) | $spl0),
  inference(superposition,[],[f346,f785])).
fof(f4598,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,'const.list.EL_2'(s('type.num.num',X106),s('type.list.list'(bool),X107))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X106))),s('type.list.list'(bool),X107)))) ) | $spl0),
  inference(forward_demodulation,[],[f4532,f271])).
fof(f4532,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X106),s('type.list.list'(bool),X107))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X106))),s('type.list.list'(bool),X107)))) ) | $spl0),
  inference(superposition,[],[f2178,f346])).
fof(f5440,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') != s(bool,'const.list.EL_2'(s('type.num.num',X95),s('type.list.list'(bool),X96))) | s(bool,X97) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X95))),s('type.list.list'(bool),X96))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X97)))) ) | $spl0),
  inference(superposition,[],[f1710,f346])).
fof(f5708,plain,(
  ( ! [X97,X98,X96] : (s(bool,'const.marker.Cong_1'(s(bool,X98))) != s(bool,'const.list.EL_2'(s('type.num.num',X96),s('type.list.list'(bool),X97))) | s(bool,X98) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X96))),s('type.list.list'(bool),X97))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X96))),s('type.list.list'(bool),X97)))) ) | $spl0),
  inference(superposition,[],[f1798,f346])).
fof(f5818,plain,(
  ( ! [X101,X99,X100] : (s(bool,X101) != s(bool,'const.list.EL_2'(s('type.num.num',X99),s('type.list.list'(bool),X100))) | s(bool,X101) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X99))),s('type.list.list'(bool),X100))) | s(bool,'F_0') = s(bool,X101)) ) | $spl0),
  inference(forward_demodulation,[],[f5749,f271])).
fof(f5749,plain,(
  ( ! [X101,X99,X100] : (s(bool,X101) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X99),s('type.list.list'(bool),X100))))) | s(bool,X101) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X99))),s('type.list.list'(bool),X100))) | s(bool,'F_0') = s(bool,X101)) ) | $spl0),
  inference(superposition,[],[f1798,f346])).
fof(f6170,plain,(
  ( ! [X103,X102,X104] : (s(bool,'const.marker.Cong_1'(s(bool,X104))) != s(bool,'const.list.EL_2'(s('type.num.num',X102),s('type.list.list'(bool),X103))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X102))),s('type.list.list'(bool),X103))) | s(bool,'T_0') = s(bool,X104)) ) | $spl0),
  inference(superposition,[],[f2047,f346])).
fof(f6297,plain,(
  ( ! [X107,X105,X106] : (s(bool,X107) != s(bool,'const.list.EL_2'(s('type.num.num',X105),s('type.list.list'(bool),X106))) | s(bool,'F_0') = s(bool,X107) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X105))),s('type.list.list'(bool),X106)))) ) | $spl0),
  inference(forward_demodulation,[],[f6212,f271])).
fof(f6212,plain,(
  ( ! [X107,X105,X106] : (s(bool,X107) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X105),s('type.list.list'(bool),X106))))) | s(bool,'F_0') = s(bool,X107) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X105))),s('type.list.list'(bool),X106)))) ) | $spl0),
  inference(superposition,[],[f2047,f346])).
fof(f6432,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.list.EL_2'(s('type.num.num',X102),s('type.list.list'(bool),X103))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X102))),s('type.list.list'(bool),X103))) | s(bool,'T_0') = s(bool,X104)) ) | $spl0),
  inference(superposition,[],[f6213,f346])).
fof(f5707,plain,(
  ( ! [X94,X95] : (s(bool,'const.marker.Cong_1'(s(bool,X95))) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X94))) | s(bool,X95) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X94))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X94)))) ) | $spl0),
  inference(superposition,[],[f1798,f338])).
fof(f6169,plain,(
  ( ! [X101,X100] : (s(bool,'const.marker.Cong_1'(s(bool,X101))) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X100))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X100))) | s(bool,'T_0') = s(bool,X101)) ) | $spl0),
  inference(superposition,[],[f2047,f338])).
fof(f1456,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X2))) | s(bool,X3) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X2))) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f338,f879])).
fof(f1458,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X6))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X6))) | s(bool,'F_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f338,f879])).
fof(f1483,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X34))) | s(bool,X35) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X34))) | s(bool,'F_0') = s(bool,X35)) ) | $spl0),
  inference(superposition,[],[f1136,f338])).
fof(f1496,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X10))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X10)))) ) | $spl0),
  inference(forward_demodulation,[],[f1460,f447])).
fof(f1460,plain,(
  ( ! [X10] : (s(bool,'const.list.HD_1'(s('type.list.list'(bool),X10))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X10)))) ) | $spl0),
  inference(superposition,[],[f338,f785])).
fof(f1562,plain,(
  ( ! [X57,X56] : (s(bool,X57) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X56))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X56))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X57)))) ) | $spl0),
  inference(superposition,[],[f1004,f338])).
fof(f1589,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),X42))) | s(bool,X43) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X42))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X43)))) ) | $spl0),
  inference(superposition,[],[f1004,f338])).
fof(f2501,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X39)))) ) | $spl0),
  inference(forward_demodulation,[],[f2480,f271])).
fof(f2480,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X39))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X39)))) ) | $spl0),
  inference(superposition,[],[f2178,f338])).
fof(f5439,plain,(
  ( ! [X94,X93] : (s(bool,'T_0') != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X93))) | s(bool,X94) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X93))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X94)))) ) | $spl0),
  inference(superposition,[],[f1710,f338])).
fof(f5817,plain,(
  ( ! [X97,X98] : (s(bool,X98) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X97))) | s(bool,X98) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X97))) | s(bool,'F_0') = s(bool,X98)) ) | $spl0),
  inference(forward_demodulation,[],[f5748,f271])).
fof(f5748,plain,(
  ( ! [X97,X98] : (s(bool,X98) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X97))))) | s(bool,X98) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X97))) | s(bool,'F_0') = s(bool,X98)) ) | $spl0),
  inference(superposition,[],[f1798,f338])).
fof(f6296,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X103))) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X103)))) ) | $spl0),
  inference(forward_demodulation,[],[f6211,f271])).
fof(f6211,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.HD_1'(s('type.list.list'(bool),X103))))) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X103)))) ) | $spl0),
  inference(superposition,[],[f2047,f338])).
fof(f6431,plain,(
  ( ! [X101,X100] : (s(bool,X101) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X100))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X100))) | s(bool,'T_0') = s(bool,X101)) ) | $spl0),
  inference(superposition,[],[f6213,f338])).
fof(f6429,plain,(
  ( ! [X94,X95,X93] : (s(bool,X93) != s(bool,X95) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X95)) ) | $spl0),
  inference(superposition,[],[f6213,f344])).
fof(f6428,plain,(
  ( ! [X92,X90,X91] : (s(bool,X90) != s(bool,X92) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X90),s('type.list.list'(bool),X91))))) | s(bool,'T_0') = s(bool,X92)) ) | $spl0),
  inference(superposition,[],[f6213,f343])).
fof(f6330,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f360,f438,f6213])).
fof(f6213,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f2047,f271])).
fof(f6294,plain,(
  ( ! [X97,X98,X96] : (s(bool,X96) != s(bool,X98) | s(bool,'F_0') = s(bool,X98) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f6209,f271])).
fof(f6209,plain,(
  ( ! [X97,X98,X96] : (s(bool,X98) != s(bool,'const.marker.Cong_1'(s(bool,X96))) | s(bool,'F_0') = s(bool,X98) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f2047,f344])).
fof(f6293,plain,(
  ( ! [X94,X95,X93] : (s(bool,X93) != s(bool,X95) | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X93),s('type.list.list'(bool),X94)))))) ) | $spl0),
  inference(forward_demodulation,[],[f6208,f271])).
fof(f6208,plain,(
  ( ! [X94,X95,X93] : (s(bool,X95) != s(bool,'const.marker.Cong_1'(s(bool,X93))) | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X93),s('type.list.list'(bool),X94)))))) ) | $spl0),
  inference(superposition,[],[f2047,f343])).
fof(f6197,plain,(
  ( ! [X68,X69] : (s(bool,X69) != s(bool,'const.marker.Cong_1'(s(bool,X68))) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X68)))) ) | $spl0),
  inference(superposition,[],[f2047,f271])).
fof(f6167,plain,(
  ( ! [X94,X95,X93] : (s(bool,X93) != s(bool,'const.marker.Cong_1'(s(bool,X95))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X95)) ) | $spl0),
  inference(superposition,[],[f2047,f344])).
fof(f6166,plain,(
  ( ! [X92,X90,X91] : (s(bool,X90) != s(bool,'const.marker.Cong_1'(s(bool,X92))) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X90),s('type.list.list'(bool),X91))))) | s(bool,'T_0') = s(bool,X92)) ) | $spl0),
  inference(superposition,[],[f2047,f343])).
fof(f5963,plain,(
  ( ! [X99,X97,X95,X98,X96] : (s(bool,X95) = s(bool,'const.list.LIST_REL_3'(s(fun(X96,fun(bool,bool)),X97),s('type.list.list'(X96),'const.list.NIL_0'),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X95),s('type.list.list'(bool),X98))))) | s(bool,X95) = s(bool,X99) | s(bool,'F_0') = s(bool,X99)) ) | $spl0),
  inference(superposition,[],[f380,f879])).
fof(f5962,plain,(
  ( ! [X94,X92,X90,X93,X91] : (s(bool,X90) = s(bool,'const.list.LIST_REL_3'(s(fun(X91,fun(bool,bool)),X92),s('type.list.list'(X91),'const.list.NIL_0'),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X90),s('type.list.list'(bool),X93))))) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X94)))) ) | $spl0),
  inference(superposition,[],[f380,f1004])).
fof(f5815,plain,(
  ( ! [X92,X90,X91] : (s(bool,X90) != s(bool,X92) | s(bool,X92) = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X92)) ) | $spl0),
  inference(forward_demodulation,[],[f5746,f271])).
fof(f5746,plain,(
  ( ! [X92,X90,X91] : (s(bool,X92) != s(bool,'const.marker.Cong_1'(s(bool,X90))) | s(bool,X92) = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X92)) ) | $spl0),
  inference(superposition,[],[f1798,f344])).
fof(f5814,plain,(
  ( ! [X88,X87,X89] : (s(bool,X87) != s(bool,X89) | s(bool,X89) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X87),s('type.list.list'(bool),X88))))) | s(bool,'F_0') = s(bool,X89)) ) | $spl0),
  inference(forward_demodulation,[],[f5745,f271])).
fof(f5745,plain,(
  ( ! [X88,X87,X89] : (s(bool,X89) != s(bool,'const.marker.Cong_1'(s(bool,X87))) | s(bool,X89) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X87),s('type.list.list'(bool),X88))))) | s(bool,'F_0') = s(bool,X89)) ) | $spl0),
  inference(superposition,[],[f1798,f343])).
fof(f5705,plain,(
  ( ! [X88,X87,X89] : (s(bool,X87) != s(bool,'const.marker.Cong_1'(s(bool,X89))) | s(bool,X89) = s(bool,'const.ind_type.FCONS_3'(s(bool,X87),s(fun('type.num.num',bool),X88),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X87),s(fun('type.num.num',bool),X88),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f1798,f344])).
fof(f5704,plain,(
  ( ! [X85,X86,X84] : (s(bool,X84) != s(bool,'const.marker.Cong_1'(s(bool,X86))) | s(bool,X86) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X84),s('type.list.list'(bool),X85))))) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X84),s('type.list.list'(bool),X85)))))) ) | $spl0),
  inference(superposition,[],[f1798,f343])).
fof(f5693,plain,(
  ( ! [X59,X60] : (s(bool,X59) != s(bool,'const.marker.Cong_1'(s(bool,X60))) | s(bool,X60) = s(bool,'const.marker.Cong_1'(s(bool,X59))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X59)))) ) | $spl0),
  inference(superposition,[],[f1798,f271])).
fof(f5438,plain,(
  ( ! [X92,X90,X91] : (s(bool,'T_0') != s(bool,X90) | s(bool,X92) = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X92)))) ) | $spl0),
  inference(superposition,[],[f1710,f344])).
fof(f5437,plain,(
  ( ! [X88,X87,X89] : (s(bool,'T_0') != s(bool,X87) | s(bool,X89) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X87),s('type.list.list'(bool),X88))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X89)))) ) | $spl0),
  inference(superposition,[],[f1710,f343])).
fof(f5426,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,X63) = s(bool,'const.marker.Cong_1'(s(bool,X62))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X63)))) ) | $spl0),
  inference(superposition,[],[f1710,f271])).
fof(f5450,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f5397,f437])).
fof(f5397,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f1710,f879])).
fof(f5378,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,X5) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X4)))) ) | $spl0),
  inference(subsumption_resolution,[],[f5321,f437])).
fof(f5321,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,X5) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X4)))) ) | $spl0),
  inference(superposition,[],[f5187,f1004])).
fof(f5377,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f5319,f437])).
fof(f5319,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f5187,f879])).
fof(f5187,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f5002,f271])).
fof(f5009,plain,(
  ( ! [X28,X29] : (s(bool,X29) = s(bool,'const.marker.Cong_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f4738,f426])).
fof(f4738,plain,(
  ( ! [X28,X29] : (p(s(bool,'F_0')) | s(bool,X29) = s(bool,'const.marker.Cong_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f546,f4389])).
fof(f5002,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X17))) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f4730,f426])).
fof(f4730,plain,(
  ( ! [X17] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X17))) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f281,f4389])).
fof(f4389,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(condensation,[],[f4388])).
fof(f4388,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X11)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f4183])).
fof(f4183,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X11)))) ) | $spl0),
  inference(superposition,[],[f4160,f1004])).
fof(f4160,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(forward_demodulation,[],[f4159,f447])).
fof(f4159,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))))) ) | $spl0),
  inference(condensation,[],[f4158])).
fof(f4158,plain,(
  ( ! [X41,X42] : (s(bool,X42) = s(bool,'const.marker.Cong_1'(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X41)) ) | $spl0),
  inference(subsumption_resolution,[],[f4120,f276])).
fof(f4120,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') != s(bool,X41) | s(bool,X42) = s(bool,'const.marker.Cong_1'(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X41)) ) | $spl0),
  inference(superposition,[],[f1429,f785])).
fof(f4155,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X8) = s(bool,'const.marker.Cong_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,X8) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X6)))) ) | $spl0),
  inference(subsumption_resolution,[],[f4103,f437])).
fof(f4103,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X8) = s(bool,'const.marker.Cong_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X6)))) ) | $spl0),
  inference(superposition,[],[f1429,f1004])).
fof(f4154,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f4101,f437])).
fof(f4101,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f1429,f879])).
fof(f3337,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X4)))) ) | $spl0),
  inference(subsumption_resolution,[],[f3289,f437])).
fof(f3289,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X4)))) ) | $spl0),
  inference(superposition,[],[f3210,f1004])).
fof(f3336,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f3287,f437])).
fof(f3287,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f3210,f879])).
fof(f3210,plain,(
  ( ! [X61] : (s(bool,'T_0') != s(bool,X61) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X61)))) ) | $spl0),
  inference(forward_demodulation,[],[f3180,f272])).
fof(f3180,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') != s(bool,X61) | s(bool,'const.marker.Cong_1'(s(bool,X61))) = s(bool,'const.bool.BOUNDED_1'(s(bool,X62)))) ) | $spl0),
  inference(superposition,[],[f3103,f271])).
fof(f3202,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,'const.bool.BOUNDED_1'(s(bool,X8))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X6)))) ) | $spl0),
  inference(subsumption_resolution,[],[f3154,f437])).
fof(f3154,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,'const.bool.BOUNDED_1'(s(bool,X8))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X6)))) ) | $spl0),
  inference(superposition,[],[f3103,f1004])).
fof(f3201,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.bool.BOUNDED_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f3152,f437])).
fof(f3152,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.bool.BOUNDED_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f3103,f879])).
fof(f3103,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | s(bool,X10) = s(bool,'const.bool.BOUNDED_1'(s(bool,X11)))) ) | $spl0),
  inference(superposition,[],[f2944,f271])).
fof(f2995,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,'const.marker.Cong_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X10))) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f2074,f276])).
fof(f2993,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.marker.Cong_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f2074,f879])).
fof(f2992,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f2074,f1004])).
fof(f2951,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) = s(bool,'const.marker.Cong_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,X37) = s(bool,'const.bool.BOUNDED_1'(s(bool,X39)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2755,f426])).
fof(f2755,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'F_0')) | s(bool,X38) = s(bool,'const.marker.Cong_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,X37) = s(bool,'const.bool.BOUNDED_1'(s(bool,X39)))) ) | $spl0),
  inference(superposition,[],[f546,f1945])).
fof(f2944,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X18))) | s(bool,X18) = s(bool,'const.bool.BOUNDED_1'(s(bool,X19)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2747,f426])).
fof(f2747,plain,(
  ( ! [X19,X18] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X18))) | s(bool,X18) = s(bool,'const.bool.BOUNDED_1'(s(bool,X19)))) ) | $spl0),
  inference(superposition,[],[f281,f1945])).
fof(f2500,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X37),s(fun('type.num.num',bool),X38),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f2479,f271])).
fof(f2479,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X37),s(fun('type.num.num',bool),X38),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f2178,f344])).
fof(f2499,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X35),s('type.list.list'(bool),X36)))))) ) | $spl0),
  inference(forward_demodulation,[],[f2478,f271])).
fof(f2478,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X35))) | s(bool,'T_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X35),s('type.list.list'(bool),X36)))))) ) | $spl0),
  inference(superposition,[],[f2178,f343])).
fof(f2374,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X10))) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f1992,f276])).
fof(f2371,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f1992,f1004])).
fof(f2360,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X36) = s(bool,'const.list.LIST_REL_3'(s(fun(X33,fun(X34,bool)),X35),s('type.list.list'(X33),'const.list.NIL_0'),s('type.list.list'(X34),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X36)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2327,f437])).
fof(f2327,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X36) = s(bool,'const.list.LIST_REL_3'(s(fun(X33,fun(X34,bool)),X35),s('type.list.list'(X33),'const.list.NIL_0'),s('type.list.list'(X34),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X36)))) ) | $spl0),
  inference(superposition,[],[f368,f1004])).
fof(f2359,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X12) = s(bool,'const.list.LIST_REL_3'(s(fun(X9,fun(X10,bool)),X11),s('type.list.list'(X9),'const.list.NIL_0'),s('type.list.list'(X10),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X12)) ) | $spl0),
  inference(subsumption_resolution,[],[f2321,f437])).
fof(f2321,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X12) = s(bool,'const.list.LIST_REL_3'(s(fun(X9,fun(X10,bool)),X11),s('type.list.list'(X9),'const.list.NIL_0'),s('type.list.list'(X10),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X12)) ) | $spl0),
  inference(superposition,[],[f368,f879])).
fof(f2178,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f2014,f276])).
fof(f2175,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X1)))) ) | $spl0),
  inference(superposition,[],[f2014,f1004])).
fof(f2156,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5))) | s(X3,X5) = s(X3,sK9(X5,X4,X3))) ) | $spl0),
  inference(subsumption_resolution,[],[f2155,f264])).
fof(f2155,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(X3,X5) = s(X3,sK9(X5,X4,X3)) | s(bool,'F_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) ) | $spl0),
  inference(forward_demodulation,[],[f2125,f447])).
fof(f2125,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) | s(X3,X5) = s(X3,sK9(X5,X4,X3)) | s(bool,'F_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) ) | $spl0),
  inference(superposition,[],[f354,f785])).
fof(f2047,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,'const.marker.Cong_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f1810,f276])).
fof(f2045,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.marker.Cong_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f1810,f879])).
fof(f2044,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f1810,f1004])).
fof(f2074,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X0)))) ) | $spl0),
  inference(condensation,[],[f2073])).
fof(f2073,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X11)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f2025])).
fof(f2025,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X11)))) ) | $spl0),
  inference(superposition,[],[f1810,f1004])).
fof(f2014,plain,(
  s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f437,f1810])).
fof(f1798,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.marker.Cong_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(equality_factoring,[],[f1004])).
fof(f1810,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1792])).
fof(f1792,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(equality_factoring,[],[f1004])).
fof(f2013,plain,(
  ( ! [X30,X31,X32] : (s(bool,X31) = s(bool,'const.marker.Cong_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f1781,f426])).
fof(f1781,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,'F_0')) | s(bool,X31) = s(bool,'const.marker.Cong_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f546,f1004])).
fof(f2008,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X11))) | s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,X12)) ) | $spl0),
  inference(subsumption_resolution,[],[f1773,f426])).
fof(f1773,plain,(
  ( ! [X12,X11] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X11))) | s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,X12)) ) | $spl0),
  inference(superposition,[],[f281,f1004])).
fof(f2007,plain,(
  ( ! [X66,X64,X67,X65] : (s(X64,X66) = s(X64,sK6(X66,X65,X64)) | s(bool,X67) = s(bool,'|dollar|exists_1'(s(fun(X64,bool),X65))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X67)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1757,f426])).
fof(f1757,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,'F_0')) | s(X64,X66) = s(X64,sK6(X66,X65,X64)) | s(bool,X67) = s(bool,'|dollar|exists_1'(s(fun(X64,bool),X65))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X67)))) ) | $spl0),
  inference(superposition,[],[f347,f1004])).
fof(f1710,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X6)))) ) | $spl0),
  inference(superposition,[],[f437,f1004])).
fof(f1992,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X0)))) ) | $spl0),
  inference(condensation,[],[f1991])).
fof(f1991,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,X84) | s(bool,'F_0') = s(bool,X83) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X84)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1699,f621])).
fof(f1699,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,X84) | s(bool,'F_0') = s(bool,X83) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X83))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X84)))) ) | $spl0),
  inference(superposition,[],[f621,f1004])).
fof(f1985,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) != s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X50)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1685,f437])).
fof(f1685,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) != s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X50)))) ) | $spl0),
  inference(superposition,[],[f1136,f1004])).
fof(f1588,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.ind_type.FCONS_3'(s(bool,X39),s(fun('type.num.num',bool),X40),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X41)))) ) | $spl0),
  inference(superposition,[],[f1004,f344])).
fof(f1587,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') = s(bool,X36) | s(bool,X38) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X36),s('type.list.list'(bool),X37))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X38)))) ) | $spl0),
  inference(superposition,[],[f1004,f343])).
fof(f1945,plain,(
  ( ! [X35,X34] : (s(bool,X35) = s(bool,'const.bool.BOUNDED_1'(s(bool,X34))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X35)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1586,f437])).
fof(f1586,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X35) = s(bool,'const.bool.BOUNDED_1'(s(bool,X34))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X35)))) ) | $spl0),
  inference(superposition,[],[f1004,f272])).
fof(f1944,plain,(
  ( ! [X30,X31] : (s(bool,X31) = s(bool,'const.marker.Cong_1'(s(bool,X30))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X30)) ) | $spl0),
  inference(subsumption_resolution,[],[f1584,f437])).
fof(f1584,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X31) = s(bool,'const.marker.Cong_1'(s(bool,X30))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X30)) ) | $spl0),
  inference(superposition,[],[f1004,f489])).
fof(f1942,plain,(
  ( ! [X26,X27] : (s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X27)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1582,f437])).
fof(f1582,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X27) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X27)))) ) | $spl0),
  inference(superposition,[],[f1004,f479])).
fof(f1566,plain,(
  ( ! [X14,X12,X13] : (s(bool,X12) = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X14))) | s(bool,'T_0') = s(bool,X12)) ) | $spl0),
  inference(superposition,[],[f1004,f276])).
fof(f1564,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X6))) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f1004,f879])).
fof(f1563,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X2))) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X3)))) ) | $spl0),
  inference(superposition,[],[f1004,f1004])).
fof(f1561,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) = s(bool,X55) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X53),s(fun('type.num.num',bool),X54),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X55)))) ) | $spl0),
  inference(superposition,[],[f1004,f344])).
fof(f1560,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) = s(bool,X52) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X50),s('type.list.list'(bool),X51))))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X52)))) ) | $spl0),
  inference(superposition,[],[f1004,f343])).
fof(f1923,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X1)))) ) | $spl0),
  inference(condensation,[],[f1922])).
fof(f1922,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X22))) | s(bool,X22) = s(bool,X24) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X24)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1521])).
fof(f1521,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X22))) | s(bool,X22) = s(bool,X24) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X24)))) ) | $spl0),
  inference(superposition,[],[f1004,f1004])).
fof(f1432,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X41) | s(bool,X43) = s(bool,'const.ind_type.FCONS_3'(s(bool,X41),s(fun('type.num.num',bool),X42),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X43)) ) | $spl0),
  inference(superposition,[],[f1136,f344])).
fof(f1431,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,X38) | s(bool,X40) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X38),s('type.list.list'(bool),X39))))) | s(bool,'F_0') = s(bool,X40)) ) | $spl0),
  inference(superposition,[],[f1136,f343])).
fof(f1429,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X34) | s(bool,X35) = s(bool,'const.marker.Cong_1'(s(bool,X34))) | s(bool,'F_0') = s(bool,X35)) ) | $spl0),
  inference(superposition,[],[f1136,f271])).
fof(f1450,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f1412,f437])).
fof(f1412,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f1136,f879])).
fof(f1397,plain,(
  ( ! [X47,X50,X48,X49] : (s(X47,X49) = s(X47,sK6(X49,X48,X47)) | s(bool,X50) = s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48))) | s(bool,'F_0') = s(bool,X50)) ) | $spl0),
  inference(subsumption_resolution,[],[f1174,f426])).
fof(f1174,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,'F_0')) | s(X47,X49) = s(X47,sK6(X49,X48,X47)) | s(bool,X50) = s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48))) | s(bool,'F_0') = s(bool,X50)) ) | $spl0),
  inference(superposition,[],[f347,f879])).
fof(f1396,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'F_0') = s(bool,X15)) ) | $spl0),
  inference(subsumption_resolution,[],[f1159,f426])).
fof(f1159,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'F_0')) | p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'F_0') = s(bool,X15)) ) | $spl0),
  inference(superposition,[],[f361,f879])).
fof(f1136,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f437,f879])).
fof(f1046,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') = s(bool,X30) | s(bool,X32) = s(bool,'const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f879,f344])).
fof(f1045,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,X29) = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X27),s('type.list.list'(bool),X28))))) | s(bool,'F_0') = s(bool,X29)) ) | $spl0),
  inference(superposition,[],[f879,f343])).
fof(f1029,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f879,f276])).
fof(f1027,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f879,f879])).
fof(f1007,plain,(
  ( ! [X39,X41,X40] : (s(bool,X39) = s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X39),s(fun('type.num.num',bool),X40),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f879,f344])).
fof(f1006,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) = s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X36),s('type.list.list'(bool),X37)))))) ) | $spl0),
  inference(superposition,[],[f879,f343])).
fof(f878,plain,(
  ( ! [X12,X11] : (s(bool,X11) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X11),s(fun('type.num.num',bool),X12),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f785,f344])).
fof(f877,plain,(
  ( ! [X10,X9] : (s(bool,X9) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,'const.list.HD_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X9),s('type.list.list'(bool),X10)))))) ) | $spl0),
  inference(superposition,[],[f785,f343])).
fof(f875,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X7)))) ) | $spl0),
  inference(superposition,[],[f785,f271])).
fof(f34588,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547,f34585,f34587])).
fof(f34586,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547,f34585])).
fof(f34573,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34572,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547])).
fof(f34572,plain,(
  ( ! [X121,X122] : (s(bool,X122) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X121)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f34041,f17030])).
fof(f34041,plain,(
  ( ! [X121,X122] : (s(bool,'T_0') != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X121))) | s(bool,X122) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X121)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14562])).
fof(f14562,plain,(
  ( ! [X121,X122] : (s(bool,'const.marker.Cong_1'(s(bool,X122))) != s(bool,'const.list.HD_1'(s('type.list.list'(bool),X121))) | s(bool,X122) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.HD_0'),s('type.list.list'(bool),X121)))) ) | $spl0),
  inference(superposition,[],[f13848,f338])).
fof(f34566,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34565,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547])).
fof(f34565,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,X133) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X130))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X131),s('type.list.list'(bool),X132)))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f34024,f17030])).
fof(f34024,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,'T_0') != s(bool,X133) | s(bool,X133) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X130))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X131),s('type.list.list'(bool),X132)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14701])).
fof(f14701,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,X133) != s(bool,'const.list.EL_2'(s('type.num.num',X130),s('type.list.list'(bool),X132))) | s(bool,X133) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X130))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X131),s('type.list.list'(bool),X132)))))) ) | $spl0),
  inference(forward_demodulation,[],[f14620,f271])).
fof(f14620,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,X133) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X130),s('type.list.list'(bool),X132))))) | s(bool,X133) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X130))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X131),s('type.list.list'(bool),X132)))))) ) | $spl0),
  inference(superposition,[],[f13848,f357])).
fof(f34564,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34563,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547])).
fof(f34563,plain,(
  ( ! [X127,X128,X129] : (s(bool,X129) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X127))),s('type.list.list'(bool),X128)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f34023,f17030])).
fof(f34023,plain,(
  ( ! [X127,X128,X129] : (s(bool,'T_0') != s(bool,X129) | s(bool,X129) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X127))),s('type.list.list'(bool),X128)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14700])).
fof(f14700,plain,(
  ( ! [X127,X128,X129] : (s(bool,X129) != s(bool,'const.list.EL_2'(s('type.num.num',X127),s('type.list.list'(bool),X128))) | s(bool,X129) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X127))),s('type.list.list'(bool),X128)))) ) | $spl0),
  inference(forward_demodulation,[],[f14619,f271])).
fof(f14619,plain,(
  ( ! [X127,X128,X129] : (s(bool,X129) != s(bool,'const.marker.Cong_1'(s(bool,'const.list.EL_2'(s('type.num.num',X127),s('type.list.list'(bool),X128))))) | s(bool,X129) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X127))),s('type.list.list'(bool),X128)))) ) | $spl0),
  inference(superposition,[],[f13848,f346])).
fof(f34562,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34561,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547])).
fof(f34561,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,X129) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X126))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X127),s('type.list.list'(bool),X128)))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f34022,f17030])).
fof(f34022,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X129))) | s(bool,X129) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X126))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X127),s('type.list.list'(bool),X128)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14564])).
fof(f14564,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'const.marker.Cong_1'(s(bool,X129))) != s(bool,'const.list.EL_2'(s('type.num.num',X126),s('type.list.list'(bool),X128))) | s(bool,X129) = s(bool,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X126))),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X127),s('type.list.list'(bool),X128)))))) ) | $spl0),
  inference(superposition,[],[f13848,f357])).
fof(f34560,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34559,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547])).
fof(f34559,plain,(
  ( ! [X125,X123,X124] : (s(bool,X125) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X123))),s('type.list.list'(bool),X124)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f34021,f17030])).
fof(f34021,plain,(
  ( ! [X125,X123,X124] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X125))) | s(bool,X125) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X123))),s('type.list.list'(bool),X124)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14563])).
fof(f14563,plain,(
  ( ! [X125,X123,X124] : (s(bool,'const.marker.Cong_1'(s(bool,X125))) != s(bool,'const.list.EL_2'(s('type.num.num',X123),s('type.list.list'(bool),X124))) | s(bool,X125) = s(bool,happ(s(fun('type.list.list'(bool),bool),'const.list.EL_1'(s('type.num.num',X123))),s('type.list.list'(bool),X124)))) ) | $spl0),
  inference(superposition,[],[f13848,f346])).
fof(f34548,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533,f34547])).
fof(f34546,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34094,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533])).
fof(f34094,plain,(
  ( ! [X461,X460,X462] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X460))),s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X461),s('type.list.list'('type.num.num'),X462)))))) ) | ($spl0 | $spl36)),
  inference(trivial_inequality_removal,[],[f34009])).
fof(f34009,plain,(
  ( ! [X461,X460,X462] : (s(bool,'T_0') != s(bool,'T_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X460))),s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X461),s('type.list.list'('type.num.num'),X462)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f11439])).
fof(f34545,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34095,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533])).
fof(f34095,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X8),s(fun('type.num.num','type.num.num'),X9),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl36)),
  inference(trivial_inequality_removal,[],[f34006])).
fof(f34006,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'T_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X8),s(fun('type.num.num','type.num.num'),X9),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f8186])).
fof(f34544,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34096,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533])).
fof(f34096,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.HD_1'(s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X6),s('type.list.list'('type.num.num'),X7)))))) ) | ($spl0 | $spl36)),
  inference(trivial_inequality_removal,[],[f34005])).
fof(f34005,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'T_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.HD_1'(s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num',X6),s('type.list.list'('type.num.num'),X7)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f8185])).
fof(f34541,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34540,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533])).
fof(f34540,plain,(
  ( ! [X227,X228,X226] : (s(bool,X228) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X226))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X227)))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f34001,f17030])).
fof(f34001,plain,(
  ( ! [X227,X228,X226] : (s(bool,'T_0') != s(bool,X228) | s(bool,X228) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X226))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X227)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f16226])).
fof(f16226,plain,(
  ( ! [X227,X228,X226] : (s(bool,X228) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X226),s('type.num.num',X227))) | s(bool,X228) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X226))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X227)))))) ) | $spl0),
  inference(forward_demodulation,[],[f16116,f271])).
fof(f16116,plain,(
  ( ! [X227,X228,X226] : (s(bool,X228) != s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X226),s('type.num.num',X227))))) | s(bool,X228) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X226))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X227)))))) ) | $spl0),
  inference(superposition,[],[f13848,f360])).
fof(f34539,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34538,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533])).
fof(f34538,plain,(
  ( ! [X225,X223,X224] : (s(bool,X225) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X223))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X224)))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f34000,f17030])).
fof(f34000,plain,(
  ( ! [X225,X223,X224] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X225))) | s(bool,X225) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X223))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X224)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f16115])).
fof(f16115,plain,(
  ( ! [X225,X223,X224] : (s(bool,'const.marker.Cong_1'(s(bool,X225))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X223),s('type.num.num',X224))) | s(bool,X225) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X223))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X224)))))) ) | $spl0),
  inference(superposition,[],[f13848,f360])).
fof(f34534,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520,f34533])).
fof(f34530,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34529,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520])).
fof(f34529,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X140) = s(bool,'const.ind_type.FCONS_3'(s(bool,X137),s(fun('type.num.num',bool),X138),s('type.num.num','const.num.SUC_1'(s('type.num.num',X139)))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33991,f17030])).
fof(f33991,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,'T_0') != s(bool,X140) | s(bool,X140) = s(bool,'const.ind_type.FCONS_3'(s(bool,X137),s(fun('type.num.num',bool),X138),s('type.num.num','const.num.SUC_1'(s('type.num.num',X139)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14703])).
fof(f14703,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X140) != s(bool,happ(s(fun('type.num.num',bool),X138),s('type.num.num',X139))) | s(bool,X140) = s(bool,'const.ind_type.FCONS_3'(s(bool,X137),s(fun('type.num.num',bool),X138),s('type.num.num','const.num.SUC_1'(s('type.num.num',X139)))))) ) | $spl0),
  inference(forward_demodulation,[],[f14622,f271])).
fof(f14622,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X140) != s(bool,'const.marker.Cong_1'(s(bool,happ(s(fun('type.num.num',bool),X138),s('type.num.num',X139))))) | s(bool,X140) = s(bool,'const.ind_type.FCONS_3'(s(bool,X137),s(fun('type.num.num',bool),X138),s('type.num.num','const.num.SUC_1'(s('type.num.num',X139)))))) ) | $spl0),
  inference(superposition,[],[f13848,f356])).
fof(f34528,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34527,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520])).
fof(f34527,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,X136) = s(bool,'const.ind_type.FCONS_3'(s(bool,X133),s(fun('type.num.num',bool),X134),s('type.num.num','const.num.SUC_1'(s('type.num.num',X135)))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33990,f17030])).
fof(f33990,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X136))) | s(bool,X136) = s(bool,'const.ind_type.FCONS_3'(s(bool,X133),s(fun('type.num.num',bool),X134),s('type.num.num','const.num.SUC_1'(s('type.num.num',X135)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14566])).
fof(f14566,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'const.marker.Cong_1'(s(bool,X136))) != s(bool,happ(s(fun('type.num.num',bool),X134),s('type.num.num',X135))) | s(bool,X136) = s(bool,'const.ind_type.FCONS_3'(s(bool,X133),s(fun('type.num.num',bool),X134),s('type.num.num','const.num.SUC_1'(s('type.num.num',X135)))))) ) | $spl0),
  inference(superposition,[],[f13848,f356])).
fof(f34521,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518,f34520])).
fof(f34519,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516,f34518])).
fof(f34517,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512,f34516])).
fof(f34515,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34514,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512])).
fof(f34514,plain,(
  ( ! [X158] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)),X158,'type.num.num'))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33981,f264])).
fof(f33981,plain,(
  ( ! [X158] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)),X158,'type.num.num'))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f15110])).
fof(f15110,plain,(
  ( ! [X158] : (~p(s(bool,happ(s(fun('type.num.num',bool),X158),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)),X158,'type.num.num'))) ) | $spl0),
  inference(forward_demodulation,[],[f15023,f13908])).
fof(f13908,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) | $spl0),
  inference(forward_demodulation,[],[f13907,f273])).
fof(f13907,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) | $spl0),
  inference(subsumption_resolution,[],[f13869,f264])).
fof(f13869,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) | $spl0),
  inference(superposition,[],[f381,f447])).
fof(f15023,plain,(
  ( ! [X158] : (~p(s(bool,happ(s(fun('type.num.num',bool),X158),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) = s('type.num.num',sK9('const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)),X158,'type.num.num'))) ) | $spl0),
  inference(superposition,[],[f354,f13908])).
fof(f34513,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508,f34512])).
fof(f34509,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504,f34508])).
fof(f34505,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500,f34504])).
fof(f34501,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494,f34500])).
fof(f34495,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488,f34494])).
fof(f34489,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476,f34488])).
fof(f34477,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474,f34476])).
fof(f34475,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472,f34474])).
fof(f34473,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470,f34472])).
fof(f34471,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464,f34470])).
fof(f34465,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462,f34464])).
fof(f34463,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460,f34462])).
fof(f34461,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458,f34460])).
fof(f34459,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454,f34458])).
fof(f34455,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448,f34454])).
fof(f34449,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446,f34448])).
fof(f34447,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444,f34446])).
fof(f34445,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442,f34444])).
fof(f34443,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440,f34442])).
fof(f34441,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438,f34440])).
fof(f34439,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436,f34438])).
fof(f34437,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434,f34436])).
fof(f34435,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432,f34434])).
fof(f34433,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430,f34432])).
fof(f34431,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428,f34430])).
fof(f34429,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426,f34428])).
fof(f34427,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424,f34426])).
fof(f34425,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422,f34424])).
fof(f34423,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420,f34422])).
fof(f34421,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418,f34420])).
fof(f34419,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416,f34418])).
fof(f34417,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414,f34416])).
fof(f34415,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412,f34414])).
fof(f34413,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410,f34412])).
fof(f34411,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408,f34410])).
fof(f34409,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402,f34408])).
fof(f34403,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127,f34402])).
fof(f34387,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34119,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34119,plain,(
  ( ! [X761,X763,X764,X760,X762] : (s(fun(X760,bool),X761) = s(fun(X760,bool),X762) | s(bool,X763) = s(bool,g(s(X760,sK19(X761,X762,bool,X760)))) | s(bool,X763) = s(bool,'const.bool.BOUNDED_1'(s(bool,X764)))) ) | ($spl0 | $spl36)),
  inference(trivial_inequality_removal,[],[f33883])).
fof(f33883,plain,(
  ( ! [X761,X763,X764,X760,X762] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X760,bool),X761) = s(fun(X760,bool),X762) | s(bool,X763) = s(bool,g(s(X760,sK19(X761,X762,bool,X760)))) | s(bool,X763) = s(bool,'const.bool.BOUNDED_1'(s(bool,X764)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f29250])).
fof(f29250,plain,(
  ( ! [X761,X763,X764,X760,X762] : (s(bool,'T_0') != s(bool,f(s(X760,sK19(X761,X762,bool,X760)))) | s(fun(X760,bool),X761) = s(fun(X760,bool),X762) | s(bool,X763) = s(bool,g(s(X760,sK19(X761,X762,bool,X760)))) | s(bool,X763) = s(bool,'const.bool.BOUNDED_1'(s(bool,X764)))) ) | $spl0),
  inference(superposition,[],[f379,f13451])).
fof(f34386,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34385,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34385,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34384])).
fof(f34384,plain,(
  ( ! [X281,X283] : (s(bool,X281) = s(bool,X283) | s(bool,X283) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33877,f13817])).
fof(f33877,plain,(
  ( ! [X281,X283] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X281))) | s(bool,X281) = s(bool,X283) | s(bool,X283) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f33049])).
fof(f33049,plain,(
  ( ! [X281,X283] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X281))) | s(bool,X281) = s(bool,X283) | s(bool,X283) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f31602,f438])).
fof(f31602,plain,(
  ( ! [X282,X281,X283] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X281))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X282),s('type.num.num','const.num.|48|_0'))) | s(bool,X281) = s(bool,X283) | s(bool,X283) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f14540,f13452])).
fof(f14540,plain,(
  ( ! [X74,X73] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X74))) | s(bool,X74) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f13848,f462])).
fof(f34373,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34372,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34372,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34371])).
fof(f34371,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34370])).
fof(f34370,plain,(
  ( ! [X74,X75,X73] : (s(bool,X74) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X73) = s(bool,X74) | s(bool,X73) = s(bool,X75)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33850,f31708])).
fof(f31708,plain,(
  ( ! [X251,X252,X250] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X251))) | s(bool,X250) = s(bool,X251) | s(bool,X250) = s(bool,X252) | s(bool,X252) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f13848,f13452])).
fof(f33850,plain,(
  ( ! [X74,X75,X73] : (s(bool,X74) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X73) = s(bool,X74) | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X75)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f27163])).
fof(f27163,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X75)))) ) | $spl0),
  inference(subsumption_resolution,[],[f25644,f437])).
fof(f25644,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X75)))) ) | $spl0),
  inference(superposition,[],[f9876,f9373])).
fof(f34358,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34357,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34357,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f34356,f17030])).
fof(f34356,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | ($spl0 | $spl36)),
  inference(duplicate_literal_removal,[],[f34355])).
fof(f34355,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34354])).
fof(f34354,plain,(
  ( ! [X212,X210,X211] : (s(bool,X210) = s(bool,X211) | s(bool,X212) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X210) = s(bool,X212)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33840,f17030])).
fof(f33840,plain,(
  ( ! [X212,X210,X211] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X211))) | s(bool,X210) = s(bool,X211) | s(bool,X212) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X210) = s(bool,X212)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f26629])).
fof(f26629,plain,(
  ( ! [X212,X210,X211] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X211))) | s(bool,X210) = s(bool,X211) | s(bool,X212) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X210) = s(bool,X212)) ) | $spl0),
  inference(superposition,[],[f13848,f9876])).
fof(f34352,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34351,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34351,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X28) = s(bool,'const.marker.Cong_1'(s(bool,X27))) | s(bool,X27) = s(bool,X29)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33836,f31711])).
fof(f31711,plain,(
  ( ! [X261,X259,X260] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X260))) | s(bool,X260) = s(bool,'const.marker.Cong_1'(s(bool,X259))) | s(bool,X259) = s(bool,X261) | s(bool,X261) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f14546,f13452])).
fof(f14546,plain,(
  ( ! [X85,X84] : (s(bool,X84) != s(bool,'const.marker.Cong_1'(s(bool,X85))) | s(bool,X85) = s(bool,'const.marker.Cong_1'(s(bool,X84)))) ) | $spl0),
  inference(superposition,[],[f13848,f271])).
fof(f33836,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X28) = s(bool,'const.marker.Cong_1'(s(bool,X27))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X28))) | s(bool,X27) = s(bool,X29)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f26470])).
fof(f26470,plain,(
  ( ! [X28,X29,X27] : (s(bool,X28) = s(bool,'const.marker.Cong_1'(s(bool,X27))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X28))) | s(bool,X29) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X27) = s(bool,X29)) ) | $spl0),
  inference(superposition,[],[f8912,f9876])).
fof(f34350,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34349,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34349,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X25) = s(bool,'const.marker.Cong_1'(s(bool,X24))) | s(bool,X24) = s(bool,X26)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33835,f31717])).
fof(f31717,plain,(
  ( ! [X277,X279,X278] : (s(bool,'T_0') != s(bool,X278) | s(bool,X278) = s(bool,'const.marker.Cong_1'(s(bool,X277))) | s(bool,X277) = s(bool,X279) | s(bool,X279) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f14826,f13452])).
fof(f14826,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,X1)))) ) | $spl0),
  inference(superposition,[],[f14546,f271])).
fof(f33835,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X25) = s(bool,'const.marker.Cong_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,X24) = s(bool,X26)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f26469])).
fof(f26469,plain,(
  ( ! [X26,X24,X25] : (s(bool,X25) = s(bool,'const.marker.Cong_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X24) = s(bool,X26)) ) | $spl0),
  inference(superposition,[],[f8650,f9876])).
fof(f34348,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34347,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34347,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34346])).
fof(f34346,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X22) = s(bool,X23)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33834,f13817])).
fof(f33834,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X22))) | s(bool,X22) = s(bool,X23)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f26468])).
fof(f26468,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X22))) | s(bool,X23) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X22) = s(bool,X23)) ) | $spl0),
  inference(superposition,[],[f2014,f9876])).
fof(f34323,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34322,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34322,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.BOUNDED_1'(s(bool,X1)))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34321])).
fof(f34321,plain,(
  ( ! [X273,X270,X272] : (s(bool,X270) = s(bool,X272) | s(bool,X270) = s(bool,'const.bool.BOUNDED_1'(s(bool,X273)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33805,f13811])).
fof(f33805,plain,(
  ( ! [X273,X270,X272] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X272))) | s(bool,X270) = s(bool,X272) | s(bool,X270) = s(bool,'const.bool.BOUNDED_1'(s(bool,X273)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f22521])).
fof(f22521,plain,(
  ( ! [X273,X270,X272] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X272))) | s(bool,X270) = s(bool,X272) | s(bool,X270) = s(bool,'const.bool.BOUNDED_1'(s(bool,X273)))) ) | $spl0),
  inference(forward_demodulation,[],[f21826,f462])).
fof(f21826,plain,(
  ( ! [X271,X273,X270,X272] : (s(bool,'const.marker.Cong_1'(s(bool,X272))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X271),s('type.num.num','const.num.|48|_0'))) | s(bool,X270) = s(bool,X272) | s(bool,X270) = s(bool,'const.bool.BOUNDED_1'(s(bool,X273)))) ) | $spl0),
  inference(superposition,[],[f13848,f16957])).
fof(f16957,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))) | s(bool,X10) = s(bool,'const.bool.BOUNDED_1'(s(bool,X12)))) ) | $spl0),
  inference(subsumption_resolution,[],[f16384,f437])).
fof(f16384,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X10) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))) | s(bool,X10) = s(bool,'const.bool.BOUNDED_1'(s(bool,X12)))) ) | $spl0),
  inference(superposition,[],[f13471,f1945])).
fof(f34311,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34310,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34310,plain,(
  ( ! [X105,X106] : (s(bool,X106) = s(bool,'const.marker.Cong_1'(s(bool,X105))) | s(bool,X105) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33775,f17030])).
fof(f33775,plain,(
  ( ! [X105,X106] : (s(bool,'T_0') != s(bool,X106) | s(bool,X106) = s(bool,'const.marker.Cong_1'(s(bool,X105))) | s(bool,X105) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f15323])).
fof(f15323,plain,(
  ( ! [X105,X106] : (s(bool,X106) = s(bool,'const.marker.Cong_1'(s(bool,X105))) | s(bool,'F_0') != s(bool,X106) | s(bool,X105) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f15247,f271])).
fof(f15247,plain,(
  ( ! [X105,X106] : (s(bool,'F_0') != s(bool,X106) | s(bool,X106) = s(bool,'const.marker.Cong_1'(s(bool,'const.marker.Cong_1'(s(bool,X105))))) | s(bool,X105) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f14826,f4389])).
fof(f34309,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34308,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34308,plain,(
  ( ! [X80,X79] : (s(bool,X80) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) ) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f34307,f17030])).
fof(f34307,plain,(
  ( ! [X80,X79] : (s(bool,X80) = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33773,f17030])).
fof(f33773,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') != s(bool,X80) | s(bool,X80) = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f15235])).
fof(f15235,plain,(
  ( ! [X80,X79] : (s(bool,'F_0') != s(bool,X80) | s(bool,X80) = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) ) | $spl0),
  inference(superposition,[],[f14826,f6980])).
fof(f34288,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34287,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34287,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34286])).
fof(f34286,plain,(
  ( ! [X19,X20] : (s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,X20)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33743,f32370])).
fof(f32370,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,X0))) | s(bool,X1) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(condensation,[],[f32369])).
fof(f32369,plain,(
  ( ! [X30,X31,X32] : (s(bool,X32) = s(bool,'const.marker.Cong_1'(s(bool,X31))) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f30792])).
fof(f30792,plain,(
  ( ! [X30,X31,X32] : (s(bool,X32) = s(bool,'const.marker.Cong_1'(s(bool,X31))) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(superposition,[],[f13452,f2308])).
fof(f33743,plain,(
  ( ! [X19,X20] : (s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X20) | s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,X20)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14834])).
fof(f14834,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,X20) | s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,X20))) | s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f14546,f8912])).
fof(f34285,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34284,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34284,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34283])).
fof(f34283,plain,(
  ( ! [X17,X18] : (s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,X18))) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33742,f32370])).
fof(f33742,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') != s(bool,X18) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,X18))) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14833])).
fof(f14833,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') != s(bool,X18) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,X18))) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f14546,f4389])).
fof(f34282,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34281,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34281,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,'const.marker.Cong_1'(s(bool,X12))) | s(bool,X11) = s(bool,'const.bool.BOUNDED_1'(s(bool,X13)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33741,f22362])).
fof(f22362,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,'const.marker.Cong_1'(s(bool,X42))) | s(bool,X43) = s(bool,'const.bool.BOUNDED_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,X42)) ) | $spl0),
  inference(subsumption_resolution,[],[f21609,f438])).
fof(f21609,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X43) = s(bool,'const.marker.Cong_1'(s(bool,X42))) | s(bool,X43) = s(bool,'const.bool.BOUNDED_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X41),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X42)) ) | $spl0),
  inference(superposition,[],[f16957,f2308])).
fof(f33741,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,X12) | s(bool,X11) = s(bool,'const.marker.Cong_1'(s(bool,X12))) | s(bool,X11) = s(bool,'const.bool.BOUNDED_1'(s(bool,X13)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14831])).
fof(f14831,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,X12) | s(bool,X11) = s(bool,'const.marker.Cong_1'(s(bool,X12))) | s(bool,X11) = s(bool,'const.bool.BOUNDED_1'(s(bool,X13)))) ) | $spl0),
  inference(superposition,[],[f14546,f1945])).
fof(f34280,plain,(
  $false | ($spl0 | $spl6 | $spl36)),
  inference(global_subsumption,[],[f34279,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34279,plain,(
  ( ! [X10,X9] : (s(bool,X9) = s(bool,'const.marker.Cong_1'(s(bool,X10))) | s(bool,X9) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(subsumption_resolution,[],[f33740,f17030])).
fof(f33740,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,X10) | s(bool,X9) = s(bool,'const.marker.Cong_1'(s(bool,X10))) | s(bool,X9) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14830])).
fof(f14830,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,X9) = s(bool,'const.marker.Cong_1'(s(bool,X10))) | s(bool,X9) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f14546,f1943])).
fof(f1943,plain,(
  ( ! [X28] : (s(bool,X28) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X28)))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1583,f437])).
fof(f1583,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X28) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X28)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1004,f806])).
fof(f806,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f798,f282])).
fof(f34276,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34275,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34275,plain,(
  ( ! [X76,X75] : (s(bool,X76) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X75)) ) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f34274,f17030])).
fof(f34274,plain,(
  ( ! [X76,X75] : (s(bool,X76) = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X75)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33729,f17030])).
fof(f33729,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X76))) | s(bool,X76) = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X75)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14743])).
fof(f14743,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X76))) | s(bool,X76) = s(bool,'const.marker.Cong_1'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X75)) ) | $spl0),
  inference(superposition,[],[f14546,f6980])).
fof(f34272,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34271,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34271,plain,(
  ( ! [X107,X108] : (s(bool,X107) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X108) = s(bool,'const.marker.Cong_1'(s(bool,X107)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33724,f17030])).
fof(f33724,plain,(
  ( ! [X107,X108] : (s(bool,X107) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X108) | s(bool,X108) = s(bool,'const.marker.Cong_1'(s(bool,X107)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14697])).
fof(f14697,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') != s(bool,X108) | s(bool,X108) = s(bool,'const.marker.Cong_1'(s(bool,X107))) | s(bool,X107) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f14610,f271])).
fof(f14610,plain,(
  ( ! [X107,X108] : (s(bool,X108) != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X108) = s(bool,'const.marker.Cong_1'(s(bool,X107))) | s(bool,X107) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f13848,f8912])).
fof(f34270,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34269,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34269,plain,(
  ( ! [X101,X99,X100] : (s(bool,X100) = s(bool,'const.marker.Cong_1'(s(bool,X99))) | s(bool,X99) = s(bool,'const.bool.BOUNDED_1'(s(bool,X101)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33723,f17030])).
fof(f33723,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,X100) | s(bool,X100) = s(bool,'const.marker.Cong_1'(s(bool,X99))) | s(bool,X99) = s(bool,'const.bool.BOUNDED_1'(s(bool,X101)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14695])).
fof(f14695,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') != s(bool,X100) | s(bool,X100) = s(bool,'const.marker.Cong_1'(s(bool,X99))) | s(bool,X99) = s(bool,'const.bool.BOUNDED_1'(s(bool,X101)))) ) | $spl0),
  inference(forward_demodulation,[],[f14607,f271])).
fof(f14607,plain,(
  ( ! [X101,X99,X100] : (s(bool,X100) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X100) = s(bool,'const.marker.Cong_1'(s(bool,X99))) | s(bool,X99) = s(bool,'const.bool.BOUNDED_1'(s(bool,X101)))) ) | $spl0),
  inference(superposition,[],[f13848,f1945])).
fof(f34268,plain,(
  $false | ($spl0 | $spl6 | $spl36)),
  inference(global_subsumption,[],[f34267,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34267,plain,(
  ( ! [X97,X98] : (s(bool,X98) = s(bool,'const.marker.Cong_1'(s(bool,X97))) | s(bool,X97) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(subsumption_resolution,[],[f33722,f17030])).
fof(f33722,plain,(
  ( ! [X97,X98] : (s(bool,'T_0') != s(bool,X98) | s(bool,X98) = s(bool,'const.marker.Cong_1'(s(bool,X97))) | s(bool,X97) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14694])).
fof(f14694,plain,(
  ( ! [X97,X98] : (s(bool,'F_0') != s(bool,X98) | s(bool,X98) = s(bool,'const.marker.Cong_1'(s(bool,X97))) | s(bool,X97) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f14606,f271])).
fof(f14606,plain,(
  ( ! [X97,X98] : (s(bool,X98) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X98) = s(bool,'const.marker.Cong_1'(s(bool,X97))) | s(bool,X97) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f13848,f1943])).
fof(f34266,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34265,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34265,plain,(
  ( ! [X83,X84] : (s(bool,X84) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33721,f17030])).
fof(f33721,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') != s(bool,X84) | s(bool,X84) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14692])).
fof(f14692,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,X84) | s(bool,X84) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f14599,f271])).
fof(f14599,plain,(
  ( ! [X83,X84] : (s(bool,X84) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X84) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(superposition,[],[f13848,f6330])).
fof(f34264,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34263,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34263,plain,(
  ( ! [X80,X79] : (s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33720,f17030])).
fof(f33720,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') != s(bool,X80) | s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14690])).
fof(f14690,plain,(
  ( ! [X80,X79] : (s(bool,'F_0') != s(bool,X80) | s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) ) | $spl0),
  inference(forward_demodulation,[],[f14597,f271])).
fof(f14597,plain,(
  ( ! [X80,X79] : (s(bool,X80) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) ) | $spl0),
  inference(superposition,[],[f13848,f6980])).
fof(f34262,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34261,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34261,plain,(
  ( ! [X78,X77] : (s(bool,X78) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X77),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33719,f17030])).
fof(f33719,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,X78) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X77),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14689])).
fof(f14689,plain,(
  ( ! [X78,X77] : (s(bool,'F_0') != s(bool,X78) | s(bool,X78) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X77),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f14596,f271])).
fof(f14596,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X78) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X77),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f13848,f462])).
fof(f34260,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34259,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34259,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,X76) = s(bool,'const.list.LIST_REL_3'(s(fun(X71,fun(X72,bool)),X73),s('type.list.list'(X71),'const.list.NIL_0'),s('type.list.list'(X72),'const.list.CONS_2'(s(X72,X74),s('type.list.list'(X72),X75)))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33718,f17030])).
fof(f33718,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'T_0') != s(bool,X76) | s(bool,X76) = s(bool,'const.list.LIST_REL_3'(s(fun(X71,fun(X72,bool)),X73),s('type.list.list'(X71),'const.list.NIL_0'),s('type.list.list'(X72),'const.list.CONS_2'(s(X72,X74),s('type.list.list'(X72),X75)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14688])).
fof(f14688,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'F_0') != s(bool,X76) | s(bool,X76) = s(bool,'const.list.LIST_REL_3'(s(fun(X71,fun(X72,bool)),X73),s('type.list.list'(X71),'const.list.NIL_0'),s('type.list.list'(X72),'const.list.CONS_2'(s(X72,X74),s('type.list.list'(X72),X75)))))) ) | $spl0),
  inference(forward_demodulation,[],[f14595,f271])).
fof(f14595,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,X76) != s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X76) = s(bool,'const.list.LIST_REL_3'(s(fun(X71,fun(X72,bool)),X73),s('type.list.list'(X71),'const.list.NIL_0'),s('type.list.list'(X72),'const.list.CONS_2'(s(X72,X74),s('type.list.list'(X72),X75)))))) ) | $spl0),
  inference(superposition,[],[f13848,f380])).
fof(f34257,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34256,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34256,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34255])).
fof(f34255,plain,(
  ( ! [X19,X20] : (s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X19) = s(bool,X20)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33710,f13452])).
fof(f33710,plain,(
  ( ! [X19,X20] : (s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X20) | s(bool,X19) = s(bool,X20)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14632])).
fof(f14632,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,X20) | s(bool,X19) = s(bool,X20) | s(bool,X19) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f13848,f8912])).
fof(f34254,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34253,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34253,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34252])).
fof(f34252,plain,(
  ( ! [X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33709,f13452])).
fof(f33709,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') != s(bool,X18) | s(bool,X17) = s(bool,X18) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14631])).
fof(f14631,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') != s(bool,X18) | s(bool,X17) = s(bool,X18) | s(bool,X17) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f13848,f4389])).
fof(f34251,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34250,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34250,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.BOUNDED_1'(s(bool,X1)))) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34249])).
fof(f34249,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X12) | s(bool,X11) = s(bool,'const.bool.BOUNDED_1'(s(bool,X13)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33708,f13451])).
fof(f33708,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,X12) | s(bool,X11) = s(bool,X12) | s(bool,X11) = s(bool,'const.bool.BOUNDED_1'(s(bool,X13)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14629])).
fof(f14629,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,X12) | s(bool,X11) = s(bool,X12) | s(bool,X11) = s(bool,'const.bool.BOUNDED_1'(s(bool,X13)))) ) | $spl0),
  inference(superposition,[],[f13848,f1945])).
fof(f34248,plain,(
  $false | ($spl0 | $spl6 | $spl36)),
  inference(global_subsumption,[],[f34247,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34247,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(condensation,[],[f34246])).
fof(f34246,plain,(
  ( ! [X10,X9] : (s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(subsumption_resolution,[],[f33707,f13450])).
fof(f13450,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,X8) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f12332,f437])).
fof(f12332,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,X8) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f9373,f1943])).
fof(f33707,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,X10) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14628])).
fof(f14628,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f13848,f1943])).
fof(f34238,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34237,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34237,plain,(
  ( ! [X103,X104] : (s(bool,X103) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X104) = s(bool,'const.marker.Cong_1'(s(bool,X103)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33690,f13476])).
fof(f33690,plain,(
  ( ! [X103,X104] : (s(bool,X103) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X104))) | s(bool,X104) = s(bool,'const.marker.Cong_1'(s(bool,X103)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14554])).
fof(f14554,plain,(
  ( ! [X103,X104] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X104))) | s(bool,X104) = s(bool,'const.marker.Cong_1'(s(bool,X103))) | s(bool,X103) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f13848,f8912])).
fof(f34236,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34235,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34235,plain,(
  ( ! [X101,X102] : (s(bool,X102) = s(bool,'const.marker.Cong_1'(s(bool,X101))) | s(bool,X101) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33689,f13476])).
fof(f33689,plain,(
  ( ! [X101,X102] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X102))) | s(bool,X102) = s(bool,'const.marker.Cong_1'(s(bool,X101))) | s(bool,X101) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14553])).
fof(f14553,plain,(
  ( ! [X101,X102] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X102))) | s(bool,X102) = s(bool,'const.marker.Cong_1'(s(bool,X101))) | s(bool,X101) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f13848,f4389])).
fof(f34234,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34233,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34233,plain,(
  ( ! [X97,X95,X96] : (s(bool,X96) = s(bool,'const.marker.Cong_1'(s(bool,X95))) | s(bool,X95) = s(bool,'const.bool.BOUNDED_1'(s(bool,X97)))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33688,f13475])).
fof(f33688,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X96))) | s(bool,X96) = s(bool,'const.marker.Cong_1'(s(bool,X95))) | s(bool,X95) = s(bool,'const.bool.BOUNDED_1'(s(bool,X97)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14551])).
fof(f14551,plain,(
  ( ! [X97,X95,X96] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X96))) | s(bool,X96) = s(bool,'const.marker.Cong_1'(s(bool,X95))) | s(bool,X95) = s(bool,'const.bool.BOUNDED_1'(s(bool,X97)))) ) | $spl0),
  inference(superposition,[],[f13848,f1945])).
fof(f34232,plain,(
  $false | ($spl0 | $spl6 | $spl36)),
  inference(global_subsumption,[],[f34231,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34231,plain,(
  ( ! [X94,X93] : (s(bool,X94) = s(bool,'const.marker.Cong_1'(s(bool,X93))) | s(bool,X93) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(subsumption_resolution,[],[f33687,f13474])).
fof(f13474,plain,(
  ( ! [X92,X91] : (s(bool,X92) = s(bool,'const.marker.Cong_1'(s(bool,X91))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X92))) | s(bool,X91) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f12398,f437])).
fof(f12398,plain,(
  ( ! [X92,X91] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X92) = s(bool,'const.marker.Cong_1'(s(bool,X91))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X92))) | s(bool,X91) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f9373,f1943])).
fof(f33687,plain,(
  ( ! [X94,X93] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X94))) | s(bool,X94) = s(bool,'const.marker.Cong_1'(s(bool,X93))) | s(bool,X93) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14550])).
fof(f14550,plain,(
  ( ! [X94,X93] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X94))) | s(bool,X94) = s(bool,'const.marker.Cong_1'(s(bool,X93))) | s(bool,X93) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f13848,f1943])).
fof(f34228,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34227,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34227,plain,(
  ( ! [X80,X79] : (s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33682,f13473])).
fof(f33682,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X80))) | s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14543])).
fof(f14543,plain,(
  ( ! [X80,X79] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X80))) | s(bool,X80) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(superposition,[],[f13848,f6330])).
fof(f34226,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34225,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34225,plain,(
  ( ! [X76,X75] : (s(bool,X76) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X75)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33681,f13472])).
fof(f33681,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X76))) | s(bool,X76) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X75)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14541])).
fof(f14541,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X76))) | s(bool,X76) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X75)) ) | $spl0),
  inference(superposition,[],[f13848,f6980])).
fof(f34224,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34223,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34223,plain,(
  ( ! [X74,X73] : (s(bool,X74) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33680,f13471])).
fof(f33680,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X74))) | s(bool,X74) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X73),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14540])).
fof(f34222,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34221,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34221,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,X72) = s(bool,'const.list.LIST_REL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.list.list'(X67),'const.list.NIL_0'),s('type.list.list'(X68),'const.list.CONS_2'(s(X68,X70),s('type.list.list'(X68),X71)))))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33679,f13470])).
fof(f33679,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X72))) | s(bool,X72) = s(bool,'const.list.LIST_REL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.list.list'(X67),'const.list.NIL_0'),s('type.list.list'(X68),'const.list.CONS_2'(s(X68,X70),s('type.list.list'(X68),X71)))))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14539])).
fof(f14539,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X72))) | s(bool,X72) = s(bool,'const.list.LIST_REL_3'(s(fun(X67,fun(X68,bool)),X69),s('type.list.list'(X67),'const.list.NIL_0'),s('type.list.list'(X68),'const.list.CONS_2'(s(X68,X70),s('type.list.list'(X68),X71)))))) ) | $spl0),
  inference(superposition,[],[f13848,f380])).
fof(f34220,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34219,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34219,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34218])).
fof(f34218,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | ($spl0 | $spl36)),
  inference(condensation,[],[f34217])).
fof(f34217,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X4) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33676,f17030])).
fof(f33676,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,X5))) | s(bool,X4) = s(bool,X5) | s(bool,X4) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f14519])).
fof(f14519,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,'const.marker.Cong_1'(s(bool,X5))) | s(bool,X4) = s(bool,X5) | s(bool,X4) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f13848,f879])).
fof(f34184,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34183,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34183,plain,(
  ( ! [X47] : (s(bool,X47) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33580,f14546])).
fof(f33580,plain,(
  ( ! [X47] : (s(bool,X47) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X47)))) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f8912])).
fof(f34181,plain,(
  $false | ($spl0 | $spl36)),
  inference(global_subsumption,[],[f34180,f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f34180,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33578,f5187])).
fof(f33578,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f8650])).
fof(f34143,plain,(
  $false | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f33454,f271])).
fof(f33454,plain,(
  s(bool,'T_0') != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | ($spl0 | $spl36)),
  inference(backward_demodulation,[],[f17030,f2014])).
fof(f34129,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33398,f264])).
fof(f33398,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(backward_demodulation,[],[f17030,f426])).
fof(f34128,plain,(
  $false | $spl36),
  inference(global_subsumption,[],[f263,f262,f264,f426,f271,f272,f273,f274,f275,f276,f277,f282,f281,f286,f285,f336,f337,f338,f343,f344,f345,f346,f356,f357,f358,f359,f364,f361,f366,f365,f367,f368,f369,f379,f380,f393,f388,f387,f386,f385,f397,f438,f437,f462,f479,f486,f489,f546,f573,f582,f621,f623,f639,f664,f698,f717,f746,f747,f764,f786,f787,f789,f792,f799,f839,f837,f840,f841,f843,f844,f845,f846,f966,f1488,f3067,f3066,f3065,f3064,f3258,f4075,f4561,f4553,f5515,f5596,f5965,f6095,f6100,f11245,f11428,f11872,f10759,f5309,f5283,f2134,f2133,f1506,f1461,f11471,f11309,f10842,f4579,f4496,f4480,f5310,f3793,f3785,f3861,f3794,f360,f4054,f4053,f4051,f4050,f4048,f4047,f396,f395,f381,f11888,f2153,f10737,f5629,f5516,f2154,f2122,f2152,f2145,f2098,f2151,f2149,f2146,f10707,f10688,f5253,f2150,f2110,f2094,f2137,f5299,f2136,f2135,f2093,f2092,f355,f354,f353,f350,f349,f4497,f2513,f793,f794,f11412,f4542,f1485,f348,f347,f3884,f2366,f6116,f5583,f383,f3277,f378,f377,f375,f374,f2537,f5286,f4607,f11533,f4541,f5287,f11532,f11411,f5665,f5582,f5664,f390,f389,f376,f372,f384,f382,f394,f5284,f5285,f2536,f2535,f4569,f11461,f5619,f4046,f4049,f4543,f11413,f5584,f392,f391,f373,f4052,f4546,f11421,f5588,f4055,f4545,f11422,f5587,f17030,f34127])).
fof(f32429,plain,(
  ( ! [X83,X81] : (s(bool,X81) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X83) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X81) = s(bool,X83)) ) | $spl40),
  inference(cnf_transformation,[],[f32429_D])).
fof(f32429_D,plain,(
  ( ! [X83,X81] : (s(bool,X81) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X83) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X81) = s(bool,X83)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f17030,plain,(
  ( ! [X154] : (s(bool,'T_0') = s(bool,X154)) ) | $spl36),
  inference(cnf_transformation,[],[f17030_D])).
fof(f17030_D,plain,(
  ( ! [X154] : (s(bool,'T_0') = s(bool,X154)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f32357,plain,(
  ( ! [X80,X81,X82] : (s(bool,'T_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,X82) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X80) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f30908])).
fof(f30908,plain,(
  ( ! [X80,X81,X82] : (s(bool,'T_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,X82) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X81) | s(bool,X80) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f13452,f13452])).
fof(f32361,plain,(
  ( ! [X85,X86,X84] : (s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X84) | s(bool,X86) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X85) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f30810])).
fof(f30810,plain,(
  ( ! [X85,X86,X84] : (s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X84) | s(bool,X86) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X84) | s(bool,X85) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f13452,f13452])).
fof(f32386,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') = s(bool,X82) | s(bool,X81) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X81) = s(bool,X83) | s(bool,X83) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f30727])).
fof(f30727,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,X81) = s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,X81) = s(bool,X83) | s(bool,X83) = s(bool,'const.marker.Cong_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f13452,f13452])).
fof(f27088,plain,(
  ( ! [X55,X53] : (s(bool,X53) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X53) = s(bool,X55) | s(bool,X55) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl38),
  inference(cnf_transformation,[],[f27088_D])).
fof(f27088_D,plain,(
  ( ! [X55,X53] : (s(bool,X53) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X53) = s(bool,X55) | s(bool,X55) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f6472,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2)) ) | $spl12),
  inference(cnf_transformation,[],[f6472_D])).
fof(f6472_D,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f26919,plain,(
  ( ! [X52,X50,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X50) = s(bool,X52) | s(bool,X50) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f25724])).
fof(f25724,plain,(
  ( ! [X52,X50,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X50) = s(bool,X52) | s(bool,X50) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X51)) ) | $spl0),
  inference(superposition,[],[f9876,f9876])).
fof(f26930,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X56) | s(bool,X57) = s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f25639])).
fof(f25639,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X56) | s(bool,X57) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X56)) ) | $spl0),
  inference(superposition,[],[f9876,f9876])).
fof(f26959,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,X54) | s(bool,X53) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X55) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X53) = s(bool,X55)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f25568])).
fof(f25568,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,X54) | s(bool,X53) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X54) | s(bool,X55) = s(bool,'const.marker.Cong_1'(s(bool,'F_0'))) | s(bool,X53) = s(bool,X55)) ) | $spl0),
  inference(superposition,[],[f9876,f9876])).
fof(f17643,plain,(
  $false | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f17493,f437])).
fof(f17493,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl30)),
  inference(condensation,[],[f17492])).
fof(f17492,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl30)),
  inference(duplicate_literal_removal,[],[f17491])).
fof(f17491,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl0 | $spl30)),
  inference(condensation,[],[f17273])).
fof(f17273,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X12) = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X12)) ) | ($spl0 | $spl30)),
  inference(backward_demodulation,[],[f17017,f1566])).
fof(f17642,plain,(
  $false | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f17497,f437])).
fof(f17497,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl30)),
  inference(condensation,[],[f17496])).
fof(f17496,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl30)),
  inference(condensation,[],[f17495])).
fof(f17495,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X3)) ) | ($spl0 | $spl30)),
  inference(duplicate_literal_removal,[],[f17494])).
fof(f17494,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl0 | $spl30)),
  inference(condensation,[],[f17272])).
fof(f17272,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | ($spl0 | $spl30)),
  inference(backward_demodulation,[],[f17017,f1564])).
fof(f17641,plain,(
  $false | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f17640,f437])).
fof(f17640,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl30)),
  inference(condensation,[],[f17639])).
fof(f17639,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,X1)) ) | ($spl0 | $spl30)),
  inference(forward_demodulation,[],[f17638,f17017])).
fof(f17638,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X1)))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f17499,f437])).
fof(f17499,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X1)))) ) | ($spl0 | $spl30)),
  inference(condensation,[],[f17498])).
fof(f17498,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X1) = s(bool,X3) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X3)))) ) | ($spl0 | $spl30)),
  inference(condensation,[],[f17271])).
fof(f17271,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X3)))) ) | ($spl0 | $spl30)),
  inference(backward_demodulation,[],[f17017,f1563])).
fof(f17634,plain,(
  $false | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f17501,f437])).
fof(f17501,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl30)),
  inference(duplicate_literal_removal,[],[f17500])).
fof(f17500,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl30)),
  inference(condensation,[],[f17267])).
fof(f17267,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,X33) | s(bool,'F_0') = s(bool,X33)) ) | ($spl0 | $spl30)),
  inference(backward_demodulation,[],[f17017,f1004])).
fof(f17022,plain,(
  ( ! [X130,X129] : (s(bool,X130) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X129),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X130)) ) | $spl34),
  inference(cnf_transformation,[],[f17022_D])).
fof(f17022_D,plain,(
  ( ! [X130,X129] : (s(bool,X130) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X129),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X130)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f17017,plain,(
  ( ! [X115] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X115)))) ) | $spl30),
  inference(cnf_transformation,[],[f17017_D])).
fof(f17017_D,plain,(
  ( ! [X115] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X115)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f17097,plain,(
  ( ! [X364,X365,X363] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X363))) | s(bool,X365) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X364),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X365)) ) | $spl0),
  inference(subsumption_resolution,[],[f17096,f14826])).
fof(f17096,plain,(
  ( ! [X364,X365,X363] : (s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X363))) | s(bool,X365) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X364),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X365) | s(bool,'T_0') = s(bool,X363)) ) | $spl0),
  inference(forward_demodulation,[],[f16532,f271])).
fof(f16532,plain,(
  ( ! [X364,X365,X363] : (s(bool,X365) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X364),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X365) | s(bool,'T_0') = s(bool,X363) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,'const.marker.Cong_1'(s(bool,X363)))))) ) | $spl0),
  inference(superposition,[],[f2047,f13471])).
fof(f17029,plain,(
  ( ! [X156,X154,X155] : (s(bool,X156) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X156) | s(bool,'T_0') = s(bool,X154)) ) | $spl0),
  inference(subsumption_resolution,[],[f16469,f13848])).
fof(f16469,plain,(
  ( ! [X156,X154,X155] : (s(bool,X156) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X156) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X154)))) ) | $spl0),
  inference(superposition,[],[f6213,f13471])).
fof(f16822,plain,(
  ( ! [X138,X139,X137] : (s(bool,X139) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X138),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X137))) | s(bool,'F_0') = s(bool,X139)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f16461])).
fof(f16461,plain,(
  ( ! [X138,X139,X137] : (s(bool,X139) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X138),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X137))) | s(bool,'F_0') = s(bool,X139) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X137)))) ) | $spl0),
  inference(superposition,[],[f3336,f13471])).
fof(f17021,plain,(
  ( ! [X130,X128,X129] : (s(bool,X130) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X129),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X130) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X128)))) ) | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f16458,f6708])).
fof(f6708,plain,(
  ( ! [X62,X63] : (s(bool,X62) != s(bool,'const.bool.BOUNDED_1'(s(bool,X63))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X62)))) ) | $spl18),
  inference(superposition,[],[f6479,f271])).
fof(f16458,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,X130) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X129),s('type.num.num','const.num.|48|_0'))) | s(bool,X128) = s(bool,'const.bool.BOUNDED_1'(s(bool,X131))) | s(bool,'F_0') = s(bool,X130) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X128)))) ) | $spl0),
  inference(superposition,[],[f3201,f13471])).
fof(f17018,plain,(
  ( ! [X116,X117] : (s(bool,X117) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X116),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X117)))) ) | $spl32),
  inference(cnf_transformation,[],[f17018_D])).
fof(f17018_D,plain,(
  ( ! [X116,X117] : (s(bool,X117) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X116),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X117)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f17016,plain,(
  ( ! [X116,X117,X115] : (s(bool,X117) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X116),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X117))) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X115)))) ) | $spl0),
  inference(subsumption_resolution,[],[f16454,f14826])).
fof(f16454,plain,(
  ( ! [X116,X117,X115] : (s(bool,X117) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X116),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X117))) | s(bool,'T_0') = s(bool,X115) | s(bool,'T_0') = s(bool,'const.marker.Cong_1'(s(bool,X115)))) ) | $spl0),
  inference(superposition,[],[f2374,f13471])).
fof(f14441,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f13951,f426])).
fof(f13951,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f437,f4073,f361])).
fof(f14440,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f13953,f4073])).
fof(f13953,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f438,f4073,f361])).
fof(f14439,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f13955,f4073])).
fof(f13955,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f7074,f4073,f361])).
fof(f14438,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f13957,f4073])).
fof(f13957,plain,(
  p(s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f2014,f4073,f361])).
fof(f14437,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f13969,f426])).
fof(f13969,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f437,f4073,f361])).
fof(f14436,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f13971,f4073])).
fof(f13971,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f438,f4073,f361])).
fof(f14435,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f13973,f4073])).
fof(f13973,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f7074,f4073,f361])).
fof(f14434,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f13975,f4073])).
fof(f13975,plain,(
  p(s(bool,'const.marker.Cong_1'(s(bool,'F_0')))) | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f2014,f4073,f361])).
fof(f13982,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4074,f4073,f396])).
fof(f13980,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f447,f4073,f281])).
fof(f13978,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f437,f4073,f286])).
fof(f13976,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4073,f2014,f4073,f361])).
fof(f13974,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4073,f7074,f4073,f361])).
fof(f13972,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4073,f438,f4073,f361])).
fof(f13970,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f426,f437,f4073,f361])).
fof(f13968,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4073,f2014,f4073,f361])).
fof(f13966,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4073,f7074,f4073,f361])).
fof(f13964,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4073,f438,f4073,f361])).
fof(f13962,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4073,f437,f4073,f361])).
fof(f13958,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4073,f2014,f4073,f361])).
fof(f13956,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4073,f7074,f4073,f361])).
fof(f13954,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4073,f438,f4073,f361])).
fof(f13952,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f426,f437,f4073,f361])).
fof(f13950,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4073,f2014,f4073,f361])).
fof(f13948,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f4073,f7074,f4073,f361])).
fof(f13946,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4073,f438,f4073,f361])).
fof(f13944,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4073,f437,f4073,f361])).
fof(f13940,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f2014,f2014,f4073,f546])).
fof(f13938,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f7074,f7074,f4073,f546])).
fof(f13936,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f438,f438,f4073,f546])).
fof(f13934,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f437,f437,f4073,f546])).
fof(f13924,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f264,f4073])).
fof(f13906,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) | $spl28),
  inference(cnf_transformation,[],[f13906_D])).
fof(f13906_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f13904,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) | ~p(s(bool,X0))) ) | $spl0),
  inference(forward_demodulation,[],[f13903,f273])).
fof(f13903,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13868,f285])).
fof(f13868,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f381,f785])).
fof(f10975,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) != s(bool,'const.bool.BOUNDED_1'(s(bool,X38))) | s(bool,X39) = s(bool,'const.bool.BOUNDED_1'(s(bool,X40)))) ) | $spl26),
  inference(cnf_transformation,[],[f10975_D])).
fof(f10975_D,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) != s(bool,'const.bool.BOUNDED_1'(s(bool,X38))) | s(bool,X39) = s(bool,'const.bool.BOUNDED_1'(s(bool,X40)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f10974,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X39) != s(bool,'const.bool.BOUNDED_1'(s(bool,X38))) | s(bool,X39) = s(bool,'const.bool.BOUNDED_1'(s(bool,X40))) | s(bool,'F_0') = s(bool,X37)) ) | $spl0),
  inference(subsumption_resolution,[],[f10861,f621])).
fof(f10861,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X39) != s(bool,'const.bool.BOUNDED_1'(s(bool,X38))) | s(bool,X39) = s(bool,'const.bool.BOUNDED_1'(s(bool,X40))) | s(bool,'F_0') = s(bool,X37) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X37)))) ) | $spl0),
  inference(superposition,[],[f3201,f1945])).
fof(f10973,plain,(
  ( ! [X35,X36] : (s(bool,X35) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,X35) = s(bool,'const.bool.BOUNDED_1'(s(bool,X36)))) ) | $spl24),
  inference(cnf_transformation,[],[f10973_D])).
fof(f10973_D,plain,(
  ( ! [X35,X36] : (s(bool,X35) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,X35) = s(bool,'const.bool.BOUNDED_1'(s(bool,X36)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f10972,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,X35) = s(bool,'const.bool.BOUNDED_1'(s(bool,X36))) | s(bool,'F_0') = s(bool,X34)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f10860,f621])).
fof(f10860,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,X35) = s(bool,'const.bool.BOUNDED_1'(s(bool,X36))) | s(bool,'F_0') = s(bool,X34) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X34)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f3201,f1943])).
fof(f10971,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,X4) = s(bool,'const.bool.BOUNDED_1'(s(bool,X5)))) ) | $spl22),
  inference(cnf_transformation,[],[f10971_D])).
fof(f10971_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,X4) = s(bool,'const.bool.BOUNDED_1'(s(bool,X5)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f10961,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,X4) = s(bool,'const.bool.BOUNDED_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f10852])).
fof(f10852,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,X4) = s(bool,'const.bool.BOUNDED_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f3201,f785])).
fof(f6679,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f343,f344,f6669,f393])).
fof(f6669,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X0,X1,X2,X4,X3)) ) | $spl12),
  inference(subsumption_resolution,[],[f6497,f426])).
fof(f6497,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | ~sP0(X0,X1,X2,X4,X3)) ) | $spl12),
  inference(backward_demodulation,[],[f6472,f396])).
fof(f6677,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f343,f343,f6669,f393])).
fof(f6673,plain,(
  $false | ($spl0 | $spl12)),
  inference(unit_resulting_resolution,[],[f4074,f6669])).
fof(f6480,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31)) ) | $spl20),
  inference(cnf_transformation,[],[f6480_D])).
fof(f6480_D,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6481,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(subsumption_resolution,[],[f6403,f621])).
fof(f6403,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X32)))) ) | $spl0),
  inference(superposition,[],[f6213,f4389])).
fof(f6434,plain,(
  ( ! [X30,X31] : (s(bool,X31) != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f6402])).
fof(f6402,plain,(
  ( ! [X30,X31] : (s(bool,X31) != s(bool,'const.marker.Cong_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X30)) ) | $spl0),
  inference(superposition,[],[f6213,f4160])).
fof(f6479,plain,(
  ( ! [X28,X29] : (s(bool,X29) != s(bool,'const.bool.BOUNDED_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X29)) ) | $spl18),
  inference(cnf_transformation,[],[f6479_D])).
fof(f6479_D,plain,(
  ( ! [X28,X29] : (s(bool,X29) != s(bool,'const.bool.BOUNDED_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f6478,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) != s(bool,'const.bool.BOUNDED_1'(s(bool,X28))) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X29)) ) | $spl0),
  inference(subsumption_resolution,[],[f6401,f621])).
fof(f6401,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) != s(bool,'const.bool.BOUNDED_1'(s(bool,X28))) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X27)))) ) | $spl0),
  inference(superposition,[],[f6213,f1945])).
fof(f6477,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,'T_0') = s(bool,X26)) ) | $spl16),
  inference(cnf_transformation,[],[f6477_D])).
fof(f6477_D,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f6476,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,'F_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X26)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f6400,f621])).
fof(f6400,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0'))) | s(bool,'F_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,'const.marker.Cong_1'(s(bool,X25)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f6213,f1943])).
fof(f6473,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl14),
  inference(cnf_transformation,[],[f6473_D])).
fof(f6473_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f6441,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f6392])).
fof(f6392,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f6213,f785])).
fof(f4072,plain,(
  sP0(sK1,sK2,sK4,sK3,'const.list.NIL_0') | $spl8),
  inference(cnf_transformation,[],[f4072_D])).
fof(f4072_D,plain,(
  sP0(sK1,sK2,sK4,sK3,'const.list.NIL_0') <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4070,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP0(sK1,sK2,sK4,sK3,'const.list.NIL_0')) ) | $spl0),
  inference(subsumption_resolution,[],[f4056,f285])).
fof(f4056,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP0(sK1,sK2,sK4,sK3,'const.list.NIL_0') | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f395,f785])).
fof(f796,plain,(
  ( ! [X0] : (s('type.list.list'(bool),X0) = s('type.list.list'(bool),sK6(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))) ) | $spl4),
  inference(cnf_transformation,[],[f796_D])).
fof(f796_D,plain,(
  ( ! [X0] : (s('type.list.list'(bool),X0) = s('type.list.list'(bool),sK6(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f798,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) | $spl6),
  inference(cnf_transformation,[],[f798_D])).
fof(f798_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f792,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(bool),bool),'const.list.HD_0')))) | s('type.list.list'(bool),X0) = s('type.list.list'(bool),sK6(X0,'const.list.EL_1'(s('type.num.num','const.num.|48|_0')),'type.list.list'(bool)))) )),
  inference(superposition,[],[f347,f275])).
fof(f430,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f429,f264])).
fof(f429,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f428,f368])).
fof(f428,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.NIL_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f423,f424])).
fof(f422,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) | $spl0),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f424,plain,(
  s('type.list.list'(sK2),'const.list.NIL_0') = s('type.list.list'(sK2),sK3) | $spl2),
  inference(cnf_transformation,[],[f424_D])).
fof(f424_D,plain,(
  s('type.list.list'(sK2),'const.list.NIL_0') = s('type.list.list'(sK2),sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f262,plain,(
  s('type.list.list'(sK2),'const.list.NIL_0') = s('type.list.list'(sK2),sK3) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))))),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  (p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) | s('type.list.list'(sK2),'const.list.NIL_0') = s('type.list.list'(sK2),sK3)) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) | s('type.list.list'(sK2),'const.list.NIL_0') != s('type.list.list'(sK2),sK3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f196])).
fof(f196,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X2)))) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2)) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X2)))) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X2)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X2)))) <~> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X2)))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X2)))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X7),X8)))) <=> s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X9),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X7),X8)))) <=> s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X8))),
  file('thm.list.LIST_REL_NIL|split|0','thm.list.LIST_REL_NIL|split|0')).
fof(f423,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) | $spl1),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f425,plain,(
  s('type.list.list'(sK2),'const.list.NIL_0') != s('type.list.list'(sK2),sK3) | $spl3),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  s('type.list.list'(sK2),'const.list.NIL_0') != s('type.list.list'(sK2),sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f263,plain,(
  s('type.list.list'(sK2),'const.list.NIL_0') != s('type.list.list'(sK2),sK3) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK4),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),sK3))))),
  inference(cnf_transformation,[],[f197])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 240
% Passive clauses: 3552
% Generated clauses: 59127
% Final active clauses: 157
% Final passive clauses: 1859
% Input formulas: 69
% Initial clauses: 152
% 
% Duplicate literals: 1700
% Trivial inequalities: 751
% Fw subsumption resolutions: 6123
% Fw demodulations: 11151
% Bw demodulations: 1407
% Condensations: 4210
% Global subsumptions: 102
% 
% Simple tautologies: 760
% Equational tautologies: 173
% Forward subsumptions: 30327
% Fw demodulations to eq. taut.: 133
% Bw demodulations to eq. taut.: 60
% 
% Unit resulting resolution: 1557
% Factoring: 1
% Forward superposition: 15991
% Backward superposition: 13873
% Self superposition: 401
% Equality factoring: 185
% Equality resolution: 70
% 
% Split clauses: 30
% Split components: 60
% SAT solver clauses: 3857
% SAT solver unit clauses: 114
% SAT solver binary clauses: 640
% 
% Sat splits: 30
% Sat splitting refutations: 141
% 
% Memory used [KB]: 32622
% Time elapsed: 3.117 s
% ------------------------------
% ------------------------------
% Success in time 4.451 s
