Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f2441,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2440,f2379,f177,f275,f273,f176,f274,f272])).
fof(f272,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f272_D])).
fof(f272_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f274,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f274_D])).
fof(f274_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f176,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f160])).
fof(f160,plain,(
  ? [X0,X1] : ((s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ? [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ~! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.numRing.num_rewrites|split|42','thm.numRing.num_rewrites|split|42')).
fof(f273,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f273_D])).
fof(f273_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f275,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f275_D])).
fof(f275_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f177,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f161])).
fof(f2379,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2346,f275])).
fof(f2346,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(equality_resolution,[],[f1599])).
fof(f1599,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | s('type.num.num',sK1) = s('type.num.num',X0)) ) | $spl0),
  inference(superposition,[],[f264,f272])).
fof(f264,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X10,X7] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10))) <=> s('type.num.num',X7) = s('type.num.num',X10))),
  file('thm.numRing.num_rewrites|split|42','thm.numeral.numeral_eq|split|6')).
fof(f2440,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f2439])).
fof(f2439,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f273,f274])).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 82
% Passive clauses: 285
% Generated clauses: 2250
% Final active clauses: 53
% Final passive clauses: 142
% Input formulas: 81
% Initial clauses: 95
% 
% Duplicate literals: 1
% Trivial inequalities: 13
% Fw subsumption resolutions: 25
% Fw demodulations: 456
% Bw demodulations: 11
% 
% Simple tautologies: 6
% Equational tautologies: 9
% Forward subsumptions: 1463
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 8
% Forward superposition: 948
% Backward superposition: 645
% Self superposition: 26
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 18
% Split components: 36
% SAT solver clauses: 24
% SAT solver unit clauses: 2
% SAT solver binary clauses: 18
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 18
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 22
% TWLsolver calls for satisfiability: 5
% 
% Memory used [KB]: 1663
% Time elapsed: 0.046 s
% ------------------------------
% ------------------------------
% Success in time 0.039 s
