Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.semi_ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.semi_ring
% SZS output start Proof for thm.semi_ring
fof(f8228,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8227,f8185,f29,f47,f30,f49,f31,f51,f32,f53,f45,f33,f54,f52,f50,f48,f46,f8223,f8221,f8225])).
fof(f8225,plain,(
  $false | ($spl0 | $spl7)),
  inference(trivial_inequality_removal,[],[f8224])).
fof(f8224,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) | ($spl0 | $spl7)),
  inference(forward_demodulation,[],[f52,f45])).
fof(f8221,plain,(
  $false | ($spl0 | $spl5)),
  inference(trivial_inequality_removal,[],[f8220])).
fof(f8220,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f50,f45])).
fof(f8223,plain,(
  $false | ($spl0 | $spl9)),
  inference(trivial_inequality_removal,[],[f8222])).
fof(f8222,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f54,f45])).
fof(f46,plain,(
  s('type.semi_ring.semi_ring'(sK0),sK1) != s('type.semi_ring.semi_ring'(sK0),sK2) | $spl1),
  inference(cnf_transformation,[],[f46_D])).
fof(f46_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),sK1) != s('type.semi_ring.semi_ring'(sK0),sK2) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f48,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | $spl3),
  inference(cnf_transformation,[],[f48_D])).
fof(f48_D,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f50,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | $spl5),
  inference(cnf_transformation,[],[f50_D])).
fof(f50_D,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f52,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | $spl7),
  inference(cnf_transformation,[],[f52_D])).
fof(f52_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f54,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | $spl9),
  inference(cnf_transformation,[],[f54_D])).
fof(f54_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f33,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s('type.semi_ring.semi_ring'(sK0),sK1) != s('type.semi_ring.semi_ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  (s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2) | (s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) & s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) & s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) & s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK2))))) & (s('type.semi_ring.semi_ring'(sK0),sK1) != s('type.semi_ring.semi_ring'(sK0),sK2) | s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f24])).
fof(f24,plain,(
  ? [X0,X1,X2] : ((s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) | (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))))) & (s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ? [X0,X1,X2] : ((s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) | (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))))) & (s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))))))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ? [X0,X1,X2] : (s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) <~> (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ~! [X0,X1,X2] : (s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) <=> (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) <=> (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.semi_ring.semi_ring'(X6),X7) = s('type.semi_ring.semi_ring'(X6),X8) <=> (s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.semi_ring.semi_ring'(X6),X7) = s('type.semi_ring.semi_ring'(X6),X8) <=> (s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X8)))))),
  file('thm.semi_ring.semi_ring_component_equality','thm.semi_ring.semi_ring_component_equality')).
fof(f45,plain,(
  s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2) | $spl0),
  inference(cnf_transformation,[],[f45_D])).
fof(f45_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f53,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | $spl8),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f32,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f51,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | $spl6),
  inference(cnf_transformation,[],[f51_D])).
fof(f51_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f31,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f49,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | $spl4),
  inference(cnf_transformation,[],[f49_D])).
fof(f49_D,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f30,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f47,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | $spl2),
  inference(cnf_transformation,[],[f47_D])).
fof(f47_D,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f29,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) | s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f8185,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f8178,f46])).
fof(f8178,plain,(
  s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2) | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(trivial_inequality_removal,[],[f8177])).
fof(f8177,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) | s('type.semi_ring.semi_ring'(sK0),sK1) = s('type.semi_ring.semi_ring'(sK0),sK2) | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(equality_resolution,[],[f2964])).
fof(f2964,plain,(
  ( ! [X0] : (s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s('type.semi_ring.semi_ring'(sK0),sK2) = s('type.semi_ring.semi_ring'(sK0),X0)) ) | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f2963,f47])).
fof(f2963,plain,(
  ( ! [X0] : (s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s('type.semi_ring.semi_ring'(sK0),sK2) = s('type.semi_ring.semi_ring'(sK0),X0) | s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),X0)))) ) | ($spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f2962,f49])).
fof(f2962,plain,(
  ( ! [X0] : (s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s('type.semi_ring.semi_ring'(sK0),sK2) = s('type.semi_ring.semi_ring'(sK0),X0) | s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),X0)))) ) | ($spl6 | $spl8)),
  inference(forward_demodulation,[],[f2949,f51])).
fof(f2949,plain,(
  ( ! [X0] : (s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s('type.semi_ring.semi_ring'(sK0),sK2) = s('type.semi_ring.semi_ring'(sK0),X0) | s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) != s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) != s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),X0))) | s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK2))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),X0)))) ) | $spl8),
  inference(superposition,[],[f41,f53])).
fof(f41,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : ((s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))))) & (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : ((s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))))) & ((s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))) | s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2)))),
  inference(nnf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2] : (s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) <=> (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ! [X0] : ! [X1,X2] : (s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) <=> (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X8,X7] : (s('type.semi_ring.semi_ring'(X6),X7) = s('type.semi_ring.semi_ring'(X6),X8) <=> (s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X8)))))),
  file('thm.semi_ring.semi_ring_component_equality','thm.semi_ring.semi_ring_47_prove_ltb')).
fof(f8227,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f8226])).
fof(f8226,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),sK1))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f48,f45])).
% SZS output end Proof for thm.semi_ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 56
% Passive clauses: 1507
% Generated clauses: 7952
% Final active clauses: 45
% Final passive clauses: 1386
% Input formulas: 8
% Initial clauses: 16
% 
% Duplicate literals: 610
% Trivial inequalities: 96
% Fw subsumption resolutions: 843
% Fw demodulations: 372
% 
% Simple tautologies: 176
% Equational tautologies: 92
% Forward subsumptions: 3246
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 25
% Forward superposition: 3123
% Backward superposition: 2244
% Self superposition: 298
% Equality factoring: 237
% Equality resolution: 13
% 
% Split clauses: 400
% Split components: 808
% SAT solver clauses: 1051
% SAT solver unit clauses: 274
% SAT solver binary clauses: 318
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 400
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1045
% TWLsolver calls for satisfiability: 39
% 
% Memory used [KB]: 4221
% Time elapsed: 0.256 s
% ------------------------------
% ------------------------------
% Success in time 0.26 s
