Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.while
% (79693)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 932
% Passive clauses: 9312
% Generated clauses: 25037
% Final active clauses: 923
% Final passive clauses: 8110
% Input formulas: 69
% Initial clauses: 247
% 
% Duplicate literals: 410
% Trivial inequalities: 77
% Fw subsumption resolutions: 2048
% Fw demodulations: 2105
% Bw demodulations: 8
% 
% Simple tautologies: 310
% Equational tautologies: 41
% Forward subsumptions: 9451
% 
% Binary resolution: 8789
% Factoring: 677
% Forward superposition: 4598
% Backward superposition: 5847
% Self superposition: 52
% Equality factoring: 28
% Equality resolution: 22
% 
% Split clauses: 426
% Split components: 854
% SAT solver clauses: 616
% SAT solver unit clauses: 21
% SAT solver binary clauses: 127
% 
% Sat splits: 426
% 
% TWLsolver clauses: 616
% TWLsolver calls for satisfiability: 359
% 
% Memory used [KB]: 13048
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.while
% (79724)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 503
% Passive clauses: 3109
% Generated clauses: 13625
% Final active clauses: 497
% Final passive clauses: 2598
% Input formulas: 69
% Initial clauses: 247
% 
% Duplicate literals: 261
% Trivial inequalities: 58
% Fw subsumption resolutions: 818
% Fw demodulations: 832
% Bw demodulations: 14
% Condensations: 332
% Global subsumptions: 20
% 
% Simple tautologies: 319
% Equational tautologies: 19
% Forward subsumptions: 4295
% 
% Binary resolution: 3098
% Factoring: 268
% Forward superposition: 1769
% Backward superposition: 2697
% Self superposition: 57
% Equality factoring: 26
% Equality resolution: 12
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 3854
% SAT solver unit clauses: 269
% SAT solver binary clauses: 109
% SAT solver learnt clauses: 40
% SAT solver learnt literals: 12
% 
% Sat splits: 5
% 
% TWLsolver clauses: 3774
% TWLsolver calls for satisfiability: 2188
% 
% Memory used [KB]: 16630
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.while
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.while
% SZS output start Proof for thm.while
fof(f8826,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f8824])).
fof(f8824,plain,(
  s('type.option.option'('type.option.option'('type.option.option'('type.num.num'))),'const.option.SOME_1'(s('type.option.option'('type.option.option'('type.num.num')),'const.option.SOME_1'(s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))))))) != s('type.option.option'('type.option.option'('type.option.option'('type.num.num'))),'const.option.SOME_1'(s('type.option.option'('type.option.option'('type.num.num')),'const.option.SOME_1'(s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13)))))))),
  inference(backward_demodulation,[],[f8571,f7592])).
fof(f7592,plain,(
  s('type.option.option'('type.option.option'('type.option.option'('type.num.num'))),'const.option.SOME_1'(s('type.option.option'('type.option.option'('type.num.num')),'const.option.SOME_1'(s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))))))) != s('type.option.option'('type.option.option'('type.option.option'('type.num.num'))),'const.option.SOME_1'(s('type.option.option'('type.option.option'('type.num.num')),'const.option.SOME_1'(s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK13)))))))))),
  inference(unit_resulting_resolution,[],[f1918,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X22] : ! [X6,X23] : (s('type.option.option'(X22),'const.option.SOME_1'(s(X22,X6))) = s('type.option.option'(X22),'const.option.SOME_1'(s(X22,X23))) <=> s(X22,X6) = s(X22,X23))),
  file('thm.while.OLEAST_EQNS|split|1','thm.option.SOME_|49||49|')).
fof(f1918,plain,(
  s('type.option.option'('type.option.option'('type.num.num')),'const.option.SOME_1'(s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))))) != s('type.option.option'('type.option.option'('type.num.num')),'const.option.SOME_1'(s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK13)))))))),
  inference(unit_resulting_resolution,[],[f344,f427])).
fof(f344,plain,(
  s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))) != s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK13)))))),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))) != s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK13)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f172])).
fof(f172,plain,(
  ? [X0] : s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))) != s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X6))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X6))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X6)))))),
  file('thm.while.OLEAST_EQNS|split|1','thm.while.OLEAST_EQNS|split|1')).
fof(f8571,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f8569,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X22] : ! [X27,X28] : s(X22,X27) = s(X22,'const.bool.COND_3'(s(bool,'T_0'),s(X22,X27),s(X22,X28)))),
  file('thm.while.OLEAST_EQNS|split|1','thm.bool.COND_CLAUSES|split|0')).
fof(f8569,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0))))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))),s('type.option.option'('type.num.num'),'const.option.NONE_0')))) )),
  inference(backward_demodulation,[],[f8111,f2664])).
fof(f2664,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0))))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK20('=_1'(s('type.num.num',X0)))),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))),s('type.option.option'('type.num.num'),'const.option.NONE_0')))) )),
  inference(superposition,[],[f437,f352])).
fof(f352,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X6)))))),
  file('thm.while.OLEAST_EQNS|split|1','thm.while.LEAST_EQ|split|1')).
fof(f437,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK20(X0)),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0] : ((~p(s(bool,sK20(X0))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK21(X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,sK20(X0)))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK20(X0)),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21])],[f279])).
fof(f279,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(rectify,[],[f278])).
fof(f278,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1)))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ? [X8] : ((p(s(bool,X8)) <=> ? [X9] : p(s(bool,happ(s(fun('type.num.num',bool),X7),s('type.num.num',X9))))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X7))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X8),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X7))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  file('thm.while.OLEAST_EQNS|split|1','thm.while.OLEAST_def')).
fof(f8111,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK20('=_1'(s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f345,f7605,f450])).
fof(f450,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X27,X28] : ((p(s(bool,X27)) => p(s(bool,X28))) => ((p(s(bool,X28)) => p(s(bool,X27))) => s(bool,X27) = s(bool,X28)))),
  file('thm.while.OLEAST_EQNS|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f7605,plain,(
  ( ! [X0] : (p(s(bool,sK20('=_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f2100,f436])).
fof(f436,plain,(
  ( ! [X2,X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,sK20(X0)))) )),
  inference(cnf_transformation,[],[f280])).
fof(f2100,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0))),s('type.num.num',X0))))) )),
  inference(forward_demodulation,[],[f2078,f352])).
fof(f2078,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0))))))),s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f352,f488])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13,X14] : (p(s(bool,happ(s(fun(X12,bool),'=_1'(s(X12,X13))),s(X12,X14)))) <=> s(X12,X13) = s(X12,X14))),
  file('thm.while.OLEAST_EQNS|split|1',arity_1_HL_EQ)).
fof(f345,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.while.OLEAST_EQNS|split|1',aHL_TRUTH)).
% SZS output end Proof for thm.while
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 447
% Passive clauses: 2532
% Generated clauses: 8228
% Final active clauses: 433
% Final passive clauses: 1806
% Input formulas: 69
% Initial clauses: 247
% 
% Duplicate literals: 9
% Trivial inequalities: 6
% Fw subsumption resolutions: 75
% Fw demodulations: 1454
% Bw demodulations: 293
% Condensations: 20
% 
% Simple tautologies: 63
% Equational tautologies: 28
% Forward subsumptions: 3995
% Fw demodulations to eq. taut.: 28
% 
% Unit resulting resolution: 4580
% Factoring: 1
% Forward superposition: 922
% Backward superposition: 836
% Self superposition: 28
% Equality factoring: 2
% Equality resolution: 7
% 
% SAT solver clauses: 2534
% SAT solver unit clauses: 1786
% SAT solver binary clauses: 366
% 
% Memory used [KB]: 8315
% Time elapsed: 0.272 s
% ------------------------------
% ------------------------------
% Success in time 1.61 s
