Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f10147,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f10146,f10133,f112,f174,f172,f9905,f115,f176,f178,f119,f179,f9858,f111,f173,f175])).
fof(f175,plain,(
  ( ! [X4] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X4)))) ) | $spl4),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  ( ! [X4] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X4)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f173,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2)))) | $spl2),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f111,plain,(
  ( ! [X4] : (p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X4)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  (! [X4] : s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X4))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2))))) & (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK3))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f80])).
fof(f80,plain,(
  ? [X0,X1,X2] : ((! [X4] : s('type.sum.sum'(X0,X1),X2) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),X2))))) & (? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),X2))))))),
  inference(rectify,[],[f79])).
fof(f79,plain,(
  ? [X0,X1,X2] : ((! [X3] : s('type.sum.sum'(X0,X1),X2) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),X2))))) & (? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),X2))))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2] : (! [X3] : s('type.sum.sum'(X0,X1),X2) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <~> p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),X2)))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ~! [X0,X1,X2] : (! [X3] : s('type.sum.sum'(X0,X1),X2) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <=> p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),X2)))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1] : ! [X2] : (! [X3] : ~s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <=> p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : (! [X9] : ~s('type.sum.sum'(X6,X7),X8) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X9))) <=> p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X6,X7),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : (! [X9] : ~s('type.sum.sum'(X6,X7),X8) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X9))) <=> p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X6,X7),X8)))))),
  file('thm.quantHeuristics.INR_NEQ_ELIM|split|0','thm.quantHeuristics.INR_NEQ_ELIM|split|0')).
fof(f9858,plain,(
  $false | ($spl4 | $spl27)),
  inference(subsumption_resolution,[],[f9840,f2163])).
fof(f2163,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2))) | $spl27),
  inference(cnf_transformation,[],[f2163_D])).
fof(f2163_D,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f9840,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2))) | $spl4),
  inference(superposition,[],[f196,f9836])).
fof(f9836,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK5(sK2,sK1,sK0)))) | $spl4),
  inference(equality_resolution,[],[f3313])).
fof(f3313,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK5(X0,sK1,sK0))))) ) | $spl4),
  inference(superposition,[],[f175,f162])).
fof(f162,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK6(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK5(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK5(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK6(X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : (? [X4] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))) | ? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ! [X2] : (? [X4] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))) | ? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7,X6] : ! [X16] : (? [X8] : s('type.sum.sum'(X7,X6),X16) = s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X8))) | ? [X13] : s('type.sum.sum'(X7,X6),X16) = s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X13))))),
  file('thm.quantHeuristics.INR_NEQ_ELIM|split|0','thm.sum.sum_CASES')).
fof(f196,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X3,X4),'const.sum.INL_1'(s(X3,X5)))))) )),
  inference(resolution,[],[f161,f123])).
fof(f123,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.quantHeuristics.INR_NEQ_ELIM|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f161,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ! [X2] : p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X7] : ! [X8] : p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X8))))))),
  file('thm.quantHeuristics.INR_NEQ_ELIM|split|0','thm.sum.ISL|split|0')).
fof(f179,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f119,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quantHeuristics.INR_NEQ_ELIM|split|0',aHL_FALSITY)).
fof(f178,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f176,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f115,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quantHeuristics.INR_NEQ_ELIM|split|0','thm.bool.NOT_CLAUSES|split|2')).
fof(f9905,plain,(
  $false | ($spl2 | $spl27)),
  inference(subsumption_resolution,[],[f9877,f2163])).
fof(f9877,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2))) | $spl2),
  inference(resolution,[],[f173,f123])).
fof(f172,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK3))) | $spl0),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f174,plain,(
  ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2)))) | $spl3),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f112,plain,(
  ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK3)))),
  inference(cnf_transformation,[],[f81])).
fof(f10133,plain,(
  $false | ($spl0 | $spl6 | $spl26)),
  inference(subsumption_resolution,[],[f10132,f186])).
fof(f186,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl6),
  inference(resolution,[],[f126,f176])).
fof(f126,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.quantHeuristics.INR_NEQ_ELIM|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f10132,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl26)),
  inference(forward_demodulation,[],[f10116,f2162])).
fof(f2162,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2))) | $spl26),
  inference(cnf_transformation,[],[f2162_D])).
fof(f2162_D,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f10116,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK1),sK2))) | $spl0),
  inference(superposition,[],[f197,f172])).
fof(f197,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(resolution,[],[f163,f127])).
fof(f127,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f86])).
fof(f163,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ! [X2] : ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X7,X6] : ! [X13] : ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X13))))))),
  file('thm.quantHeuristics.INR_NEQ_ELIM|split|0','thm.sum.ISL|split|1')).
fof(f10146,plain,(
  $false | ($spl3 | $spl6 | $spl26)),
  inference(subsumption_resolution,[],[f10145,f176])).
fof(f10145,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl26)),
  inference(forward_demodulation,[],[f174,f2162])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 121
% Passive clauses: 1841
% Generated clauses: 9867
% Final active clauses: 94
% Final passive clauses: 1451
% Input formulas: 30
% Initial clauses: 59
% 
% Duplicate literals: 332
% Trivial inequalities: 193
% Fw subsumption resolutions: 1434
% Fw demodulations: 1632
% Bw demodulations: 16
% 
% Simple tautologies: 191
% Equational tautologies: 36
% Forward subsumptions: 4097
% 
% Binary resolution: 614
% Factoring: 3
% Forward superposition: 2497
% Backward superposition: 2870
% Self superposition: 77
% Equality factoring: 58
% Equality resolution: 20
% 
% Split clauses: 72
% Split components: 144
% SAT solver clauses: 241
% SAT solver unit clauses: 18
% SAT solver binary clauses: 77
% SAT solver learnt clauses: 4
% SAT solver learnt literals: 4
% 
% Sat splits: 72
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 237
% TWLsolver calls for satisfiability: 47
% 
% Memory used [KB]: 4349
% Time elapsed: 0.192 s
% ------------------------------
% ------------------------------
% Success in time 0.179 s
