Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.while
% (111443)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 218
% Passive clauses: 4470
% Generated clauses: 29049
% Final active clauses: 212
% Final passive clauses: 4059
% Input formulas: 31
% Initial clauses: 81
% 
% Duplicate literals: 765
% Trivial inequalities: 1222
% Fw subsumption resolutions: 3012
% Fw demodulations: 7519
% Bw demodulations: 6
% 
% Simple tautologies: 514
% Equational tautologies: 31
% Forward subsumptions: 10372
% 
% Binary resolution: 477
% Forward superposition: 7857
% Backward superposition: 7795
% Self superposition: 75
% Equality factoring: 55
% Equality resolution: 42
% 
% Split clauses: 606
% Split components: 1245
% SAT solver clauses: 1344
% SAT solver unit clauses: 25
% SAT solver binary clauses: 304
% 
% Sat splits: 606
% 
% TWLsolver clauses: 1333
% TWLsolver calls for satisfiability: 164
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.while
% (111461)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 139
% Passive clauses: 1406
% Generated clauses: 14144
% Final active clauses: 128
% Final passive clauses: 1203
% Input formulas: 31
% Initial clauses: 81
% 
% Duplicate literals: 493
% Trivial inequalities: 352
% Fw subsumption resolutions: 944
% Fw demodulations: 2614
% Bw demodulations: 73
% Condensations: 750
% Global subsumptions: 8
% 
% Simple tautologies: 358
% Equational tautologies: 27
% Forward subsumptions: 4839
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 300
% Forward superposition: 3739
% Backward superposition: 3255
% Self superposition: 54
% Equality factoring: 46
% Equality resolution: 26
% 
% Split clauses: 146
% Split components: 312
% SAT solver clauses: 2118
% SAT solver unit clauses: 76
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 16
% 
% Sat splits: 146
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2086
% TWLsolver calls for satisfiability: 376
% 
% Memory used [KB]: 10746
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.while
% (111479)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 808
% Passive clauses: 12829
% Generated clauses: 389472
% Final active clauses: 203
% Final passive clauses: 1687
% Input formulas: 31
% Initial clauses: 81
% Discarded non-redundant clauses: 4053
% 
% Duplicate literals: 3958
% Trivial inequalities: 4689
% Fw subsumption resolutions: 11549
% Fw demodulations: 188631
% Bw demodulations: 6309
% Condensations: 5025
% Global subsumptions: 503
% 
% Simple tautologies: 1837
% Equational tautologies: 1894
% Forward subsumptions: 157440
% Bw demodulations to eq. taut.: 448
% 
% Unit resulting resolution: 78963
% Forward superposition: 43556
% Backward superposition: 37100
% Self superposition: 210
% Equality factoring: 122
% Equality resolution: 96
% 
% Split clauses: 224
% Split components: 465
% SAT solver clauses: 14393
% SAT solver unit clauses: 890
% SAT solver binary clauses: 3837
% 
% Sat splits: 224
% Sat splitting refutations: 610
% 
% Memory used [KB]: 70105
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.while
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.while
% SZS output start Proof for thm.while
fof(f480,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f252,f296,f253,f298,f254,f297,f295,f259,f260,f261,f262,f263,f302,f301,f264,f266,f277,f299,f368,f370,f371,f300,f373,f304,f375,f303,f386,f305,f406,f451,f450,f437,f472,f479])).
fof(f479,plain,(
  $false | $spl0),
  inference(equality_resolution,[],[f364])).
fof(f364,plain,(
  ( ! [X9] : (s('type.num.num',sK0) != s('type.num.num',X9)) ) | $spl0),
  inference(resolution,[],[f295,f160])).
fof(f160,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X23,X24,X25] : (p(s(bool,happ(s(fun(X23,bool),'=_1'(s(X23,X24))),s(X23,X25)))) <=> s(X23,X24) = s(X23,X25))),
  file('thm.while.LEAST_EQ|split|1',arity_1_HL_EQ)).
fof(f472,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f441,f305])).
fof(f441,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl18),
  inference(resolution,[],[f305,f179])).
fof(f179,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X8)))),
  file('thm.while.LEAST_EQ|split|1','thm.bool.IMP_CLAUSES|split|4')).
fof(f437,plain,(
  $false | $spl18),
  inference(resolution,[],[f305,f208])).
fof(f208,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.while.LEAST_EQ|split|1',aHL_FALSITY)).
fof(f450,plain,(
  $false | $spl18),
  inference(resolution,[],[f305,f250])).
fof(f250,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))))))))),
  inference(equality_resolution,[],[f220])).
fof(f220,plain,(
  ( ! [X0] : (s('type.num.num',sK0) != s('type.num.num',X0) | ~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))))))))) )),
  inference(superposition,[],[f130,f159])).
fof(f159,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2) | ~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f130,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f71])).
fof(f71,plain,(
  ? [X0] : s('type.num.num',X0) != s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ~! [X0] : s('type.num.num',X0) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num',X6) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',X6)))))),
  file('thm.while.LEAST_EQ|split|1','thm.while.LEAST_EQ|split|1')).
fof(f451,plain,(
  $false | ($spl0 | $spl18)),
  inference(resolution,[],[f305,f295])).
fof(f406,plain,(
  ( ! [X17,X15] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X17)))) ) | $spl20),
  inference(cnf_transformation,[],[f406_D])).
fof(f406_D,plain,(
  ( ! [X17,X15] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X17)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f305,plain,(
  ( ! [X20] : (p(s(bool,X20))) ) | $spl18),
  inference(cnf_transformation,[],[f305_D])).
fof(f305_D,plain,(
  ( ! [X20] : (p(s(bool,X20))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f386,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X51)))) ) | $spl0),
  inference(resolution,[],[f295,f211])).
fof(f211,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f135])).
fof(f135,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X14,X9,X15] : ((s(bool,X6) = s(bool,X14) & (p(s(bool,X14)) => s(bool,X9) = s(bool,X15))) => ((p(s(bool,X6)) => p(s(bool,X9))) <=> (p(s(bool,X14)) => p(s(bool,X15)))))),
  file('thm.while.LEAST_EQ|split|1','thm.bool.IMP_CONG')).
fof(f303,plain,(
  ( ! [X13,X11] : (~p(s(bool,X11)) | s(bool,X11) != s(bool,X13)) ) | $spl14),
  inference(cnf_transformation,[],[f303_D])).
fof(f303_D,plain,(
  ( ! [X13,X11] : (~p(s(bool,X11)) | s(bool,X11) != s(bool,X13)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f375,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X43))) | s(bool,X41) != s(bool,X44)) ) | $spl0),
  inference(resolution,[],[f295,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f304,plain,(
  ( ! [X22] : (~p(s(bool,X22))) ) | $spl16),
  inference(cnf_transformation,[],[f304_D])).
fof(f304_D,plain,(
  ( ! [X22] : (~p(s(bool,X22))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f373,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X37)))) ) | $spl0),
  inference(resolution,[],[f295,f137])).
fof(f137,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f300,plain,(
  ( ! [X7,X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,X7)) ) | $spl8),
  inference(cnf_transformation,[],[f300_D])).
fof(f300_D,plain,(
  ( ! [X7,X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,X7)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f371,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X29)))) ) | $spl0),
  inference(resolution,[],[f295,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f370,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X24))) | s(bool,X22) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f295,f134])).
fof(f134,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f368,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X17)))) ) | $spl0),
  inference(resolution,[],[f295,f132])).
fof(f132,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f299,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))))) ) | $spl6),
  inference(cnf_transformation,[],[f299_D])).
fof(f299_D,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f277,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))))) )),
  inference(resolution,[],[f250,f211])).
fof(f266,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))))))) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f250,f138])).
fof(f264,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))))) )),
  inference(resolution,[],[f250,f137])).
fof(f301,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))))) ) | $spl10),
  inference(cnf_transformation,[],[f301_D])).
fof(f301_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f302,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,X8) != s(bool,X9) | p(s(bool,X9))) ) | $spl12),
  inference(cnf_transformation,[],[f302_D])).
fof(f302_D,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,X8) != s(bool,X9) | p(s(bool,X9))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f263,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X19) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))))))) | s(bool,X17) != s(bool,X18)) )),
  inference(resolution,[],[f250,f136])).
fof(f262,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))))) )),
  inference(resolution,[],[f250,f136])).
fof(f261,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))))))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f250,f134])).
fof(f260,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X10) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))))))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f250,f132])).
fof(f259,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0)))))))) )),
  inference(resolution,[],[f250,f132])).
fof(f295,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X1))))) ) | $spl0),
  inference(cnf_transformation,[],[f295_D])).
fof(f295_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X1))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f297,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0))))))) | $spl5),
  inference(cnf_transformation,[],[f297_D])).
fof(f297_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f254,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X3))))) )),
  inference(resolution,[],[f250,f158])).
fof(f158,plain,(
  ( ! [X4,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X1)))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK3(X1,X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X4))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : (! [X4] : ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X4)))) | (! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',sK3(X1,X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK3(X1,X0))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK3(X1,X0)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X1)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f76])).
fof(f76,plain,(
  ! [X0,X1] : (! [X4] : ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X4)))) | ? [X2] : (! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X2)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X1)))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ((! [X4] : ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X4)))) | ? [X2] : ((! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X2))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X1)))))))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ((? [X4] : p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X4)))) & ! [X2] : ((! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X2))))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X1)))))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ((? [X4] : p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X4)))) & ! [X2] : ((! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X2))))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X1)))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X19,X20] : ((? [X21] : p(s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',X21)))) & ! [X21] : ((! [X22] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',X21)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',X22))))) & p(s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',X21))))) => p(s(bool,happ(s(fun('type.num.num',bool),X19),s('type.num.num',X21)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X19),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X20)))))))),
  file('thm.while.LEAST_EQ|split|1','thm.while.LEAST_ELIM')).
fof(f298,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0))))))) | $spl4),
  inference(cnf_transformation,[],[f298_D])).
fof(f298_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0))))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f253,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X2))))) )),
  inference(resolution,[],[f250,f157])).
fof(f157,plain,(
  ( ! [X4,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X1)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK3(X1,X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X4))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f296,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0)))))))) ) | $spl2),
  inference(cnf_transformation,[],[f296_D])).
fof(f296_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0)))))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f252,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK3('=_1'(s('type.num.num',sK0)),'=_1'(s('type.num.num',sK0))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'=_1'(s('type.num.num',sK0))),s('type.num.num',X1))))) )),
  inference(resolution,[],[f250,f156])).
fof(f156,plain,(
  ( ! [X4,X0,X3,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X1)))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X3)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',sK3(X1,X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X4))))) )),
  inference(cnf_transformation,[],[f99])).
% SZS output end Proof for thm.while
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 51
% Passive clauses: 89
% Generated clauses: 262
% Final active clauses: 50
% Final passive clauses: 79
% Input formulas: 31
% Initial clauses: 81
% 
% Selected by SInE selection: 31
% SInE iterations: 1
% 
% Duplicate literals: 8
% Trivial inequalities: 1
% Fw subsumption resolutions: 20
% Fw demodulations: 6
% 
% Simple tautologies: 45
% Equational tautologies: 1
% Forward subsumptions: 106
% 
% Binary resolution: 90
% Forward superposition: 128
% Equality resolution: 4
% 
% Split clauses: 17
% Split components: 36
% SAT solver clauses: 23
% SAT solver unit clauses: 4
% SAT solver binary clauses: 11
% 
% Sat splits: 17
% Sat splitting refutations: 5
% 
% Memory used [KB]: 1023
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
% Success in time 11.73 s
