Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f1860,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1852,f103,f110])).
fof(f110,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f103,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f88])).
fof(f88,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7] : (s(X6,X7) = s(X6,X7) <=> p(s(bool,'T_0')))),
  file('thm.list.LAST_CONS|split|1','thm.bool.REFL_CLAUSE')).
fof(f1852,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1851,f94])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X13,X14] : ~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X14),s('type.list.list'(X6),X13)))),
  file('thm.list.LAST_CONS|split|1','thm.list.NOT_CONS_NIL')).
fof(f1851,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))) | $spl6),
  inference(subsumption_resolution,[],[f1805,f110])).
fof(f1805,plain,(
  ~p(s(bool,'T_0')) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))),
  inference(superposition,[],[f95,f1789])).
fof(f1789,plain,(
  s(bool,'T_0') = s(bool,sK4('const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)),sK1,sK0))),
  inference(trivial_inequality_removal,[],[f1788])).
fof(f1788,plain,(
  s(sK0,'const.list.LAST_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))))) != s(sK0,'const.list.LAST_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))))) | s(bool,'T_0') = s(bool,sK4('const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)),sK1,sK0))),
  inference(superposition,[],[f67,f425])).
fof(f425,plain,(
  ( ! [X14,X15,X13] : (s(X15,'const.list.LAST_1'(s('type.list.list'(X15),X13))) = s(X15,'const.list.LAST_1'(s('type.list.list'(X15),'const.list.CONS_2'(s(X15,X14),s('type.list.list'(X15),X13))))) | s(bool,'T_0') = s(bool,sK4(X13,X14,X15))) )),
  inference(forward_demodulation,[],[f388,f92])).
fof(f92,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X10,X11] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X10),s(X6,X11)))),
  file('thm.list.LAST_CONS|split|1','thm.bool.COND_CLAUSES|split|1')).
fof(f388,plain,(
  ( ! [X14,X15,X13] : (s(X15,'const.list.LAST_1'(s('type.list.list'(X15),'const.list.CONS_2'(s(X15,X14),s('type.list.list'(X15),X13))))) = s(X15,'const.bool.COND_3'(s(bool,'F_0'),s(X15,X14),s(X15,'const.list.LAST_1'(s('type.list.list'(X15),X13))))) | s(bool,'T_0') = s(bool,sK4(X13,X14,X15))) )),
  inference(superposition,[],[f97,f71])).
fof(f71,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.LAST_CONS|split|1',aHL_BOOL_CASES)).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s(X0,'const.bool.COND_3'(s(bool,sK4(X2,X1,X0)),s(X0,X1),s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,sK4(X2,X1,X0))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | p(s(bool,sK4(X2,X1,X0)))) & s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s(X0,'const.bool.COND_3'(s(bool,sK4(X2,X1,X0)),s(X0,X1),s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : ? [X3] : ((~p(s(bool,X3)) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | p(s(bool,X3))) & s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,X1),s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X2))))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : ? [X3] : (((~p(s(bool,X3)) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | p(s(bool,X3)))) & s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,X1),s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2] : ? [X3] : ((p(s(bool,X3)) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,X1),s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X2))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1,X2] : ? [X3] : ((p(s(bool,X3)) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) & s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))) = s(X0,'const.bool.COND_3'(s(bool,X3),s(X0,X1),s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X15,X12] : ? [X16] : ((p(s(bool,X16)) <=> s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X12)) & s(X6,'const.list.LAST_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X15),s('type.list.list'(X6),X12))))) = s(X6,'const.bool.COND_3'(s(bool,X16),s(X6,X15),s(X6,'const.list.LAST_1'(s('type.list.list'(X6),X12))))))),
  file('thm.list.LAST_CONS|split|1','thm.list.LAST_DEF')).
fof(f67,plain,(
  s(sK0,'const.list.LAST_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))))) != s(sK0,'const.list.LAST_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK1),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))))))),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  s(sK0,'const.list.LAST_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))))) != s(sK0,'const.list.LAST_1'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK1),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f44])).
fof(f44,plain,(
  ? [X0,X1,X2,X3] : s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) != s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3] : s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s(X0,'const.list.LAST_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,'const.list.LAST_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9))))) = s(X6,'const.list.LAST_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,'const.list.LAST_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9))))) = s(X6,'const.list.LAST_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))))))),
  file('thm.list.LAST_CONS|split|1','thm.list.LAST_CONS|split|1')).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK4(X2,X1,X0))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f65])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 37
% Passive clauses: 485
% Generated clauses: 1729
% Final active clauses: 34
% Final passive clauses: 422
% Input formulas: 19
% Initial clauses: 36
% 
% Duplicate literals: 132
% Trivial inequalities: 15
% Fw subsumption resolutions: 129
% Fw demodulations: 79
% Bw demodulations: 1
% 
% Simple tautologies: 48
% Equational tautologies: 30
% Forward subsumptions: 803
% 
% Binary resolution: 41
% Factoring: 1
% Forward superposition: 519
% Backward superposition: 646
% Self superposition: 66
% Equality factoring: 52
% Equality resolution: 3
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 31
% SAT solver unit clauses: 5
% SAT solver binary clauses: 15
% SAT solver learnt clauses: 1
% 
% Sat splits: 12
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 30
% TWLsolver calls for satisfiability: 15
% 
% Memory used [KB]: 1791
% Time elapsed: 0.051 s
% ------------------------------
% ------------------------------
% Success in time 0.056 s
