Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (58629)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1087
% Passive clauses: 10015
% Generated clauses: 27179
% Final active clauses: 980
% Final passive clauses: 8173
% Input formulas: 49
% Initial clauses: 209
% 
% Duplicate literals: 381
% Trivial inequalities: 136
% Fw subsumption resolutions: 3261
% Fw demodulations: 1719
% Bw demodulations: 86
% 
% Simple tautologies: 257
% Equational tautologies: 38
% Forward subsumptions: 10512
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 10309
% Factoring: 732
% Forward superposition: 3105
% Backward superposition: 6810
% Self superposition: 29
% Equality factoring: 26
% Equality resolution: 12
% 
% Split clauses: 855
% Split components: 1713
% SAT solver clauses: 1276
% SAT solver unit clauses: 10
% SAT solver binary clauses: 188
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 855
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1273
% TWLsolver calls for satisfiability: 517
% 
% Memory used [KB]: 11641
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (58649)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 506
% Passive clauses: 2585
% Generated clauses: 16217
% Final active clauses: 445
% Final passive clauses: 2002
% Input formulas: 49
% Initial clauses: 209
% 
% Duplicate literals: 307
% Trivial inequalities: 82
% Fw subsumption resolutions: 1460
% Fw demodulations: 1836
% Bw demodulations: 87
% Condensations: 401
% Global subsumptions: 13
% 
% Simple tautologies: 312
% Equational tautologies: 28
% Forward subsumptions: 5758
% 
% Binary resolution: 3746
% Factoring: 122
% Forward superposition: 2273
% Backward superposition: 2820
% Self superposition: 41
% Equality factoring: 34
% Equality resolution: 9
% 
% Split clauses: 111
% Split components: 224
% SAT solver clauses: 3514
% SAT solver unit clauses: 196
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 12
% 
% Sat splits: 111
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 3460
% TWLsolver calls for satisfiability: 2013
% 
% Memory used [KB]: 14072
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (58670)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3234
% Passive clauses: 18673
% Generated clauses: 101875
% Final active clauses: 2775
% Final passive clauses: 6941
% Input formulas: 49
% Initial clauses: 209
% Discarded non-redundant clauses: 6587
% 
% Duplicate literals: 707
% Trivial inequalities: 54
% Fw subsumption resolutions: 1698
% Fw demodulations: 21849
% Bw demodulations: 2510
% Condensations: 629
% Global subsumptions: 35
% 
% Simple tautologies: 105
% Equational tautologies: 267
% Forward subsumptions: 55545
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 61823
% Factoring: 3
% Forward superposition: 6509
% Backward superposition: 5467
% Self superposition: 33
% Equality factoring: 32
% Equality resolution: 2
% 
% Split clauses: 153
% Split components: 313
% SAT solver clauses: 20169
% SAT solver unit clauses: 833
% SAT solver binary clauses: 6141
% 
% Sat splits: 153
% Sat splitting refutations: 1152
% 
% Memory used [KB]: 52962
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 183 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f1618,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f486,f490,f491,f493,f488,f276,f494,f496,f608,f1137,f1347,f618,f487,f1473,f1474,f1604,f492,f610,f1616,f1617])).
fof(f1617,plain,(
  $false | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f1138,f487])).
fof(f1138,plain,(
  ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl11),
  inference(resolution,[],[f610,f309])).
fof(f309,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))))) & (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))))) & ((~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1))))) | p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X6,fun(X6,bool)),X7)))) & p(s(bool,'const.relation.trichotomous_1'(s(fun(X6,fun(X6,bool)),X7))))))),
  file('thm.relation.WeakLinearOrder_dichotomy','thm.relation.WeakLinearOrder')).
fof(f1616,plain,(
  $false | ($spl3 | $spl12)),
  inference(subsumption_resolution,[],[f1604,f1601])).
fof(f1601,plain,(
  s(sK12,sK14) != s(sK12,sK15) | $spl3),
  inference(subsumption_resolution,[],[f1471,f495])).
fof(f495,plain,(
  p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(subsumption_resolution,[],[f275,f308])).
fof(f308,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f199])).
fof(f275,plain,(
  p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | (p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) & ! [X4,X5] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X4))),s(sK12,X5)))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X5))),s(sK12,X4))))))) & (~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | (~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) & ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f168])).
fof(f168,plain,(
  ? [X0,X1] : ((p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ! [X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4))))))) & (~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ? [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(rectify,[],[f167])).
fof(f167,plain,(
  ? [X0,X1] : ((p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))))) & (~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ? [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ? [X0,X1] : ((p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))))) & (~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ? [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))))))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ? [X0,X1] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <~> (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ~! [X0,X1] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X6,fun(X6,bool)),X7)))) & ! [X8,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X6,fun(X6,bool)),X7)))) & ! [X8,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8)))))))),
  file('thm.relation.WeakLinearOrder_dichotomy','thm.relation.WeakLinearOrder_dichotomy')).
fof(f1471,plain,(
  s(sK12,sK14) != s(sK12,sK15) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl3),
  inference(resolution,[],[f490,f305])).
fof(f305,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : ((s(X0,X2) != s(X0,X3) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : ((s(X0,X2) != s(X0,X3) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))) | s(X0,X2) = s(X0,X3))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : (s(X0,X2) = s(X0,X3) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) => ! [X2,X3] : (s(X0,X2) = s(X0,X3) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) => ! [X2,X3] : (s(X0,X2) = s(X0,X3) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.WeakOrder_1'(s(fun(X6,fun(X6,bool)),X7)))) => ! [X14,X19] : (s(X6,X14) = s(X6,X19) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X14))),s(X6,X19)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X19))),s(X6,X14)))))))),
  file('thm.relation.WeakLinearOrder_dichotomy','thm.relation.WeakOrder_EQ')).
fof(f610,plain,(
  ~p(s(bool,'const.relation.trichotomous_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl11),
  inference(cnf_transformation,[],[f610_D])).
fof(f610_D,plain,(
  ~p(s(bool,'const.relation.trichotomous_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f492,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | $spl4),
  inference(cnf_transformation,[],[f492_D])).
fof(f492_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1604,plain,(
  s(sK12,sK14) = s(sK12,sK15) | $spl12),
  inference(cnf_transformation,[],[f1604_D])).
fof(f1604_D,plain,(
  s(sK12,sK14) = s(sK12,sK15) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1474,plain,(
  s(sK12,sK14) = s(sK12,sK15) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,'const.relation.trichotomous_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl3),
  inference(resolution,[],[f490,f301])).
fof(f301,plain,(
  ( ! [X4,X0,X5,X1] : (s(X0,X4) = s(X0,X5) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK29(X1,X0)))),s(X0,sK30(X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK30(X1,X0)))),s(X0,sK29(X1,X0))))) & s(X0,sK29(X1,X0)) != s(X0,sK30(X1,X0))) | p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29,sK30])],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) | s(X0,X4) = s(X0,X5))) & (? [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3))) & (? [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) | s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.trichotomous_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> ! [X8,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8)))) | s(X6,X8) = s(X6,X9)))),
  file('thm.relation.WeakLinearOrder_dichotomy','thm.relation.trichotomous')).
fof(f1473,plain,(
  s(sK12,sK14) = s(sK12,sK15) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,'const.relation.trichotomous_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl3),
  inference(resolution,[],[f490,f301])).
fof(f487,plain,(
  p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl0),
  inference(cnf_transformation,[],[f487_D])).
fof(f487_D,plain,(
  p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f618,plain,(
  p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.trichotomous_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(resolution,[],[f495,f310])).
fof(f310,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f199])).
fof(f1347,plain,(
  $false | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f1041,f610])).
fof(f1041,plain,(
  p(s(bool,'const.relation.trichotomous_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl6),
  inference(subsumption_resolution,[],[f758,f302])).
fof(f302,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK29(X1,X0)))),s(X0,sK30(X1,X0)))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f758,plain,(
  p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK29(sK13,sK12)))),s(sK12,sK30(sK13,sK12))))) | p(s(bool,'const.relation.trichotomous_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl6),
  inference(resolution,[],[f494,f303])).
fof(f303,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.trichotomous_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK30(X1,X0)))),s(X0,sK29(X1,X0)))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f1137,plain,(
  $false | $spl9),
  inference(subsumption_resolution,[],[f608,f495])).
fof(f608,plain,(
  ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl9),
  inference(cnf_transformation,[],[f608_D])).
fof(f608_D,plain,(
  ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f496,plain,(
  ~p(s(bool,'const.relation.trichotomous_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl1),
  inference(resolution,[],[f488,f310])).
fof(f494,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X5))),s(sK12,X4)))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X4))),s(sK12,X5))))) ) | $spl6),
  inference(cnf_transformation,[],[f494_D])).
fof(f494_D,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X5))),s(sK12,X4)))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X4))),s(sK12,X5))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f276,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X5))),s(sK12,X4)))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,X4))),s(sK12,X5)))) | p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13))))) )),
  inference(cnf_transformation,[],[f169])).
fof(f488,plain,(
  ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | $spl1),
  inference(cnf_transformation,[],[f488_D])).
fof(f488_D,plain,(
  ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f493,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | $spl5),
  inference(cnf_transformation,[],[f493_D])).
fof(f493_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f491,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(subsumption_resolution,[],[f277,f308])).
fof(f277,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK14))),s(sK12,sK15)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(cnf_transformation,[],[f169])).
fof(f490,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | $spl3),
  inference(cnf_transformation,[],[f490_D])).
fof(f490_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f486,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(subsumption_resolution,[],[f278,f308])).
fof(f278,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun(sK12,fun(sK12,bool)),sK13),s(sK12,sK15))),s(sK12,sK14)))) | ~p(s(bool,'const.relation.WeakOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13)))) | ~p(s(bool,'const.relation.WeakLinearOrder_1'(s(fun(sK12,fun(sK12,bool)),sK13))))),
  inference(cnf_transformation,[],[f169])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 170
% Passive clauses: 878
% Generated clauses: 1148
% Final active clauses: 158
% Final passive clauses: 365
% Input formulas: 49
% Initial clauses: 206
% 
% Selected by SInE selection: 46
% SInE iterations: 1
% 
% Duplicate literals: 18
% Trivial inequalities: 1
% Fw subsumption resolutions: 76
% 
% Simple tautologies: 77
% Equational tautologies: 1
% Forward subsumptions: 157
% 
% Binary resolution: 750
% Factoring: 1
% Forward superposition: 278
% 
% Split clauses: 7
% Split components: 16
% SAT solver clauses: 12
% SAT solver unit clauses: 1
% SAT solver binary clauses: 7
% 
% Sat splits: 7
% Sat splitting refutations: 4
% 
% Memory used [KB]: 1791
% Time elapsed: 0.051 s
% ------------------------------
% ------------------------------
% Success in time 11.723 s
