Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (115846)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 778
% Passive clauses: 9563
% Generated clauses: 19822
% Final active clauses: 747
% Final passive clauses: 8348
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 263
% Trivial inequalities: 39
% Fw subsumption resolutions: 1212
% Fw demodulations: 1448
% Bw demodulations: 30
% 
% Simple tautologies: 236
% Equational tautologies: 62
% Forward subsumptions: 6799
% 
% Binary resolution: 10954
% Factoring: 1461
% Forward superposition: 2240
% Backward superposition: 1781
% Self superposition: 23
% Equality factoring: 4
% Equality resolution: 19
% 
% Split clauses: 237
% Split components: 477
% SAT solver clauses: 568
% SAT solver unit clauses: 17
% SAT solver binary clauses: 75
% 
% Sat splits: 237
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 567
% TWLsolver calls for satisfiability: 363
% 
% Memory used [KB]: 10874
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (115866)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 431
% Passive clauses: 2426
% Generated clauses: 13524
% Final active clauses: 427
% Final passive clauses: 1994
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 258
% Trivial inequalities: 38
% Fw subsumption resolutions: 776
% Fw demodulations: 1030
% Bw demodulations: 5
% Condensations: 546
% Global subsumptions: 23
% 
% Simple tautologies: 287
% Equational tautologies: 17
% Forward subsumptions: 4936
% 
% Binary resolution: 2390
% Factoring: 99
% Forward superposition: 2796
% Backward superposition: 2799
% Self superposition: 44
% Equality factoring: 28
% Equality resolution: 11
% 
% Split clauses: 44
% Split components: 88
% SAT solver clauses: 3241
% SAT solver unit clauses: 265
% SAT solver binary clauses: 110
% SAT solver learnt clauses: 46
% SAT solver learnt literals: 19
% 
% Sat splits: 44
% 
% TWLsolver clauses: 3150
% TWLsolver calls for satisfiability: 2449
% 
% Memory used [KB]: 16247
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% (115883)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1670
% Passive clauses: 8749
% Generated clauses: 311928
% Final active clauses: 1445
% Final passive clauses: 3341
% Input formulas: 62
% Initial clauses: 251
% Discarded non-redundant clauses: 3859
% 
% Duplicate literals: 776
% Trivial inequalities: 70
% Fw subsumption resolutions: 2416
% Fw demodulations: 174288
% Bw demodulations: 1651
% Condensations: 2015
% Global subsumptions: 1522
% 
% Simple tautologies: 137
% Equational tautologies: 83
% Forward subsumptions: 86463
% Fw demodulations to eq. taut.: 30892
% Bw demodulations to eq. taut.: 44
% 
% Unit resulting resolution: 98518
% Factoring: 1
% Forward superposition: 20492
% Backward superposition: 8123
% Self superposition: 72
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 107
% Split components: 219
% SAT solver clauses: 13429
% SAT solver unit clauses: 4810
% SAT solver binary clauses: 4146
% 
% Sat splits: 107
% Sat splitting refutations: 2942
% 
% Memory used [KB]: 174496
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quantHeuristics
% (116157)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2287
% Passive clauses: 32738
% Generated clauses: 246073
% Final active clauses: 2120
% Final passive clauses: 28101
% Input formulas: 62
% Initial clauses: 206
% 
% Selected by SInE selection: 51
% SInE iterations: 1
% 
% Duplicate literals: 7349
% Trivial inequalities: 980
% Fw subsumption resolutions: 34113
% Fw demodulations: 12944
% Bw demodulations: 140
% 
% Simple tautologies: 8902
% Equational tautologies: 1013
% Forward subsumptions: 132039
% 
% Binary resolution: 176657
% Factoring: 214
% Forward superposition: 7107
% Backward superposition: 5401
% Self superposition: 9
% Equality factoring: 2
% Equality resolution: 85
% 
% Split clauses: 17433
% Split components: 38547
% SAT solver clauses: 19074
% SAT solver unit clauses: 76
% SAT solver binary clauses: 1060
% 
% Sat splits: 17433
% Sat splitting refutations: 112
% 
% Memory used [KB]: 64988
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.quantHeuristics
% (116279)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 324
% Passive clauses: 3135
% Generated clauses: 12707
% Final active clauses: 314
% Final passive clauses: 2709
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 182
% Trivial inequalities: 29
% Fw subsumption resolutions: 796
% Fw demodulations: 2679
% Bw demodulations: 39
% Global subsumptions: 27
% 
% Simple tautologies: 232
% Equational tautologies: 19
% Forward subsumptions: 4466
% Fw demodulations to eq. taut.: 876
% 
% Unit resulting resolution: 2717
% Factoring: 2
% Forward superposition: 3135
% Backward superposition: 2648
% Self superposition: 120
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 8
% Split components: 17
% SAT solver clauses: 3275
% SAT solver unit clauses: 381
% SAT solver binary clauses: 430
% 
% Sat splits: 8
% Sat splitting refutations: 103
% 
% Memory used [KB]: 13944
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.quantHeuristics
% (116303)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1172
% Passive clauses: 5521
% Generated clauses: 51098
% Final active clauses: 1151
% Final passive clauses: 4407
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 3194
% Trivial inequalities: 147
% Fw subsumption resolutions: 5026
% Bw subsumption resolutions: 17
% Fw demodulations: 8949
% Bw demodulations: 12
% Condensations: 4084
% Global subsumptions: 2
% 
% Simple tautologies: 3849
% Equational tautologies: 163
% Forward subsumptions: 20435
% Backward subsumptions: 115
% 
% Binary resolution: 20502
% Factoring: 1
% Forward superposition: 5203
% Backward superposition: 4188
% Self superposition: 4
% Equality factoring: 4
% Equality resolution: 9
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 5527
% SAT solver unit clauses: 263
% SAT solver binary clauses: 1290
% 
% Sat splits: 2
% 
% Memory used [KB]: 21875
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f22526,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f343,f600,f342,f602,f341,f604,f914,f961,f908,f963,f906,f905,f904,f967,f903,f965,f902,f901,f969,f899,f898,f897,f973,f896,f971,f895,f894,f975,f892,f891,f890,f979,f889,f977,f888,f887,f981,f874,f873,f864,f863,f862,f853,f852,f851,f842,f841,f840,f833,f832,f831,f826,f825,f816,f815,f814,f805,f804,f803,f1003,f790,f789,f784,f783,f781,f780,f779,f1065,f778,f1063,f777,f776,f1067,f768,f1069,f766,f765,f764,f1073,f763,f1071,f762,f761,f1075,f757,f1077,f742,f1079,f739,f738,f737,f1083,f736,f1081,f735,f734,f1085,f731,f730,f729,f728,f727,f726,f725,f724,f723,f1093,f722,f721,f720,f1089,f719,f718,f717,f716,f715,f714,f1097,f713,f1087,f712,f711,f1091,f710,f709,f708,f707,f1095,f706,f705,f1099,f703,f702,f701,f1103,f700,f1101,f699,f698,f1105,f681,f680,f679,f1109,f678,f1107,f677,f676,f1111,f674,f673,f672,f1115,f671,f1113,f670,f669,f1117,f668,f667,f666,f1121,f665,f1119,f664,f663,f1123,f662,f661,f660,f659,f658,f657,f656,f655,f654,f1133,f653,f1131,f652,f651,f650,f1129,f649,f1125,f648,f647,f1127,f641,f640,f638,f637,f609,f608,f960,f607,f606,f917,f918,f1176,f1135,f1195,f1194,f1201,f1193,f1203,f1187,f1186,f1205,f1185,f1207,f1519,f1566,f1513,f1568,f1511,f1510,f1509,f1572,f1508,f1570,f1507,f1506,f1574,f1504,f1503,f1502,f1578,f1501,f1576,f1500,f1499,f1580,f1497,f1496,f1495,f1584,f1494,f1582,f1493,f1492,f1586,f1479,f1478,f1469,f1458,f1447,f1438,f1437,f1436,f1431,f1430,f1421,f1410,f1395,f1394,f1389,f1388,f1386,f1385,f1384,f1668,f1383,f1666,f1382,f1381,f1670,f1373,f1672,f1371,f1370,f1369,f1676,f1368,f1674,f1367,f1366,f1678,f1362,f1680,f1347,f1682,f1344,f1343,f1342,f1686,f1341,f1684,f1340,f1339,f1688,f1336,f1335,f1334,f1333,f1332,f1331,f1330,f1329,f1328,f1696,f1327,f1326,f1325,f1692,f1324,f1323,f1322,f1321,f1320,f1319,f1700,f1318,f1690,f1317,f1316,f1694,f1315,f1314,f1313,f1312,f1698,f1311,f1310,f1702,f1308,f1307,f1306,f1706,f1305,f1704,f1304,f1303,f1708,f1300,f1710,f1299,f1284,f1283,f1282,f1716,f1281,f1714,f1280,f1279,f1718,f1277,f1276,f1275,f1722,f1274,f1720,f1273,f1272,f1724,f1271,f1270,f1269,f1728,f1268,f1726,f1267,f1266,f1712,f1265,f1264,f1263,f1262,f1261,f1260,f1259,f1258,f1257,f1738,f1256,f1736,f1255,f1254,f1253,f1734,f1252,f1730,f1251,f1250,f1732,f1244,f1243,f1241,f1240,f1212,f1211,f1565,f1210,f1209,f1522,f1523,f1781,f1740,f2191,f2190,f2183,f2182,f2158,f2157,f2156,f2155,f2264,f2154,f2153,f2262,f2152,f2151,f2266,f2148,f2147,f2146,f2145,f2268,f2142,f2141,f2140,f2139,f2270,f2136,f2135,f2134,f2133,f2272,f2124,f2123,f2122,f2121,f2120,f2119,f2118,f2117,f2100,f2099,f2096,f2095,f2084,f2083,f2082,f2081,f2072,f2070,f2060,f2058,f2048,f2047,f2046,f2045,f2044,f2043,f2042,f2041,f2032,f2031,f2030,f2029,f2028,f2027,f2026,f2025,f2016,f2014,f2004,f2002,f1992,f1991,f1990,f1989,f1988,f1987,f1986,f1985,f1978,f1977,f1976,f1975,f2362,f1972,f1971,f1970,f1969,f2366,f1968,f1967,f1966,f2364,f1965,f2368,f1962,f1961,f1960,f1959,f2370,f1948,f1947,f1946,f1945,f2374,f1944,f1943,f1942,f2372,f1941,f2376,f1940,f1939,f1938,f1937,f2384,f1936,f1935,f1934,f2380,f1933,f2388,f1932,f2378,f1931,f2386,f1930,f2382,f1929,f2390,f1927,f1926,f1925,f1924,f2394,f1923,f1922,f2392,f1921,f1920,f2396,f1915,f1914,f1913,f1912,f2398,f1897,f1896,f1895,f1894,f2400,f1879,f2402,f1878,f2404,f1877,f1876,f1875,f1874,f2412,f1873,f1872,f2408,f1871,f1870,f2416,f1869,f2406,f1868,f2410,f1867,f2414,f1866,f2418,f1863,f1862,f1861,f1860,f2420,f1857,f1856,f1855,f1854,f2422,f1853,f1852,f1851,f1850,f2424,f1849,f1848,f1847,f1846,f1845,f1844,f1843,f1842,f1841,f1840,f2434,f1839,f1838,f2432,f1837,f1836,f2430,f1835,f2426,f1834,f2428,f1833,f1832,f1831,f1830,f1829,f1828,f2438,f1827,f2436,f1826,f2440,f1824,f1822,f1793,f1792,f1791,f1790,f2241,f1787,f2484,f2200,f2201,f2442,f2600,f2599,f2598,f2603,f2612,f2604,f2610,f2605,f2608,f2615,f2614,f2618,f2823,f2822,f2821,f2825,f2835,f2826,f2833,f2827,f2831,f3002,f3008,f3010,f3066,f3072,f3077,f3094,f3100,f3105,f3128,f3132,f3138,f3142,f3148,f3152,f3163,f3167,f3180,f3179,f3186,f3195,f3194,f3198,f3206,f3205,f3209,f3634,f3640,f3642,f3694,f3700,f3705,f3722,f3728,f3733,f3756,f3760,f3766,f3770,f3776,f3780,f4452,f4451,f5484,f3971,f3970,f6381,f7572,f7571,f8983,f6930,f6929,f10283,f11730,f10872,f10871,f11380,f11736,f11381,f11734,f11382,f11741,f11740,f11742,f11732,f11749,f11748,f12291,f12642,f12292,f12640,f12293,f12645,f12644,f13213,f13258,f13212,f13260,f13208,f13262,f13207,f13264,f13203,f13295,f13171,f13301,f13167,f13166,f13165,f13307,f13162,f13305,f13161,f13160,f13311,f13156,f13155,f13154,f13317,f13151,f13315,f13150,f13149,f13321,f13145,f13144,f13143,f13327,f13140,f13325,f13139,f13138,f13331,f13114,f13113,f13112,f13111,f13110,f13109,f13108,f13096,f13082,f13068,f13054,f13042,f13041,f13040,f13033,f13032,f13020,f13006,f12992,f12966,f12965,f12964,f12963,f12962,f12961,f12960,f12953,f12952,f12946,f12945,f12944,f13483,f12943,f13481,f12942,f12941,f13485,f12931,f13489,f12927,f12926,f12925,f13495,f12922,f13493,f12921,f12920,f13499,f12890,f13509,f12867,f13515,f12852,f12851,f12850,f13523,f12849,f13521,f12848,f12847,f13525,f12844,f12843,f12842,f12841,f12840,f12839,f12838,f12837,f12836,f13533,f12835,f12834,f12833,f13529,f12832,f12831,f12830,f12829,f12828,f12827,f13537,f12826,f13527,f12825,f12824,f13531,f12823,f12822,f12821,f12820,f13535,f12819,f12818,f13539,f12814,f12813,f12812,f13545,f12809,f13543,f12808,f12807,f13549,f12801,f12800,f12799,f13557,f12797,f12796,f12795,f13561,f12793,f13555,f12792,f12791,f13563,f12789,f13559,f12788,f12787,f13565,f12783,f13567,f12781,f12776,f13571,f12775,f13573,f12769,f13575,f12768,f13577,f12760,f12759,f12758,f13583,f12755,f13581,f12754,f12753,f13587,f12749,f12748,f12747,f13593,f12744,f13591,f12743,f12742,f13597,f12741,f13599,f12740,f13601,f12739,f13603,f12738,f13605,f12737,f12736,f12735,f13609,f12734,f13607,f12733,f12732,f13569,f12731,f12730,f12729,f12728,f12727,f12726,f12725,f12724,f12723,f13619,f12722,f13617,f12721,f12720,f12719,f13615,f12718,f13611,f12717,f12716,f13613,f12710,f12709,f12707,f12706,f12656,f12655,f13269,f12654,f12653,f13217,f13218,f13706,f13621,f14182,f14188,f14195,f14222,f14226,f14234,f14238,f14373,f14379,f14386,f14413,f14417,f14425,f14429,f14553,f14552,f14559,f14551,f14550,f14561,f14585,f14591,f14584,f14593,f14605,f14606,f14615,f14614,f14613,f14618,f14612,f14620,f14779,f14778,f14782,f14777,f14776,f14784,f14898,f14897,f14901,f14915,f14916,f14608,f14918,f15166,f15172,f15174,f15230,f15236,f15238,f15247,f15253,f15255,f15300,f15302,f15307,f15309,f15314,f15316,f15321,f15325,f15330,f15334,f15345,f15353,f15341,f15356,f16432,f16431,f18124,f15618,f15617,f19748,f20450,f20458,f20447,f20460,f20958,f20964,f20967,f20986,f20988,f20994,f20996,f21014,f21016,f21782,f21781,f21779,f21778,f21777,f21776,f21827,f21775,f21825,f21774,f21829,f21773,f21772,f21833,f21771,f21831,f21770,f21835,f21764,f21763,f21839,f21762,f21837,f21761,f21841,f21760,f21843,f21759,f21845,f21757,f21756,f21849,f21755,f21847,f21754,f21851,f21753,f21853,f21752,f21855,f21859,f21739,f21738,f21878,f21730,f21729,f21884,f21888,f21890,f21892,f21894,f21896,f21694,f21693,f21692,f21691,f21918,f21687,f21686,f21916,f21685,f21684,f21923,f21679,f21678,f21677,f21676,f21927,f21671,f21670,f21669,f21668,f21931,f21663,f21662,f21661,f21660,f21935,f21939,f21943,f21947,f21951,f21643,f21642,f21641,f21640,f21639,f21638,f21637,f21636,f21635,f21634,f21633,f21632,f21631,f21630,f21629,f21628,f21967,f21607,f21606,f21603,f21602,f21990,f21994,f21589,f21588,f21587,f21586,f22012,f21573,f21571,f22027,f21557,f21555,f22042,f21541,f21539,f22051,f22055,f22059,f22063,f21521,f21520,f21519,f21518,f21517,f21516,f21515,f21514,f21513,f21512,f21511,f21510,f22079,f22083,f22087,f22091,f21493,f21492,f21491,f21490,f21489,f21488,f21487,f21486,f21485,f21484,f21483,f21482,f21481,f21480,f21479,f21478,f22113,f21465,f21463,f22128,f21449,f21447,f22143,f21433,f21431,f22152,f22156,f22160,f22164,f21861,f21413,f21412,f21411,f21410,f21409,f21408,f21407,f21406,f21405,f21404,f21403,f21402,f21388,f21387,f21386,f21385,f22183,f21379,f21378,f21377,f21376,f22188,f21374,f21373,f21372,f22186,f21371,f22191,f21366,f21365,f21364,f21363,f22195,f21359,f21358,f22199,f21334,f21333,f21332,f21331,f22211,f21327,f21326,f21325,f22209,f21324,f22216,f21323,f21322,f21321,f21320,f22224,f21317,f21316,f21315,f22220,f21314,f22230,f21311,f22218,f21310,f22228,f21309,f22222,f21308,f22234,f21302,f21301,f21300,f21299,f22242,f21294,f21293,f22240,f21292,f21291,f22248,f21287,f21286,f21821,f21282,f21819,f21281,f21823,f21258,f21257,f21256,f21255,f22256,f21224,f21223,f21222,f21221,f22270,f21202,f22276,f21201,f22278,f21200,f22280,f21194,f22282,f21193,f22284,f21192,f22286,f21190,f22288,f21188,f22290,f21187,f22292,f21185,f22294,f21178,f21177,f21176,f21175,f22308,f21168,f21167,f22304,f21166,f21165,f22318,f21158,f22302,f21157,f22306,f21156,f22316,f21155,f22326,f21150,f21149,f21148,f21147,f22330,f21142,f21141,f21140,f21139,f22334,f21138,f21137,f21136,f21135,f22336,f21134,f22340,f21133,f22338,f21132,f21131,f21130,f21129,f22342,f21128,f22344,f21127,f22346,f21126,f21125,f21124,f21123,f22348,f21122,f21121,f21120,f21119,f21118,f21117,f21116,f21115,f21114,f21113,f22358,f21112,f21111,f22356,f21110,f21109,f22354,f21108,f22350,f21107,f22352,f21103,f21102,f21098,f21097,f21093,f21092,f22365,f21088,f22363,f21087,f22370,f21085,f21083,f21026,f21025,f21024,f21023,f21874,f21019,f21784,f22523,f21786,f21787,f21886,f22525])).
fof(f22525,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f21886,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl628),
  inference(cnf_transformation,[],[f21886_D])).
fof(f21886_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl628),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl628])])).
fof(f21787,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(resolution,[],[f2484,f368])).
fof(f368,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.EQ_CLAUSES|split|2')).
fof(f21786,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(resolution,[],[f2484,f366])).
fof(f366,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.EQ_CLAUSES|split|3')).
fof(f22523,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | $spl752),
  inference(cnf_transformation,[],[f22523_D])).
fof(f22523_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) <=> ~$spl752),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl752])])).
fof(f21784,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | $spl307),
  inference(resolution,[],[f2484,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK31(X6,X3,X2,X1,X0))))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK32(X3,X2,X1,X0))))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31,sK32])],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7,X6] : ! [X8,X11] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X11)))) <=> ! [X20] : (~p(s(bool,happ(s(fun(X7,bool),X11),s(X7,X20)))) => ? [X21] : ~p(s(bool,happ(s(fun(X7,bool),X11),s(X7,happ(s(fun(X6,X7),X8),s(X6,X21))))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|1')).
fof(f21019,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f476])).
fof(f21874,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl627),
  inference(cnf_transformation,[],[f21874_D])).
fof(f21874_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl627),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl627])])).
fof(f21023,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f361])).
fof(f361,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.EQ_CLAUSES|split|1')).
fof(f21024,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f363])).
fof(f363,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.EQ_CLAUSES|split|0')).
fof(f21025,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f366])).
fof(f21026,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f368])).
fof(f21083,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f431])).
fof(f431,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.IMP_ANTISYM_AX')).
fof(f352,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7',aHL_FALSITY)).
fof(f21085,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f431])).
fof(f22370,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl751),
  inference(cnf_transformation,[],[f22370_D])).
fof(f22370_D,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl751),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl751])])).
fof(f21087,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f437])).
fof(f437,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f22363,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl747),
  inference(cnf_transformation,[],[f22363_D])).
fof(f22363_D,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl747),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl747])])).
fof(f21088,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f437])).
fof(f22365,plain,(
  ~sP0('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl749),
  inference(cnf_transformation,[],[f22365_D])).
fof(f22365_D,plain,(
  ~sP0('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl749),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl749])])).
fof(f21092,plain,(
  ~sP0('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f437])).
fof(f21093,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f437])).
fof(f21097,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f441])).
fof(f441,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f178,f191])).
fof(f178,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X7,X6] : ((~(p(s(bool,X6)) | p(s(bool,X7))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X7)) => p(s(bool,'F_0')))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.sat.sat_9_prove_ltb')).
fof(f21098,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f441])).
fof(f21102,plain,(
  ~sP0('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f441])).
fof(f21103,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f441])).
fof(f22352,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl739),
  inference(cnf_transformation,[],[f22352_D])).
fof(f22352_D,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl739),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl739])])).
fof(f21107,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f452])).
fof(f452,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f266])).
fof(f266,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f22350,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl737),
  inference(cnf_transformation,[],[f22350_D])).
fof(f22350_D,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl737),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl737])])).
fof(f21108,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f452])).
fof(f22354,plain,(
  ~sP1('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl741),
  inference(cnf_transformation,[],[f22354_D])).
fof(f22354_D,plain,(
  ~sP1('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl741),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl741])])).
fof(f21109,plain,(
  ~sP1('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f452])).
fof(f21110,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f452])).
fof(f22356,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl742),
  inference(cnf_transformation,[],[f22356_D])).
fof(f22356_D,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl742),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl742])])).
fof(f21111,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f454])).
fof(f454,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f345,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.TRUTH')).
fof(f21112,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f454])).
fof(f344,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7',aHL_TRUTH)).
fof(f22358,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl744),
  inference(cnf_transformation,[],[f22358_D])).
fof(f22358_D,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl744),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl744])])).
fof(f21113,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f457])).
fof(f457,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f21114,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f457])).
fof(f21115,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f458])).
fof(f458,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f127,f193])).
fof(f127,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X25,X26] : ((p(s(bool,X26)) <=> ~p(s(bool,X25))) <=> ((p(s(bool,X26)) | p(s(bool,X25))) & (~p(s(bool,X25)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.sat.sat_29_prove_ltb')).
fof(f21116,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f458])).
fof(f21117,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f459])).
fof(f459,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f21118,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f459])).
fof(f21119,plain,(
  ~sP1('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f461])).
fof(f461,plain,(
  ( ! [X0,X1] : (~sP1(X1,X0) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f268])).
fof(f21120,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f461])).
fof(f21121,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f461])).
fof(f21122,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f461])).
fof(f22348,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl735),
  inference(cnf_transformation,[],[f22348_D])).
fof(f22348_D,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl735),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl735])])).
fof(f21123,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f488])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f291])).
fof(f291,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f21124,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f488])).
fof(f21125,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f488])).
fof(f21126,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f488])).
fof(f22346,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl732),
  inference(cnf_transformation,[],[f22346_D])).
fof(f22346_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl732),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl732])])).
fof(f21127,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f683,f1864,f1858,f2484,f490])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f1858,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f494])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f294])).
fof(f294,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f1864,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f498])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f297])).
fof(f297,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f683,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f501])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f22344,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl730),
  inference(cnf_transformation,[],[f22344_D])).
fof(f22344_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl730),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl730])])).
fof(f21128,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f1286,f1864,f1858,f2484,f490])).
fof(f1286,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f501])).
fof(f22342,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl728),
  inference(cnf_transformation,[],[f22342_D])).
fof(f22342_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl728),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl728])])).
fof(f21129,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307 | $spl362)),
  inference(unit_resulting_resolution,[],[f12764,f1864,f1858,f2484,f490])).
fof(f12764,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f501])).
fof(f21130,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f683,f682,f1858,f2484,f490])).
fof(f682,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f497])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f298])).
fof(f21131,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f1286,f1285,f1858,f2484,f490])).
fof(f1285,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f497])).
fof(f21132,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307 | $spl362)),
  inference(unit_resulting_resolution,[],[f12764,f12763,f1858,f2484,f490])).
fof(f12763,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f497])).
fof(f22338,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl724),
  inference(cnf_transformation,[],[f22338_D])).
fof(f22338_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl724),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl724])])).
fof(f21133,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f684,f1864,f675,f2484,f491])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f675,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f493])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f295])).
fof(f684,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f502])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f300])).
fof(f22340,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl726),
  inference(cnf_transformation,[],[f22340_D])).
fof(f22340_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl726),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl726])])).
fof(f21134,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f1287,f1864,f1278,f2484,f491])).
fof(f1278,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f493])).
fof(f1287,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f502])).
fof(f22336,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl722),
  inference(cnf_transformation,[],[f22336_D])).
fof(f22336_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl722),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl722])])).
fof(f21135,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307 | $spl362)),
  inference(unit_resulting_resolution,[],[f12765,f1864,f12752,f2484,f491])).
fof(f12752,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f493])).
fof(f12765,plain,(
  ( ! [X0,X1] : (sP2(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f502])).
fof(f21136,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f1287,f1864,f1858,f2484,f491])).
fof(f21137,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f684,f1864,f1858,f2484,f491])).
fof(f21138,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307 | $spl362)),
  inference(unit_resulting_resolution,[],[f12765,f1864,f1858,f2484,f491])).
fof(f22334,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl721),
  inference(cnf_transformation,[],[f22334_D])).
fof(f22334_D,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl721),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl721])])).
fof(f21139,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f492])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f21140,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f492])).
fof(f21141,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f492])).
fof(f21142,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f492])).
fof(f22330,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl719),
  inference(cnf_transformation,[],[f22330_D])).
fof(f22330_D,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl719),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl719])])).
fof(f21147,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f496])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f21148,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f496])).
fof(f21149,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f496])).
fof(f21150,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f496])).
fof(f22326,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') | $spl717),
  inference(cnf_transformation,[],[f22326_D])).
fof(f22326_D,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') <=> ~$spl717),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl717])])).
fof(f21155,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f500])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f300])).
fof(f22316,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl713),
  inference(cnf_transformation,[],[f22316_D])).
fof(f22316_D,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl713),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl713])])).
fof(f21156,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f500])).
fof(f22306,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl709),
  inference(cnf_transformation,[],[f22306_D])).
fof(f22306_D,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl709),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl709])])).
fof(f21157,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f500])).
fof(f22302,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl705),
  inference(cnf_transformation,[],[f22302_D])).
fof(f22302_D,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl705),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl705])])).
fof(f21158,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f500])).
fof(f22318,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl715),
  inference(cnf_transformation,[],[f22318_D])).
fof(f22318_D,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl715),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl715])])).
fof(f21165,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f500])).
fof(f21166,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f500])).
fof(f22304,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl707),
  inference(cnf_transformation,[],[f22304_D])).
fof(f22304_D,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl707),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl707])])).
fof(f21167,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f500])).
fof(f21168,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f500])).
fof(f22308,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl711),
  inference(cnf_transformation,[],[f22308_D])).
fof(f22308_D,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl711),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl711])])).
fof(f21175,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f500])).
fof(f21176,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f500])).
fof(f21177,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f500])).
fof(f21178,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f500])).
fof(f22294,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','F_0') | $spl702),
  inference(cnf_transformation,[],[f22294_D])).
fof(f22294_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','F_0') <=> ~$spl702),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl702])])).
fof(f21185,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','F_0') | ($spl29 | $spl307)),
  inference(unit_resulting_resolution,[],[f1003,f2484,f504])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f144,f198,f197,f196,f195])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> s(bool,X24) = s(bool,X25)) <=> ((p(s(bool,X26)) | p(s(bool,X25)) | p(s(bool,X24))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X25))) & (p(s(bool,X25)) | ~p(s(bool,X24)) | ~p(s(bool,X26))) & (p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.sat.sat_2|49|_prove_ltb')).
fof(f22292,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl700),
  inference(cnf_transformation,[],[f22292_D])).
fof(f22292_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl700),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl700])])).
fof(f21187,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl211 | $spl307)),
  inference(unit_resulting_resolution,[],[f2241,f2484,f504])).
fof(f22290,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | $spl698),
  inference(cnf_transformation,[],[f22290_D])).
fof(f22290_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') <=> ~$spl698),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl698])])).
fof(f21188,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | ($spl29 | $spl307)),
  inference(unit_resulting_resolution,[],[f1003,f2484,f504])).
fof(f22288,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl696),
  inference(cnf_transformation,[],[f22288_D])).
fof(f22288_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl696),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl696])])).
fof(f21190,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | ($spl211 | $spl307)),
  inference(unit_resulting_resolution,[],[f2241,f2484,f504])).
fof(f22286,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl695),
  inference(cnf_transformation,[],[f22286_D])).
fof(f22286_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl695),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl695])])).
fof(f21192,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f1135,f2484,f507])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f301])).
fof(f22284,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl693),
  inference(cnf_transformation,[],[f22284_D])).
fof(f22284_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl693),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl693])])).
fof(f21193,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f1740,f2484,f507])).
fof(f22282,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl691),
  inference(cnf_transformation,[],[f22282_D])).
fof(f22282_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl691),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl691])])).
fof(f21194,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f2442,f2484,f507])).
fof(f22280,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl689),
  inference(cnf_transformation,[],[f22280_D])).
fof(f22280_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl689),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl689])])).
fof(f21200,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f1135,f2484,f507])).
fof(f22278,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl687),
  inference(cnf_transformation,[],[f22278_D])).
fof(f22278_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl687),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl687])])).
fof(f21201,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f1740,f2484,f507])).
fof(f22276,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl685),
  inference(cnf_transformation,[],[f22276_D])).
fof(f22276_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl685),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl685])])).
fof(f21202,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f2442,f2484,f507])).
fof(f22270,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl683),
  inference(cnf_transformation,[],[f22270_D])).
fof(f22270_D,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl683),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl683])])).
fof(f21221,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f21222,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f515])).
fof(f21223,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f515])).
fof(f21224,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f515])).
fof(f22256,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl681),
  inference(cnf_transformation,[],[f22256_D])).
fof(f22256_D,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') <=> ~$spl681),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl681])])).
fof(f21255,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f524])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f146,f201,f200])).
fof(f201,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) & p(s(bool,X24)))) <=> ((p(s(bool,X26)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X26))) & (p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.sat.sat_23_prove_ltb')).
fof(f21256,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f524])).
fof(f21257,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f524])).
fof(f21258,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f524])).
fof(f21823,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl590),
  inference(cnf_transformation,[],[f21823_D])).
fof(f21823_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl590),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl590])])).
fof(f21281,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f1928,f2484,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f310])).
fof(f310,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f1928,plain,(
  ( ! [X0,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f21819,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl586),
  inference(cnf_transformation,[],[f21819_D])).
fof(f21819_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl586),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl586])])).
fof(f21282,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f1928,f2484,f531])).
fof(f21821,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl588),
  inference(cnf_transformation,[],[f21821_D])).
fof(f21821_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f21286,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f1928,f2484,f531])).
fof(f21287,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f1928,f2484,f531])).
fof(f22248,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl679),
  inference(cnf_transformation,[],[f22248_D])).
fof(f22248_D,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl679),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl679])])).
fof(f21291,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f21292,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f532])).
fof(f22240,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl675),
  inference(cnf_transformation,[],[f22240_D])).
fof(f22240_D,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl675),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl675])])).
fof(f21293,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f532])).
fof(f21294,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f532])).
fof(f22242,plain,(
  ~sP8('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl677),
  inference(cnf_transformation,[],[f22242_D])).
fof(f22242_D,plain,(
  ~sP8('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl677),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl677])])).
fof(f21299,plain,(
  ~sP8('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f532])).
fof(f21300,plain,(
  ~sP8('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f532])).
fof(f21301,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f532])).
fof(f21302,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f532])).
fof(f22234,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl672),
  inference(cnf_transformation,[],[f22234_D])).
fof(f22234_D,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl672),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl672])])).
fof(f21308,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f536])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f148,f204,f203])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) | p(s(bool,X24)))) <=> ((p(s(bool,X26)) | ~p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.sat.sat_25_prove_ltb')).
fof(f22222,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl664),
  inference(cnf_transformation,[],[f22222_D])).
fof(f22222_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl664),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl664])])).
fof(f21309,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f536])).
fof(f22228,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl668),
  inference(cnf_transformation,[],[f22228_D])).
fof(f22228_D,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl668),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl668])])).
fof(f21310,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f536])).
fof(f22218,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl660),
  inference(cnf_transformation,[],[f22218_D])).
fof(f22218_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl660),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl660])])).
fof(f21311,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f536])).
fof(f22230,plain,(
  sP9('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl670),
  inference(cnf_transformation,[],[f22230_D])).
fof(f22230_D,plain,(
  sP9('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl670),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl670])])).
fof(f21314,plain,(
  sP9('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f536])).
fof(f22220,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl662),
  inference(cnf_transformation,[],[f22220_D])).
fof(f22220_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl662),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl662])])).
fof(f21315,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f536])).
fof(f21316,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f536])).
fof(f21317,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f536])).
fof(f22224,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') | $spl666),
  inference(cnf_transformation,[],[f22224_D])).
fof(f22224_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') <=> ~$spl666),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl666])])).
fof(f21320,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f536])).
fof(f21321,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f536])).
fof(f21322,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f536])).
fof(f21323,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f536])).
fof(f22216,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl659),
  inference(cnf_transformation,[],[f22216_D])).
fof(f22216_D,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl659),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl659])])).
fof(f21324,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f2484,f539])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f316])).
fof(f22209,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl655),
  inference(cnf_transformation,[],[f22209_D])).
fof(f22209_D,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl655),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl655])])).
fof(f21325,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f2484,f539])).
fof(f21326,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f2484,f539])).
fof(f21327,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f2484,f539])).
fof(f22211,plain,(
  ~sP9('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl657),
  inference(cnf_transformation,[],[f22211_D])).
fof(f22211_D,plain,(
  ~sP9('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl657),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl657])])).
fof(f21331,plain,(
  ~sP9('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f2484,f539])).
fof(f21332,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f2484,f539])).
fof(f21333,plain,(
  ~sP9('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f2484,f539])).
fof(f21334,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f2484,f539])).
fof(f22199,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl652),
  inference(cnf_transformation,[],[f22199_D])).
fof(f22199_D,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl652),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl652])])).
fof(f21358,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1964,f2484,f548])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f318])).
fof(f318,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1964,plain,(
  ( ! [X0,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f552])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f321])).
fof(f321,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f21359,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f1964,f2484,f548])).
fof(f22195,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl651),
  inference(cnf_transformation,[],[f22195_D])).
fof(f22195_D,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl651),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl651])])).
fof(f21363,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f549])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f21364,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f549])).
fof(f21365,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f549])).
fof(f21366,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f549])).
fof(f22191,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl648),
  inference(cnf_transformation,[],[f22191_D])).
fof(f22191_D,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl648),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl648])])).
fof(f21371,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f2484,f553])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f188,f207,f206])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) => p(s(bool,X24)))) <=> ((p(s(bool,X26)) | p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24))) & (~p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.sat.sat_27_prove_ltb')).
fof(f22186,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl644),
  inference(cnf_transformation,[],[f22186_D])).
fof(f22186_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl644),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl644])])).
fof(f21372,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f2484,f553])).
fof(f21373,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f2484,f553])).
fof(f21374,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f2484,f553])).
fof(f22188,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | $spl646),
  inference(cnf_transformation,[],[f22188_D])).
fof(f22188_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') <=> ~$spl646),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl646])])).
fof(f21376,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f553])).
fof(f21377,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f553])).
fof(f21378,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f553])).
fof(f21379,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f553])).
fof(f22183,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl643),
  inference(cnf_transformation,[],[f22183_D])).
fof(f22183_D,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl643),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl643])])).
fof(f21385,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f556])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f21386,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f556])).
fof(f21387,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f556])).
fof(f21388,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f556])).
fof(f21402,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f352,f2442,f2484,f560])).
fof(f560,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X13,X18,X16,X19] : ((s(bool,X13) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X16) = s(bool,X19))) => ((p(s(bool,X13)) => p(s(bool,X16))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.IMP_CONG')).
fof(f21403,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f352,f2442,f2484,f560])).
fof(f21404,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f2484,f560])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X14,X15] : s(X6,X14) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X14),s(X6,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.COND_CLAUSES|split|0')).
fof(f21405,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f2484,f560])).
fof(f21406,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f421,f2484,f560])).
fof(f21407,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f421,f2484,f560])).
fof(f21408,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f2484,f560])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X14,X15] : s(X6,X15) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X14),s(X6,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.COND_CLAUSES|split|1')).
fof(f21409,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f2484,f560])).
fof(f21410,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f422,f2484,f560])).
fof(f21411,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f422,f2484,f560])).
fof(f21412,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f591,f2484,f560])).
fof(f591,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(fun(X4,X2),X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,happ(s(fun(X4,X2),X7),s(X4,X9))))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(fun(X4,X2),X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,happ(s(fun(X4,X2),X7),s(X4,X9))))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8,X9] : s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(fun(X4,X2),X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,happ(s(fun(X4,X2),X7),s(X4,X9))))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X31,X32,X33,X34,X35] : ! [X36,X37,X38,X39,X40] : s(X31,happ(s(fun(X35,X31),'_dst_x0x1x1x1_4'(s(fun(X32,fun(X33,fun(X34,X31))),X36),s(fun(X35,X32),X37),s(fun(X35,X33),X38),s(fun(X35,X34),X39))),s(X35,X40))) = s(X31,happ(s(fun(X34,X31),happ(s(fun(X33,fun(X34,X31)),happ(s(fun(X32,fun(X33,fun(X34,X31))),X36),s(X32,happ(s(fun(X35,X32),X37),s(X35,X40))))),s(X33,happ(s(fun(X35,X33),X38),s(X35,X40))))),s(X34,happ(s(fun(X35,X34),X39),s(X35,X40)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','_dst_x0x1x1x1_ax')).
fof(f21413,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f591,f2484,f560])).
fof(f21861,plain,(
  ~p(s(bool,'T_0')) | $spl625),
  inference(cnf_transformation,[],[f21861_D])).
fof(f21861_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl625),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl625])])).
fof(f22164,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22163,f422])).
fof(f22163,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22162,f11795])).
fof(f11795,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,X0)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f1208,f431])).
fof(f1208,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,X0))))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK33(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33])],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X7,X6] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X22),s(fun(X7,bool),X23)))) <=> ! [X21] : ~p(s(bool,happ(s(fun(X7,bool),X23),s(X7,happ(s(fun(X6,X7),X22),s(X6,X21)))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f22162,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22161,f10917])).
fof(f10917,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f605,f434])).
fof(f434,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f605,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,X0))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK24(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24])],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7,X6] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X22),s(fun(X7,bool),X23)))) <=> ! [X21] : p(s(bool,happ(s(fun(X7,bool),X23),s(X7,happ(s(fun(X6,X7),X22),s(X6,X21)))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.quantHeuristics.GUESS_REWRITES|split|2')).
fof(f22161,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl304 | $spl307)),
  inference(forward_demodulation,[],[f21426,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : ! [X28,X29,X30] : s(X6,'const.bool.COND_3'(s(bool,X28),s(X6,X29),s(X6,X30))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),happ(s(fun(bool,fun(X6,fun(X6,X6))),'const.bool.COND_0'),s(bool,X28))),s(X6,X29))),s(X6,X30)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','ARITY_0_3_const.bool.COND')).
fof(f21426,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2442,f591,f2484,f560])).
fof(f22160,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22159,f422])).
fof(f22159,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22158,f11795])).
fof(f22158,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22157,f10917])).
fof(f22157,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21427,f426])).
fof(f21427,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f421,f591,f2484,f560])).
fof(f22156,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22155,f422])).
fof(f22155,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22154,f11795])).
fof(f22154,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22153,f10917])).
fof(f22153,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21428,f426])).
fof(f21428,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f422,f591,f2484,f560])).
fof(f22152,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22151,f422])).
fof(f22151,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22150,f11795])).
fof(f22150,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22149,f10917])).
fof(f22149,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21429,f426])).
fof(f21429,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f591,f2484,f560])).
fof(f21431,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f421,f2484,f560])).
fof(f21433,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f421,f2484,f560])).
fof(f22143,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22142,f422])).
fof(f22142,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22141,f11795])).
fof(f22141,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22140,f10917])).
fof(f22140,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21439,f426])).
fof(f21439,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f421,f2484,f560])).
fof(f21447,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f422,f2484,f560])).
fof(f21449,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f422,f2484,f560])).
fof(f22128,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22127,f422])).
fof(f22127,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22126,f11795])).
fof(f22126,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22125,f10917])).
fof(f22125,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21455,f426])).
fof(f21455,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f422,f2484,f560])).
fof(f21463,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f591,f2484,f560])).
fof(f21465,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f591,f2484,f560])).
fof(f22113,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22112,f422])).
fof(f22112,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22111,f11795])).
fof(f22111,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22110,f10917])).
fof(f22110,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21471,f426])).
fof(f21471,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f591,f2484,f560])).
fof(f21478,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f345,f1135,f2484,f562])).
fof(f562,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f21479,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f345,f1135,f2484,f562])).
fof(f21480,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f344,f1135,f2484,f562])).
fof(f21481,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f344,f1135,f2484,f562])).
fof(f21482,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f345,f1740,f2484,f562])).
fof(f21483,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f345,f1740,f2484,f562])).
fof(f21484,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f344,f1740,f2484,f562])).
fof(f21485,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f344,f1740,f2484,f562])).
fof(f21486,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f421,f2484,f562])).
fof(f21487,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f421,f2484,f562])).
fof(f21488,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f421,f2484,f562])).
fof(f21489,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f421,f2484,f562])).
fof(f21490,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f422,f2484,f562])).
fof(f21491,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f422,f2484,f562])).
fof(f21492,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f422,f2484,f562])).
fof(f21493,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f422,f2484,f562])).
fof(f22091,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f22090,f422])).
fof(f22090,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f22089,f11795])).
fof(f22089,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f22088,f10917])).
fof(f22088,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl102 | $spl307)),
  inference(forward_demodulation,[],[f21506,f426])).
fof(f21506,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1135,f591,f2484,f562])).
fof(f22087,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f22086,f422])).
fof(f22086,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f22085,f11795])).
fof(f22085,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f22084,f10917])).
fof(f22084,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl208 | $spl307)),
  inference(forward_demodulation,[],[f21507,f426])).
fof(f21507,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1740,f591,f2484,f562])).
fof(f22083,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22082,f422])).
fof(f22082,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22081,f11795])).
fof(f22081,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22080,f10917])).
fof(f22080,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21508,f426])).
fof(f21508,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f421,f591,f2484,f562])).
fof(f22079,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22078,f422])).
fof(f22078,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22077,f11795])).
fof(f22077,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22076,f10917])).
fof(f22076,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21509,f426])).
fof(f21509,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f422,f591,f2484,f562])).
fof(f21510,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f352,f2442,f2484,f564])).
fof(f564,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f21511,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f352,f2442,f2484,f564])).
fof(f21512,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f2484,f564])).
fof(f21513,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f2484,f564])).
fof(f21514,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f421,f2484,f564])).
fof(f21515,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f421,f2484,f564])).
fof(f21516,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f2484,f564])).
fof(f21517,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f2484,f564])).
fof(f21518,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f422,f2484,f564])).
fof(f21519,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f422,f2484,f564])).
fof(f21520,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f591,f2484,f564])).
fof(f21521,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f591,f2484,f564])).
fof(f22063,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22062,f422])).
fof(f22062,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22061,f11795])).
fof(f22061,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22060,f10917])).
fof(f22060,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl304 | $spl307)),
  inference(forward_demodulation,[],[f21534,f426])).
fof(f21534,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2442,f591,f2484,f564])).
fof(f22059,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22058,f422])).
fof(f22058,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22057,f11795])).
fof(f22057,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22056,f10917])).
fof(f22056,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21535,f426])).
fof(f21535,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f421,f591,f2484,f564])).
fof(f22055,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22054,f422])).
fof(f22054,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22053,f11795])).
fof(f22053,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22052,f10917])).
fof(f22052,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21536,f426])).
fof(f21536,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f422,f591,f2484,f564])).
fof(f22051,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22050,f422])).
fof(f22050,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22049,f11795])).
fof(f22049,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22048,f10917])).
fof(f22048,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21537,f426])).
fof(f21537,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f591,f2484,f564])).
fof(f21539,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f421,f2484,f564])).
fof(f21541,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f421,f2484,f564])).
fof(f22042,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22041,f422])).
fof(f22041,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22040,f11795])).
fof(f22040,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22039,f10917])).
fof(f22039,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21547,f426])).
fof(f21547,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f421,f2484,f564])).
fof(f21555,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f422,f2484,f564])).
fof(f21557,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f422,f2484,f564])).
fof(f22027,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22026,f422])).
fof(f22026,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22025,f11795])).
fof(f22025,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22024,f10917])).
fof(f22024,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21563,f426])).
fof(f21563,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f422,f2484,f564])).
fof(f21571,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f591,f2484,f564])).
fof(f21573,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f591,f2484,f564])).
fof(f22012,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22011,f422])).
fof(f22011,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22010,f11795])).
fof(f22010,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22009,f10917])).
fof(f22009,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21579,f426])).
fof(f21579,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f591,f2484,f564])).
fof(f21586,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f345,f2484,f565])).
fof(f565,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f326])).
fof(f21587,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f345,f2484,f565])).
fof(f21588,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f344,f2484,f565])).
fof(f21589,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f344,f2484,f565])).
fof(f21994,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21993,f422])).
fof(f21993,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21992,f11795])).
fof(f21992,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f21991,f10917])).
fof(f21991,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21598,f426])).
fof(f21598,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f591,f2484,f565])).
fof(f21990,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21989,f422])).
fof(f21989,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21988,f11795])).
fof(f21988,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f21987,f10917])).
fof(f21987,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21599,f426])).
fof(f21599,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f591,f2484,f565])).
fof(f21602,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f345,f2484,f565])).
fof(f21603,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f345,f2484,f565])).
fof(f21606,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f344,f2484,f565])).
fof(f21607,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f344,f2484,f565])).
fof(f21967,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21966,f422])).
fof(f21966,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21965,f11795])).
fof(f21965,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f21964,f10917])).
fof(f21964,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21627,f426])).
fof(f21627,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f591,f2484,f565])).
fof(f21628,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f345,f1135,f2484,f566])).
fof(f566,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f21629,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f345,f1135,f2484,f566])).
fof(f21630,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f344,f1135,f2484,f566])).
fof(f21631,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f344,f1135,f2484,f566])).
fof(f21632,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f345,f1740,f2484,f566])).
fof(f21633,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f345,f1740,f2484,f566])).
fof(f21634,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f344,f1740,f2484,f566])).
fof(f21635,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f344,f1740,f2484,f566])).
fof(f21636,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f421,f2484,f566])).
fof(f21637,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f421,f2484,f566])).
fof(f21638,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f421,f2484,f566])).
fof(f21639,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f421,f2484,f566])).
fof(f21640,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f422,f2484,f566])).
fof(f21641,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f422,f2484,f566])).
fof(f21642,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f422,f2484,f566])).
fof(f21643,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f422,f2484,f566])).
fof(f21951,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f21950,f422])).
fof(f21950,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f21949,f11795])).
fof(f21949,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f21948,f10917])).
fof(f21948,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl102 | $spl307)),
  inference(forward_demodulation,[],[f21656,f426])).
fof(f21656,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1135,f591,f2484,f566])).
fof(f21947,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f21946,f422])).
fof(f21946,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f21945,f11795])).
fof(f21945,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f21944,f10917])).
fof(f21944,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl208 | $spl307)),
  inference(forward_demodulation,[],[f21657,f426])).
fof(f21657,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1740,f591,f2484,f566])).
fof(f21943,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21942,f422])).
fof(f21942,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21941,f11795])).
fof(f21941,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f21940,f10917])).
fof(f21940,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21658,f426])).
fof(f21658,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f421,f591,f2484,f566])).
fof(f21939,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21938,f422])).
fof(f21938,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21937,f11795])).
fof(f21937,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f21936,f10917])).
fof(f21936,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21659,f426])).
fof(f21659,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f422,f591,f2484,f566])).
fof(f21935,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl641),
  inference(cnf_transformation,[],[f21935_D])).
fof(f21935_D,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl641),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl641])])).
fof(f21660,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f567])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (~sP15(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP15(X0,X1,X2)))),
  inference(rectify,[],[f328])).
fof(f328,plain,(
  ! [X3,X1,X2] : ((~sP15(X3,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP15(X3,X1,X2)))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X3,X1,X2] : ((~sP15(X3,X1,X2) | (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3)))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP15(X3,X1,X2)))),
  inference(nnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X3,X1,X2] : (sP15(X3,X1,X2) <=> (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f21661,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f567])).
fof(f21662,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f567])).
fof(f21663,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f567])).
fof(f21931,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl639),
  inference(cnf_transformation,[],[f21931_D])).
fof(f21931_D,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl639),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl639])])).
fof(f21668,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f571])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (~sP14(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1,X2)))),
  inference(rectify,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X3] : ((~sP14(X0,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1,X3)))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X3] : ((~sP14(X0,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1,X3)))),
  inference(nnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X3] : (sP14(X0,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f21669,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f571])).
fof(f21670,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f571])).
fof(f21671,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f571])).
fof(f21927,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl637),
  inference(cnf_transformation,[],[f21927_D])).
fof(f21927_D,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl637),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl637])])).
fof(f21676,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f575])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP13(X0,X1,X2)))),
  inference(rectify,[],[f334])).
fof(f334,plain,(
  ! [X2,X1,X3] : ((~sP13(X2,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP13(X2,X1,X3)))),
  inference(flattening,[],[f333])).
fof(f333,plain,(
  ! [X2,X1,X3] : ((~sP13(X2,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP13(X2,X1,X3)))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X2,X1,X3] : (sP13(X2,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f21677,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f575])).
fof(f21678,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f575])).
fof(f21679,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f575])).
fof(f21923,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl635),
  inference(cnf_transformation,[],[f21923_D])).
fof(f21923_D,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl635),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl635])])).
fof(f21684,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f579])).
fof(f579,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f337])).
fof(f337,plain,(
  ! [X0,X2,X3] : ((~sP12(X0,X2,X3) | p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP12(X0,X2,X3)))),
  inference(flattening,[],[f336])).
fof(f336,plain,(
  ! [X0,X2,X3] : ((~sP12(X0,X2,X3) | (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP12(X0,X2,X3)))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X2,X3] : (sP12(X0,X2,X3) <=> (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f21685,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f579])).
fof(f21916,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl631),
  inference(cnf_transformation,[],[f21916_D])).
fof(f21916_D,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl631),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl631])])).
fof(f21686,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f579])).
fof(f21687,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f579])).
fof(f21918,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl633),
  inference(cnf_transformation,[],[f21918_D])).
fof(f21918_D,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl633),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl633])])).
fof(f21691,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f579])).
fof(f21692,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f579])).
fof(f21693,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f579])).
fof(f21694,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f579])).
fof(f21896,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl2 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21895,f422])).
fof(f21895,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl2 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21721,f2442])).
fof(f21721,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2150,f2159,f2144,f893,f2484,f590])).
fof(f590,plain,(
  ( ! [X2,X0,X3,X1] : (~sP15(X3,X1,X2) | p(s(bool,X3)) | s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | ~sP14(X0,X1,X3) | ~sP13(X2,X1,X3) | ~sP12(X0,X2,X3)) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (sP12(X0,X2,X3) & sP13(X2,X1,X3) & sP14(X0,X1,X3) & sP15(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3))))) & (~sP12(X0,X2,X3) | ~sP13(X2,X1,X3) | ~sP14(X0,X1,X3) | ~sP15(X3,X1,X2) | (~p(s(bool,X2)) & ~p(s(bool,X0)) & p(s(bool,X3))) | s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (sP12(X0,X2,X3) & sP13(X2,X1,X3) & sP14(X0,X1,X3) & sP15(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3))))) & ((~sP12(X0,X2,X3) | ~sP13(X2,X1,X3) | ~sP14(X0,X1,X3) | ~sP15(X3,X1,X2) | (~p(s(bool,X2)) & ~p(s(bool,X0)) & p(s(bool,X3)))) | s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))))),
  inference(nnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> (sP12(X0,X2,X3) & sP13(X2,X1,X3) & sP14(X0,X1,X3) & sP15(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(definition_folding,[],[f153,f212,f211,f210,f209])).
fof(f153,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X27,X24,X25,X26] : (s(bool,X26) = s(bool,'const.bool.COND_3'(s(bool,X25),s(bool,X24),s(bool,X27))) <=> ((p(s(bool,X26)) | p(s(bool,X25)) | ~p(s(bool,X27))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X27))) & (~p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26))) & (p(s(bool,X25)) | p(s(bool,X27)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.sat.sat_3|49|_prove_ltb')).
fof(f893,plain,(
  ( ! [X0,X1] : (sP15(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f569])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP15(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f329])).
fof(f2144,plain,(
  ( ! [X0,X1] : (sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f574])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f332])).
fof(f2159,plain,(
  ( ! [X0,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f582])).
fof(f582,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f338])).
fof(f2150,plain,(
  ( ! [X0,X1] : (sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f578])).
fof(f578,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f21894,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21893,f422])).
fof(f21893,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21722,f2442])).
fof(f21722,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2150,f2159,f2144,f1498,f2484,f590])).
fof(f1498,plain,(
  ( ! [X0,X1] : (sP15(X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f569])).
fof(f21892,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21891,f422])).
fof(f21891,plain,(
  ( ! [X0] : (s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))) ) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21723,f2442])).
fof(f21723,plain,(
  ( ! [X0] : (s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,X0),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))))) ) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2144,f2159,f2150,f2137,f2484,f590])).
fof(f2137,plain,(
  ( ! [X0,X1] : (sP15(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f568])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f329])).
fof(f21890,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21889,f422])).
fof(f21889,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'F_0'),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21724,f2442])).
fof(f21724,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2143,f2159,f2150,f2137,f2484,f590])).
fof(f2143,plain,(
  ( ! [X0,X1] : (sP14(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f573])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f332])).
fof(f21888,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21887,f422])).
fof(f21887,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'F_0'),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21725,f2442])).
fof(f21725,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2159,f2143,f2149,f2137,f2484,f590])).
fof(f2149,plain,(
  ( ! [X0,X1] : (sP13(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f577])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f335])).
fof(f21884,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21883,f422])).
fof(f21883,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'F_0'),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f21726,f2442])).
fof(f21726,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2159,f2144,f2149,f2137,f2484,f590])).
fof(f21729,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f2484,f598])).
fof(f598,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f563])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f21730,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f2484,f598])).
fof(f21878,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21877,f422])).
fof(f21877,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21876,f11795])).
fof(f21876,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f21875,f10917])).
fof(f21875,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21735,f426])).
fof(f21735,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f591,f2484,f598])).
fof(f21738,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f598])).
fof(f21739,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f598])).
fof(f21859,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21858,f422])).
fof(f21858,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f21857,f11795])).
fof(f21857,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0')))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f21856,f10917])).
fof(f21856,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f21751,f426])).
fof(f21751,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f2484,f598])).
fof(f21855,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl622),
  inference(cnf_transformation,[],[f21855_D])).
fof(f21855_D,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl622),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl622])])).
fof(f21752,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f931])).
fof(f931,plain,(
  ( ! [X6] : (sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X6) | p(s(bool,X6))) ) | $spl2),
  inference(resolution,[],[f602,f454])).
fof(f21853,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl620),
  inference(cnf_transformation,[],[f21853_D])).
fof(f21853_D,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl620),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl620])])).
fof(f21753,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f932])).
fof(f932,plain,(
  ( ! [X7] : (sP1(X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | p(s(bool,X7))) ) | $spl2),
  inference(resolution,[],[f602,f457])).
fof(f21851,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl618),
  inference(cnf_transformation,[],[f21851_D])).
fof(f21851_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl618),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl618])])).
fof(f21754,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f948])).
fof(f948,plain,(
  ( ! [X35,X36] : (sP11(X36,X35,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl2),
  inference(resolution,[],[f602,f553])).
fof(f21847,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl614),
  inference(cnf_transformation,[],[f21847_D])).
fof(f21847_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl614),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl614])])).
fof(f21755,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f948])).
fof(f21849,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl616),
  inference(cnf_transformation,[],[f21849_D])).
fof(f21849_D,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl616),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl616])])).
fof(f21756,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f948])).
fof(f21757,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f948])).
fof(f21845,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl612),
  inference(cnf_transformation,[],[f21845_D])).
fof(f21845_D,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl612),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl612])])).
fof(f21759,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f1536])).
fof(f1536,plain,(
  ( ! [X6] : (sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X6) | p(s(bool,X6))) ) | $spl4),
  inference(resolution,[],[f604,f454])).
fof(f21843,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl610),
  inference(cnf_transformation,[],[f21843_D])).
fof(f21843_D,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl610),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl610])])).
fof(f21760,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f1537])).
fof(f1537,plain,(
  ( ! [X7] : (sP1(X7,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | p(s(bool,X7))) ) | $spl4),
  inference(resolution,[],[f604,f457])).
fof(f21841,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl608),
  inference(cnf_transformation,[],[f21841_D])).
fof(f21841_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl608),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl608])])).
fof(f21761,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f1553])).
fof(f1553,plain,(
  ( ! [X35,X36] : (sP11(X36,X35,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl4),
  inference(resolution,[],[f604,f553])).
fof(f21837,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl604),
  inference(cnf_transformation,[],[f21837_D])).
fof(f21837_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f21762,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f1553])).
fof(f21839,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl606),
  inference(cnf_transformation,[],[f21839_D])).
fof(f21839_D,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl606),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl606])])).
fof(f21763,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f1553])).
fof(f21764,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f1553])).
fof(f21835,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl602),
  inference(cnf_transformation,[],[f21835_D])).
fof(f21835_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl602),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl602])])).
fof(f21770,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2219])).
fof(f2219,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X27) | p(s(bool,X27)) | p(s(bool,X26))) ) | $spl1),
  inference(resolution,[],[f600,f536])).
fof(f21831,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl598),
  inference(cnf_transformation,[],[f21831_D])).
fof(f21831_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl598])])).
fof(f21771,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2219])).
fof(f21833,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl600),
  inference(cnf_transformation,[],[f21833_D])).
fof(f21833_D,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl600),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl600])])).
fof(f21772,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2219])).
fof(f21773,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2219])).
fof(f21829,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl596),
  inference(cnf_transformation,[],[f21829_D])).
fof(f21829_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f21774,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2220])).
fof(f2220,plain,(
  ( ! [X28,X29] : (sP9(X28,X29,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl1),
  inference(resolution,[],[f600,f536])).
fof(f21825,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl592),
  inference(cnf_transformation,[],[f21825_D])).
fof(f21825_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f21775,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2220])).
fof(f21827,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl594),
  inference(cnf_transformation,[],[f21827_D])).
fof(f21827_D,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl594])])).
fof(f21776,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2220])).
fof(f21777,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2220])).
fof(f21778,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2221])).
fof(f2221,plain,(
  ( ! [X30,X31] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X30,X31) | p(s(bool,X30)) | p(s(bool,X31))) ) | $spl1),
  inference(resolution,[],[f600,f536])).
fof(f21779,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2221])).
fof(f21781,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2221])).
fof(f21782,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2221])).
fof(f21016,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl585),
  inference(cnf_transformation,[],[f21016_D])).
fof(f21016_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl585),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl585])])).
fof(f21014,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl277),
  inference(unit_resulting_resolution,[],[f2414,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f20996,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl583),
  inference(cnf_transformation,[],[f20996_D])).
fof(f20996_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl583),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl583])])).
fof(f20994,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl273),
  inference(unit_resulting_resolution,[],[f2410,f485])).
fof(f20988,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl581),
  inference(cnf_transformation,[],[f20988_D])).
fof(f20988_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl581),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl581])])).
fof(f20986,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl271),
  inference(unit_resulting_resolution,[],[f2408,f485])).
fof(f20967,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl578),
  inference(cnf_transformation,[],[f20967_D])).
fof(f20967_D,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f20964,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl248),
  inference(resolution,[],[f2386,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f311])).
fof(f20958,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl248),
  inference(unit_resulting_resolution,[],[f2386,f527])).
fof(f20460,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl576),
  inference(cnf_transformation,[],[f20460_D])).
fof(f20460_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f20447,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl211)),
  inference(unit_resulting_resolution,[],[f2241,f2211])).
fof(f2211,plain,(
  ( ! [X10,X11] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X10,X11) | s(bool,X10) = s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f600,f504])).
fof(f20458,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl574),
  inference(cnf_transformation,[],[f20458_D])).
fof(f20458_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f20450,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | ($spl1 | $spl211)),
  inference(unit_resulting_resolution,[],[f2241,f2211])).
fof(f19748,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl572),
  inference(cnf_transformation,[],[f19748_D])).
fof(f19748_D,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f15617,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl304),
  inference(unit_resulting_resolution,[],[f345,f2442,f505])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f301])).
fof(f15618,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl304),
  inference(unit_resulting_resolution,[],[f344,f2442,f505])).
fof(f18124,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl570),
  inference(cnf_transformation,[],[f18124_D])).
fof(f18124_D,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f16431,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl304),
  inference(unit_resulting_resolution,[],[f345,f2442,f505])).
fof(f16432,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl304),
  inference(unit_resulting_resolution,[],[f344,f2442,f505])).
fof(f15356,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl568),
  inference(cnf_transformation,[],[f15356_D])).
fof(f15356_D,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f15341,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl211),
  inference(unit_resulting_resolution,[],[f352,f2241,f504])).
fof(f15353,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl566),
  inference(cnf_transformation,[],[f15353_D])).
fof(f15353_D,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f15345,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl211),
  inference(unit_resulting_resolution,[],[f352,f2241,f504])).
fof(f15334,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl565),
  inference(cnf_transformation,[],[f15334_D])).
fof(f15334_D,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl565),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl565])])).
fof(f15330,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl285),
  inference(unit_resulting_resolution,[],[f2422,f487])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f15325,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl563),
  inference(cnf_transformation,[],[f15325_D])).
fof(f15325_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') <=> ~$spl563),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl563])])).
fof(f15321,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl283),
  inference(unit_resulting_resolution,[],[f2420,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f15316,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl561),
  inference(cnf_transformation,[],[f15316_D])).
fof(f15316_D,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl561),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl561])])).
fof(f15314,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl281),
  inference(unit_resulting_resolution,[],[f2418,f485])).
fof(f15309,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl559),
  inference(cnf_transformation,[],[f15309_D])).
fof(f15309_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl559),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl559])])).
fof(f15307,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl279),
  inference(unit_resulting_resolution,[],[f2416,f485])).
fof(f15302,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl557),
  inference(cnf_transformation,[],[f15302_D])).
fof(f15302_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') <=> ~$spl557),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl557])])).
fof(f15300,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl275),
  inference(unit_resulting_resolution,[],[f2412,f485])).
fof(f15255,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl554),
  inference(cnf_transformation,[],[f15255_D])).
fof(f15255_D,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f15253,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl252),
  inference(resolution,[],[f2390,f527])).
fof(f15247,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl252),
  inference(unit_resulting_resolution,[],[f2390,f527])).
fof(f15238,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl552),
  inference(cnf_transformation,[],[f15238_D])).
fof(f15238_D,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f15236,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl250),
  inference(resolution,[],[f2388,f527])).
fof(f15230,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl250),
  inference(unit_resulting_resolution,[],[f2388,f527])).
fof(f15174,plain,(
  sP10('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl550),
  inference(cnf_transformation,[],[f15174_D])).
fof(f15174_D,plain,(
  sP10('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f15172,plain,(
  sP10('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl230),
  inference(resolution,[],[f2368,f544])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f319])).
fof(f15166,plain,(
  sP10('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl230),
  inference(unit_resulting_resolution,[],[f2368,f544])).
fof(f14918,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))),
  introduced(theory_axiom,[])).
fof(f14608,plain,(
  p(s(bool,'F_0')) | $spl536),
  inference(cnf_transformation,[],[f14608_D])).
fof(f14608_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl536])])).
fof(f14916,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl548),
  inference(cnf_transformation,[],[f14916_D])).
fof(f14916_D,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f14915,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) | p(s(bool,'F_0')) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | ($spl1 | $spl304)),
  inference(forward_demodulation,[],[f14914,f2442])).
fof(f14914,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'F_0')) | s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | ($spl1 | $spl304)),
  inference(forward_demodulation,[],[f14908,f2442])).
fof(f14908,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl1),
  inference(resolution,[],[f2138,f590])).
fof(f2138,plain,(
  ( ! [X0,X1] : (sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f570])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f14901,plain,(
  sP5('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl546),
  inference(cnf_transformation,[],[f14901_D])).
fof(f14901_D,plain,(
  sP5('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f14897,plain,(
  sP5('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f1858,f1864,f12765,f491])).
fof(f14898,plain,(
  sP5('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f12752,f1864,f12765,f491])).
fof(f14784,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl544),
  inference(cnf_transformation,[],[f14784_D])).
fof(f14784_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f14776,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f12752,f12763,f12764,f489])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f14777,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f12752,f12763,f12764,f489])).
fof(f14782,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl542),
  inference(cnf_transformation,[],[f14782_D])).
fof(f14782_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f14778,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f1858,f12763,f12764,f490])).
fof(f14779,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f1858,f1864,f12764,f490])).
fof(f14620,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl540),
  inference(cnf_transformation,[],[f14620_D])).
fof(f14620_D,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f14612,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f352,f684,f1858,f1864,f491])).
fof(f14618,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl538),
  inference(cnf_transformation,[],[f14618_D])).
fof(f14618_D,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f14613,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f352,f1287,f1858,f1864,f491])).
fof(f14614,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f352,f684,f675,f1864,f491])).
fof(f14615,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f352,f1287,f1278,f1864,f491])).
fof(f14606,plain,(
  ( ! [X6,X7] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X6,X7) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl534),
  inference(cnf_transformation,[],[f14606_D])).
fof(f14606_D,plain,(
  ( ! [X6,X7] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X6,X7) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) <=> ~$spl534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl534])])).
fof(f14605,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X6,X7) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | ($spl1 | $spl304)),
  inference(forward_demodulation,[],[f14602,f2442])).
fof(f14602,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X6,X7) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl1),
  inference(resolution,[],[f1859,f491])).
fof(f1859,plain,(
  ( ! [X0,X1] : (sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f495])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f14593,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl532),
  inference(cnf_transformation,[],[f14593_D])).
fof(f14593_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl532),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl532])])).
fof(f14584,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f352,f683,f682,f1858,f490])).
fof(f14591,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl530),
  inference(cnf_transformation,[],[f14591_D])).
fof(f14591_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl530),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl530])])).
fof(f14585,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f352,f1286,f1285,f1858,f490])).
fof(f14561,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl528),
  inference(cnf_transformation,[],[f14561_D])).
fof(f14561_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f14550,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f683,f682,f12752,f489])).
fof(f14551,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f683,f682,f12752,f489])).
fof(f14559,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl526),
  inference(cnf_transformation,[],[f14559_D])).
fof(f14559_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f14552,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f1286,f1285,f12752,f489])).
fof(f14553,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f1286,f1285,f12752,f489])).
fof(f14429,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl525),
  inference(cnf_transformation,[],[f14429_D])).
fof(f14429_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl525),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl525])])).
fof(f14425,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl189),
  inference(unit_resulting_resolution,[],[f1720,f487])).
fof(f14417,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl523),
  inference(cnf_transformation,[],[f14417_D])).
fof(f14417_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl523),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl523])])).
fof(f14413,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl183),
  inference(unit_resulting_resolution,[],[f1714,f486])).
fof(f14386,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl520),
  inference(cnf_transformation,[],[f14386_D])).
fof(f14386_D,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f14379,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl162),
  inference(resolution,[],[f1694,f508])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f303])).
fof(f303,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f201])).
fof(f14373,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl162),
  inference(unit_resulting_resolution,[],[f1694,f508])).
fof(f14238,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl519),
  inference(cnf_transformation,[],[f14238_D])).
fof(f14238_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl519),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl519])])).
fof(f14234,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl81),
  inference(unit_resulting_resolution,[],[f1113,f487])).
fof(f14226,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl517),
  inference(cnf_transformation,[],[f14226_D])).
fof(f14226_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl517),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl517])])).
fof(f14222,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl75),
  inference(unit_resulting_resolution,[],[f1107,f486])).
fof(f14195,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl514),
  inference(cnf_transformation,[],[f14195_D])).
fof(f14195_D,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f14188,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl58),
  inference(resolution,[],[f1091,f508])).
fof(f14182,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl58),
  inference(unit_resulting_resolution,[],[f1091,f508])).
fof(f13621,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl512),
  inference(cnf_transformation,[],[f13621_D])).
fof(f13621_D,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f13706,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(global_subsumption,[],[f343,f342,f341,f11732,f13203,f13295,f13177,f13176,f13171,f13170,f13167,f13166,f13165,f13162,f13161,f13160,f13159,f13156,f13155,f13154,f13151,f13150,f13149,f13148,f13145,f13144,f13143,f13140,f13139,f13138,f13114,f13113,f13108,f13096,f13082,f13068,f13054,f13042,f13041,f13040,f13033,f13032,f13020,f13006,f12992,f12966,f12965,f12960,f12953,f12952,f12949,f12946,f12945,f12944,f12943,f12942,f12941,f12940,f12939,f12936,f12934,f12931,f12930,f12927,f12926,f12925,f12922,f12921,f12920,f12896,f12891,f12890,f12889,f12888,f12883,f12879,f12878,f12875,f12873,f12872,f12867,f12852,f12851,f12850,f12849,f12848,f12847,f12846,f12845,f12844,f12843,f12842,f12841,f12840,f12839,f12838,f12837,f12836,f12835,f12834,f12833,f12832,f12831,f12830,f12829,f12828,f12827,f12826,f12825,f12824,f12823,f12822,f12821,f12820,f12819,f12818,f12817,f12814,f12813,f12812,f12809,f12808,f12807,f12780,f12779,f12778,f12777,f12774,f12773,f12772,f12771,f12770,f12767,f12766,f12765,f12764,f12763,f12760,f12759,f12758,f12755,f12754,f12753,f12752,f12749,f12748,f12747,f12744,f12743,f12742,f12737,f12736,f12735,f12734,f12733,f12732,f12731,f12730,f12729,f12728,f12727,f12726,f12725,f12724,f12723,f12722,f12721,f12720,f12719,f12718,f12717,f12716,f12715,f12714,f12710,f12709,f12707,f12706,f13623,f13625,f13627,f12656,f12655,f12654,f12653,f13217,f13218,f13225,f13230,f13231,f13232,f13235,f13242,f13245,f13246,f13248,f13249,f13250,f13256])).
fof(f13256,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(superposition,[],[f11732,f354])).
fof(f354,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7',aHL_BOOL_CASES)).
fof(f13250,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X39,X40)) ) | $spl362),
  inference(resolution,[],[f11732,f555])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f13249,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X38,X37)) ) | $spl362),
  inference(resolution,[],[f11732,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f13248,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | sP11(X36,X35,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | p(s(bool,X36))) ) | $spl362),
  inference(resolution,[],[f11732,f553])).
fof(f13246,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X31,X32)) ) | $spl362),
  inference(resolution,[],[f11732,f538])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f13245,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X30,X29)) ) | $spl362),
  inference(resolution,[],[f11732,f537])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f13242,plain,(
  ( ! [X24,X23] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X23,X24)) ) | $spl362),
  inference(resolution,[],[f11732,f521])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f308])).
fof(f13235,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X10)) ) | $spl362),
  inference(resolution,[],[f11732,f460])).
fof(f460,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f13232,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(resolution,[],[f11732,f457])).
fof(f13231,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X6)) ) | $spl362),
  inference(resolution,[],[f11732,f454])).
fof(f13230,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~sP1(X5,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(resolution,[],[f11732,f453])).
fof(f453,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f13225,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(resolution,[],[f11732,f434])).
fof(f13627,plain,(
  ( ! [X2,X0,X5,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,X1),s(X0,X5)))) ) | $spl362),
  inference(forward_demodulation,[],[f13626,f422])).
fof(f13626,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))))),s(X0,X1),s(X0,X5)))) ) | $spl362),
  inference(forward_demodulation,[],[f12703,f422])).
fof(f12703,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X1))),s(X0,X5)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f422,f422,f11732,f428])).
fof(f428,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X2)) | s(X0,X3) != s(X0,X4) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X11,X10,X13,X18,X16,X19] : ((s(bool,X10) = s(bool,X11) & (p(s(bool,X10)) => s(X6,X13) = s(X6,X18)) & (~p(s(bool,X10)) => s(X6,X16) = s(X6,X19))) => s(X6,'const.bool.COND_3'(s(bool,X10),s(X6,X18),s(X6,X19))) = s(X6,'const.bool.COND_3'(s(bool,X11),s(X6,X13),s(X6,X16))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.COND_CONG')).
fof(f13625,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl362),
  inference(forward_demodulation,[],[f13624,f422])).
fof(f13624,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl362),
  inference(forward_demodulation,[],[f12704,f426])).
fof(f12704,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3))),s(X0,X6)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f426,f422,f11732,f428])).
fof(f13623,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl362),
  inference(forward_demodulation,[],[f13622,f422])).
fof(f13622,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl362),
  inference(forward_demodulation,[],[f12705,f591])).
fof(f12705,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))))),s(X0,happ(s(fun(X4,X0),happ(s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(X2,happ(s(fun(X1,X2),X6),s(X1,X9))))),s(X3,happ(s(fun(X1,X3),X7),s(X1,X9))))),s(X4,happ(s(fun(X1,X4),X8),s(X1,X9))))),s(X0,X12)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f591,f422,f11732,f428])).
fof(f12714,plain,(
  ( ! [X0] : (sP0(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f443])).
fof(f443,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f12715,plain,(
  ( ! [X0] : (sP0('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f444])).
fof(f444,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f12766,plain,(
  ( ! [X0,X1] : (sP2('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0,X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f503])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f300])).
fof(f12767,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f418,f11732,f505])).
fof(f418,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.EXISTS_DEF')).
fof(f12770,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f421,f11732,f505])).
fof(f12771,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f422,f11732,f505])).
fof(f12772,plain,(
  ( ! [X2,X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f426,f11732,f505])).
fof(f12773,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),happ(s(fun(X0,bool),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(fun(X0,X1),X5),s(fun(X0,X2),X6),s(fun(X0,X3),X7))),s(X0,X8)),happ(s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),happ(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(X1,happ(s(fun(X0,X1),X5),s(X0,X8))))),s(X2,happ(s(fun(X0,X2),X6),s(X0,X8))))),s(X3,happ(s(fun(X0,X3),X7),s(X0,X8)))))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f591,f11732,f505])).
fof(f12774,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f418,f11732,f505])).
fof(f12777,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f421,f11732,f505])).
fof(f12778,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f422,f11732,f505])).
fof(f12779,plain,(
  ( ! [X2,X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f426,f11732,f505])).
fof(f12780,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),happ(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,happ(s(fun(X4,X1),X7),s(X4,X6))))),s(X0,happ(s(fun(X4,X0),X8),s(X4,X6)))),happ(s(fun(X4,bool),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7),s(fun(X4,X0),X8))),s(X4,X6)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f591,f11732,f505])).
fof(f12817,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f516])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f306])).
fof(f12845,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0,'F_0')) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f522])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f308])).
fof(f12846,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0',X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f523])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f308])).
fof(f12872,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f533])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f12873,plain,(
  ( ! [X0,X1] : (sP8(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f534])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f12875,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0,'T_0')) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f537])).
fof(f12878,plain,(
  ( ! [X0] : (sP9('T_0',X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f537])).
fof(f12879,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f537])).
fof(f12883,plain,(
  ( ! [X0] : (sP9('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f538])).
fof(f12888,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0',X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f538])).
fof(f12889,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f538])).
fof(f12891,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f540])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f12896,plain,(
  ( ! [X0] : (~sP9('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f541])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f12930,plain,(
  ( ! [X0,X1] : (sP10(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f551])).
fof(f551,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f322])).
fof(f12934,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0,'F_0')) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f554])).
fof(f12936,plain,(
  ( ! [X0] : (sP11('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f555])).
fof(f12939,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0',X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f555])).
fof(f12940,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f555])).
fof(f12949,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f558])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f324])).
fof(f13148,plain,(
  ( ! [X0,X1] : (sP15(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f569])).
fof(f13159,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP14(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f332])).
fof(f13170,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f576])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP13(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f335])).
fof(f13176,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f580])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f338])).
fof(f13177,plain,(
  ( ! [X0,X1] : (sP12(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK20)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f581])).
fof(f581,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f338])).
fof(f13218,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(resolution,[],[f11732,f364])).
fof(f364,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f13217,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(resolution,[],[f11732,f362])).
fof(f362,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f220])).
fof(f12653,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f362])).
fof(f12654,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f364])).
fof(f13269,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl389),
  inference(cnf_transformation,[],[f13269_D])).
fof(f13269_D,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl389),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl389])])).
fof(f12655,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f365])).
fof(f365,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f12656,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f367])).
fof(f367,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f12706,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f434])).
fof(f12707,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f434])).
fof(f12709,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f434])).
fof(f12710,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f434])).
fof(f13613,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl505),
  inference(cnf_transformation,[],[f13613_D])).
fof(f13613_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl505),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl505])])).
fof(f12716,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f453])).
fof(f12717,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f453])).
fof(f13611,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl503),
  inference(cnf_transformation,[],[f13611_D])).
fof(f13611_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl503),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl503])])).
fof(f12718,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f453])).
fof(f13615,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl507),
  inference(cnf_transformation,[],[f13615_D])).
fof(f13615_D,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl507),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl507])])).
fof(f12719,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f453])).
fof(f12720,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f453])).
fof(f12721,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f453])).
fof(f13617,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl508),
  inference(cnf_transformation,[],[f13617_D])).
fof(f13617_D,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f12722,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f454])).
fof(f13619,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl510),
  inference(cnf_transformation,[],[f13619_D])).
fof(f13619_D,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f12723,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f457])).
fof(f12724,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f458])).
fof(f12725,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f459])).
fof(f12726,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f460])).
fof(f12727,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f460])).
fof(f12728,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f460])).
fof(f12729,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f460])).
fof(f12730,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f460])).
fof(f12731,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f460])).
fof(f13569,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl469),
  inference(cnf_transformation,[],[f13569_D])).
fof(f13569_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl469),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl469])])).
fof(f12732,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f488])).
fof(f12733,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f488])).
fof(f13607,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl499),
  inference(cnf_transformation,[],[f13607_D])).
fof(f13607_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl499),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl499])])).
fof(f12734,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f488])).
fof(f13609,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl501),
  inference(cnf_transformation,[],[f13609_D])).
fof(f13609_D,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl501),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl501])])).
fof(f12735,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f488])).
fof(f12736,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f488])).
fof(f12737,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f488])).
fof(f13605,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl496),
  inference(cnf_transformation,[],[f13605_D])).
fof(f13605_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f12738,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f684,f682,f675,f11732,f489])).
fof(f13603,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl494),
  inference(cnf_transformation,[],[f13603_D])).
fof(f13603_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f12739,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl2 | $spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f684,f1285,f675,f11732,f489])).
fof(f13601,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl492),
  inference(cnf_transformation,[],[f13601_D])).
fof(f13601_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f12740,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl2 | $spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f1287,f682,f1278,f11732,f489])).
fof(f13599,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl490),
  inference(cnf_transformation,[],[f13599_D])).
fof(f13599_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f12741,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f1287,f1285,f1278,f11732,f489])).
fof(f13597,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl489),
  inference(cnf_transformation,[],[f13597_D])).
fof(f13597_D,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl489),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl489])])).
fof(f12742,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f492])).
fof(f12743,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f492])).
fof(f13591,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl485),
  inference(cnf_transformation,[],[f13591_D])).
fof(f13591_D,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl485),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl485])])).
fof(f12744,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f492])).
fof(f13593,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl487),
  inference(cnf_transformation,[],[f13593_D])).
fof(f13593_D,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl487),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl487])])).
fof(f12747,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f492])).
fof(f12748,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f492])).
fof(f12749,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f492])).
fof(f13587,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl483),
  inference(cnf_transformation,[],[f13587_D])).
fof(f13587_D,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl483),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl483])])).
fof(f12753,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f496])).
fof(f12754,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f496])).
fof(f13581,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl479),
  inference(cnf_transformation,[],[f13581_D])).
fof(f13581_D,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl479),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl479])])).
fof(f12755,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f496])).
fof(f13583,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl481),
  inference(cnf_transformation,[],[f13583_D])).
fof(f13583_D,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl481),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl481])])).
fof(f12758,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f496])).
fof(f12759,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f496])).
fof(f12760,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f496])).
fof(f13577,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl476),
  inference(cnf_transformation,[],[f13577_D])).
fof(f13577_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f12768,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f1135,f11732,f505])).
fof(f13575,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl474),
  inference(cnf_transformation,[],[f13575_D])).
fof(f13575_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f12769,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f1740,f11732,f505])).
fof(f13573,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl472),
  inference(cnf_transformation,[],[f13573_D])).
fof(f13573_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f12775,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f1135,f11732,f505])).
fof(f13571,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl470),
  inference(cnf_transformation,[],[f13571_D])).
fof(f13571_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f12776,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f1740,f11732,f505])).
fof(f12781,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | ($spl29 | $spl362)),
  inference(unit_resulting_resolution,[],[f1003,f11732,f506])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f301])).
fof(f13567,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl467),
  inference(cnf_transformation,[],[f13567_D])).
fof(f13567_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') <=> ~$spl467),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl467])])).
fof(f12783,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | ($spl29 | $spl362)),
  inference(unit_resulting_resolution,[],[f1003,f11732,f506])).
fof(f13565,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl464),
  inference(cnf_transformation,[],[f13565_D])).
fof(f13565_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f12787,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f704,f11732,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (~sP6(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f304])).
fof(f704,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f516])).
fof(f12788,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f704,f11732,f511])).
fof(f13559,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl458),
  inference(cnf_transformation,[],[f13559_D])).
fof(f13559_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f12789,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f704,f11732,f511])).
fof(f13563,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl462),
  inference(cnf_transformation,[],[f13563_D])).
fof(f13563_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f12791,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f1309,f11732,f511])).
fof(f1309,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f516])).
fof(f12792,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f1309,f11732,f511])).
fof(f13555,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl454),
  inference(cnf_transformation,[],[f13555_D])).
fof(f13555_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f12793,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f1309,f11732,f511])).
fof(f13561,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl460),
  inference(cnf_transformation,[],[f13561_D])).
fof(f13561_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f12795,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f704,f11732,f511])).
fof(f12796,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f704,f11732,f511])).
fof(f12797,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f704,f11732,f511])).
fof(f13557,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl456),
  inference(cnf_transformation,[],[f13557_D])).
fof(f13557_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f12799,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f1309,f11732,f511])).
fof(f12800,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f1309,f11732,f511])).
fof(f12801,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f1309,f11732,f511])).
fof(f13549,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl453),
  inference(cnf_transformation,[],[f13549_D])).
fof(f13549_D,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl453),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl453])])).
fof(f12807,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f515])).
fof(f12808,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f515])).
fof(f13543,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl449),
  inference(cnf_transformation,[],[f13543_D])).
fof(f13543_D,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl449),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl449])])).
fof(f12809,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f515])).
fof(f13545,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl451),
  inference(cnf_transformation,[],[f13545_D])).
fof(f13545_D,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl451),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl451])])).
fof(f12812,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f515])).
fof(f12813,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f515])).
fof(f12814,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f515])).
fof(f13539,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl446),
  inference(cnf_transformation,[],[f13539_D])).
fof(f13539_D,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f12818,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f345,f11732,f521])).
fof(f12819,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f345,f11732,f521])).
fof(f13535,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl442),
  inference(cnf_transformation,[],[f13535_D])).
fof(f13535_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f12820,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f345,f11732,f521])).
fof(f12821,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f344,f11732,f521])).
fof(f12822,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f344,f11732,f521])).
fof(f12823,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f344,f11732,f521])).
fof(f13531,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl438),
  inference(cnf_transformation,[],[f13531_D])).
fof(f13531_D,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f12824,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f11732,f521])).
fof(f12825,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f11732,f521])).
fof(f13527,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl434),
  inference(cnf_transformation,[],[f13527_D])).
fof(f13527_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f12826,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f11732,f521])).
fof(f13537,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl444),
  inference(cnf_transformation,[],[f13537_D])).
fof(f13537_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','T_0') <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f12827,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f345,f11732,f521])).
fof(f12828,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f345,f11732,f521])).
fof(f12829,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f345,f11732,f521])).
fof(f12830,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f344,f11732,f521])).
fof(f12831,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f344,f11732,f521])).
fof(f12832,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f344,f11732,f521])).
fof(f13529,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl436),
  inference(cnf_transformation,[],[f13529_D])).
fof(f13529_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f12833,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f11732,f521])).
fof(f12834,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f11732,f521])).
fof(f12835,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f11732,f521])).
fof(f13533,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl440),
  inference(cnf_transformation,[],[f13533_D])).
fof(f13533_D,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f12836,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f345,f11732,f521])).
fof(f12837,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f345,f11732,f521])).
fof(f12838,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f345,f11732,f521])).
fof(f12839,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f344,f11732,f521])).
fof(f12840,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f344,f11732,f521])).
fof(f12841,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f344,f11732,f521])).
fof(f12842,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f11732,f521])).
fof(f12843,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f11732,f521])).
fof(f12844,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f11732,f521])).
fof(f13525,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl433),
  inference(cnf_transformation,[],[f13525_D])).
fof(f13525_D,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl433),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl433])])).
fof(f12847,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f524])).
fof(f12848,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f524])).
fof(f13521,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl429),
  inference(cnf_transformation,[],[f13521_D])).
fof(f13521_D,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl429),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl429])])).
fof(f12849,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f524])).
fof(f13523,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl431),
  inference(cnf_transformation,[],[f13523_D])).
fof(f13523_D,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl431),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl431])])).
fof(f12850,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f524])).
fof(f12851,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f524])).
fof(f12852,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f524])).
fof(f13515,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl427),
  inference(cnf_transformation,[],[f13515_D])).
fof(f13515_D,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') <=> ~$spl427),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl427])])).
fof(f12867,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f532])).
fof(f13509,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl425),
  inference(cnf_transformation,[],[f13509_D])).
fof(f13509_D,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') <=> ~$spl425),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl425])])).
fof(f12890,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f539])).
fof(f13499,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl423),
  inference(cnf_transformation,[],[f13499_D])).
fof(f13499_D,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') <=> ~$spl423),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl423])])).
fof(f12920,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f549])).
fof(f12921,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f549])).
fof(f13493,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl419),
  inference(cnf_transformation,[],[f13493_D])).
fof(f13493_D,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl419),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl419])])).
fof(f12922,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f549])).
fof(f13495,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl421),
  inference(cnf_transformation,[],[f13495_D])).
fof(f13495_D,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl421),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl421])])).
fof(f12925,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f549])).
fof(f12926,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f549])).
fof(f12927,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f549])).
fof(f13489,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl416),
  inference(cnf_transformation,[],[f13489_D])).
fof(f13489_D,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f12931,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f553])).
fof(f13485,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl415),
  inference(cnf_transformation,[],[f13485_D])).
fof(f13485_D,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') <=> ~$spl415),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl415])])).
fof(f12941,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f556])).
fof(f12942,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f556])).
fof(f13481,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl411),
  inference(cnf_transformation,[],[f13481_D])).
fof(f13481_D,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl411),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl411])])).
fof(f12943,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f556])).
fof(f13483,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl413),
  inference(cnf_transformation,[],[f13483_D])).
fof(f13483_D,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl413),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl413])])).
fof(f12944,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f556])).
fof(f12945,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f556])).
fof(f12946,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f556])).
fof(f12952,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f11732,f560])).
fof(f12953,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f11732,f560])).
fof(f12960,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f418,f11732,f562])).
fof(f12961,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f345,f1135,f11732,f562])).
fof(f12962,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f344,f1135,f11732,f562])).
fof(f12963,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f345,f1740,f11732,f562])).
fof(f12964,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f344,f1740,f11732,f562])).
fof(f12965,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f11732,f562])).
fof(f12966,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f11732,f562])).
fof(f12992,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f418,f11732,f562])).
fof(f13006,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f421,f11732,f562])).
fof(f13020,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f422,f11732,f562])).
fof(f13032,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f11732,f564])).
fof(f13033,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f11732,f564])).
fof(f13040,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f11732,f565])).
fof(f13041,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f11732,f565])).
fof(f13042,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f11732,f565])).
fof(f13054,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f11732,f565])).
fof(f13068,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f418,f11732,f566])).
fof(f13082,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f421,f11732,f566])).
fof(f13096,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f422,f11732,f566])).
fof(f13108,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f418,f11732,f566])).
fof(f13109,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f345,f1135,f11732,f566])).
fof(f13110,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f344,f1135,f11732,f566])).
fof(f13111,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f345,f1740,f11732,f566])).
fof(f13112,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f344,f1740,f11732,f566])).
fof(f13113,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f11732,f566])).
fof(f13114,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f11732,f566])).
fof(f13331,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl409),
  inference(cnf_transformation,[],[f13331_D])).
fof(f13331_D,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') <=> ~$spl409),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl409])])).
fof(f13138,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f567])).
fof(f13139,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f567])).
fof(f13325,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl405),
  inference(cnf_transformation,[],[f13325_D])).
fof(f13325_D,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl405),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl405])])).
fof(f13140,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f567])).
fof(f13327,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl407),
  inference(cnf_transformation,[],[f13327_D])).
fof(f13327_D,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl407),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl407])])).
fof(f13143,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f567])).
fof(f13144,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f567])).
fof(f13145,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f567])).
fof(f13321,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl403),
  inference(cnf_transformation,[],[f13321_D])).
fof(f13321_D,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl403),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl403])])).
fof(f13149,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f571])).
fof(f13150,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f571])).
fof(f13315,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl399),
  inference(cnf_transformation,[],[f13315_D])).
fof(f13315_D,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl399),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl399])])).
fof(f13151,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f571])).
fof(f13317,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl401),
  inference(cnf_transformation,[],[f13317_D])).
fof(f13317_D,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl401),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl401])])).
fof(f13154,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f571])).
fof(f13155,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f571])).
fof(f13156,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f571])).
fof(f13311,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl397),
  inference(cnf_transformation,[],[f13311_D])).
fof(f13311_D,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl397),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl397])])).
fof(f13160,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f575])).
fof(f13161,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f575])).
fof(f13305,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl393),
  inference(cnf_transformation,[],[f13305_D])).
fof(f13305_D,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl393),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl393])])).
fof(f13162,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f575])).
fof(f13307,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl395),
  inference(cnf_transformation,[],[f13307_D])).
fof(f13307_D,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl395),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl395])])).
fof(f13165,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f575])).
fof(f13166,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f575])).
fof(f13167,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f575])).
fof(f13301,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl391),
  inference(cnf_transformation,[],[f13301_D])).
fof(f13301_D,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') <=> ~$spl391),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl391])])).
fof(f13171,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f579])).
fof(f13295,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(forward_demodulation,[],[f13178,f422])).
fof(f13178,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f587])).
fof(f587,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | p(s(bool,X0)) | p(s(bool,X2)) | ~p(s(bool,X3))) )),
  inference(cnf_transformation,[],[f340])).
fof(f13203,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f598])).
fof(f13264,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl387),
  inference(cnf_transformation,[],[f13264_D])).
fof(f13264_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl387),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl387])])).
fof(f13207,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f930])).
fof(f930,plain,(
  ( ! [X5] : (~sP1(X5,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | ~p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f602,f453])).
fof(f13262,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl385),
  inference(cnf_transformation,[],[f13262_D])).
fof(f13262_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl385),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl385])])).
fof(f13208,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f935])).
fof(f935,plain,(
  ( ! [X10] : (~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X10) | ~p(s(bool,X10))) ) | $spl2),
  inference(resolution,[],[f602,f460])).
fof(f13260,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl383),
  inference(cnf_transformation,[],[f13260_D])).
fof(f13260_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl383),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl383])])).
fof(f13212,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f1535])).
fof(f1535,plain,(
  ( ! [X5] : (~sP1(X5,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ~p(s(bool,X5))) ) | $spl4),
  inference(resolution,[],[f604,f453])).
fof(f13258,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | $spl381),
  inference(cnf_transformation,[],[f13258_D])).
fof(f13258_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) <=> ~$spl381),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl381])])).
fof(f13213,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK20))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f1540])).
fof(f1540,plain,(
  ( ! [X10] : (~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X10) | ~p(s(bool,X10))) ) | $spl4),
  inference(resolution,[],[f604,f460])).
fof(f12644,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl377),
  inference(cnf_transformation,[],[f12644_D])).
fof(f12644_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) <=> ~$spl377),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl377])])).
fof(f12645,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,X0))))) ) | $spl378),
  inference(cnf_transformation,[],[f12645_D])).
fof(f12645_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,X0))))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f12293,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,X0)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))) ) | $spl4),
  inference(resolution,[],[f1208,f468])).
fof(f468,plain,(
  ( ! [X6,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X3,X2,X1,X0))))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28])],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X7,X6] : ! [X8,X11] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X11)))) <=> ! [X20] : (p(s(bool,happ(s(fun(X7,bool),X11),s(X7,X20)))) => ? [X21] : p(s(bool,happ(s(fun(X7,bool),X11),s(X7,happ(s(fun(X6,X7),X8),s(X6,X21))))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|0')).
fof(f12640,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl373),
  inference(cnf_transformation,[],[f12640_D])).
fof(f12640_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f12292,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(resolution,[],[f1208,f463])).
fof(f12642,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl374),
  inference(cnf_transformation,[],[f12642_D])).
fof(f12642_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f12291,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(resolution,[],[f1208,f476])).
fof(f11748,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(unit_resulting_resolution,[],[f1208,f476])).
fof(f11749,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(unit_resulting_resolution,[],[f1208,f463])).
fof(f11732,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) | $spl362),
  inference(cnf_transformation,[],[f11732_D])).
fof(f11732_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f11742,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) | $spl2),
  inference(forward_demodulation,[],[f11383,f418])).
fof(f11383,plain,(
  p(s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,'const.min.@_1'(s(fun(sK17,bool),sK20)))))) | $spl2),
  inference(resolution,[],[f605,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11,X13] : (p(s(bool,happ(s(fun(X6,bool),X11),s(X6,X13)))) => p(s(bool,happ(s(fun(X6,bool),X11),s(X6,'const.min.@_1'(s(fun(X6,bool),X11)))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.bool.SELECT_AX')).
fof(f11740,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | $spl369),
  inference(cnf_transformation,[],[f11740_D])).
fof(f11740_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f11741,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,X0))))) ) | $spl370),
  inference(cnf_transformation,[],[f11741_D])).
fof(f11741_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,X0))))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f11382,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,X0)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))))) ) | $spl2),
  inference(resolution,[],[f605,f474])).
fof(f474,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK31(X6,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f11734,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | $spl364),
  inference(cnf_transformation,[],[f11734_D])).
fof(f11734_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f11381,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | $spl2),
  inference(resolution,[],[f605,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f11736,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | $spl367),
  inference(cnf_transformation,[],[f11736_D])).
fof(f11736_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) <=> ~$spl367),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl367])])).
fof(f11380,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | $spl2),
  inference(resolution,[],[f605,f477])).
fof(f10871,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | $spl2),
  inference(unit_resulting_resolution,[],[f605,f477])).
fof(f10872,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | $spl2),
  inference(unit_resulting_resolution,[],[f605,f470])).
fof(f11730,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK20)))) | $spl2),
  inference(forward_demodulation,[],[f10873,f418])).
fof(f10873,plain,(
  p(s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,'const.min.@_1'(s(fun(sK17,bool),sK20)))))) | $spl2),
  inference(unit_resulting_resolution,[],[f605,f423])).
fof(f10283,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl360),
  inference(cnf_transformation,[],[f10283_D])).
fof(f10283_D,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f6929,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl208),
  inference(unit_resulting_resolution,[],[f345,f1740,f505])).
fof(f6930,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl208),
  inference(unit_resulting_resolution,[],[f344,f1740,f505])).
fof(f8983,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl358),
  inference(cnf_transformation,[],[f8983_D])).
fof(f8983_D,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f7571,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl208),
  inference(unit_resulting_resolution,[],[f345,f1740,f505])).
fof(f7572,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl208),
  inference(unit_resulting_resolution,[],[f344,f1740,f505])).
fof(f6381,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl356),
  inference(cnf_transformation,[],[f6381_D])).
fof(f6381_D,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f3970,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl102),
  inference(unit_resulting_resolution,[],[f345,f1135,f505])).
fof(f3971,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl102),
  inference(unit_resulting_resolution,[],[f344,f1135,f505])).
fof(f5484,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl354),
  inference(cnf_transformation,[],[f5484_D])).
fof(f5484_D,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f4451,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl102),
  inference(unit_resulting_resolution,[],[f345,f1135,f505])).
fof(f4452,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl102),
  inference(unit_resulting_resolution,[],[f344,f1135,f505])).
fof(f3780,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl353),
  inference(cnf_transformation,[],[f3780_D])).
fof(f3780_D,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl353),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl353])])).
fof(f3776,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl191),
  inference(unit_resulting_resolution,[],[f1722,f487])).
fof(f3770,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl351),
  inference(cnf_transformation,[],[f3770_D])).
fof(f3770_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl351),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl351])])).
fof(f3766,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl187),
  inference(unit_resulting_resolution,[],[f1718,f486])).
fof(f3760,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl349),
  inference(cnf_transformation,[],[f3760_D])).
fof(f3760_D,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f3756,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl185),
  inference(unit_resulting_resolution,[],[f1716,f486])).
fof(f3733,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl346),
  inference(cnf_transformation,[],[f3733_D])).
fof(f3733_D,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f3728,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl170),
  inference(resolution,[],[f1702,f508])).
fof(f3722,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1702,f508])).
fof(f3705,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl344),
  inference(cnf_transformation,[],[f3705_D])).
fof(f3705_D,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f3700,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl164),
  inference(resolution,[],[f1696,f508])).
fof(f3694,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl164),
  inference(unit_resulting_resolution,[],[f1696,f508])).
fof(f3642,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl342),
  inference(cnf_transformation,[],[f3642_D])).
fof(f3642_D,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f3640,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl140),
  inference(resolution,[],[f1672,f544])).
fof(f3634,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl140),
  inference(unit_resulting_resolution,[],[f1672,f544])).
fof(f3209,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl340),
  inference(cnf_transformation,[],[f3209_D])).
fof(f3209_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f3205,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f1278,f1285,f1286,f489])).
fof(f3206,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f1278,f1285,f1286,f489])).
fof(f3198,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl338),
  inference(cnf_transformation,[],[f3198_D])).
fof(f3198_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f3194,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f345,f684,f675,f1285,f489])).
fof(f3195,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f344,f684,f675,f1285,f489])).
fof(f3186,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl336),
  inference(cnf_transformation,[],[f3186_D])).
fof(f3186_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f3179,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f345,f683,f682,f1278,f489])).
fof(f3180,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f344,f683,f682,f1278,f489])).
fof(f3167,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl335),
  inference(cnf_transformation,[],[f3167_D])).
fof(f3167_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f3163,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl85),
  inference(unit_resulting_resolution,[],[f1117,f487])).
fof(f3152,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl333),
  inference(cnf_transformation,[],[f3152_D])).
fof(f3152_D,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl333),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl333])])).
fof(f3148,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl83),
  inference(unit_resulting_resolution,[],[f1115,f487])).
fof(f3142,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl331),
  inference(cnf_transformation,[],[f3142_D])).
fof(f3142_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl331),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl331])])).
fof(f3138,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl79),
  inference(unit_resulting_resolution,[],[f1111,f486])).
fof(f3132,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl329),
  inference(cnf_transformation,[],[f3132_D])).
fof(f3132_D,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f3128,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl77),
  inference(unit_resulting_resolution,[],[f1109,f486])).
fof(f3105,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl326),
  inference(cnf_transformation,[],[f3105_D])).
fof(f3105_D,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f3100,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl66),
  inference(resolution,[],[f1099,f508])).
fof(f3094,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl66),
  inference(unit_resulting_resolution,[],[f1099,f508])).
fof(f3077,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl324),
  inference(cnf_transformation,[],[f3077_D])).
fof(f3077_D,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f3072,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl60),
  inference(resolution,[],[f1093,f508])).
fof(f3066,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1093,f508])).
fof(f3010,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl322),
  inference(cnf_transformation,[],[f3010_D])).
fof(f3010_D,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f3008,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl36),
  inference(resolution,[],[f1069,f544])).
fof(f3002,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl36),
  inference(unit_resulting_resolution,[],[f1069,f544])).
fof(f2831,plain,(
  sP4('F_0','T_0','F_0') | $spl316),
  inference(cnf_transformation,[],[f2831_D])).
fof(f2831_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f2827,plain,(
  sP4('F_0','T_0','F_0') | $spl110),
  inference(resolution,[],[f1207,f487])).
fof(f2833,plain,(
  sP3('F_0','T_0','F_0') | $spl318),
  inference(cnf_transformation,[],[f2833_D])).
fof(f2833_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f2826,plain,(
  sP3('F_0','T_0','F_0') | $spl110),
  inference(resolution,[],[f1207,f486])).
fof(f2835,plain,(
  sP2('T_0','F_0','F_0') | $spl320),
  inference(cnf_transformation,[],[f2835_D])).
fof(f2835_D,plain,(
  sP2('T_0','F_0','F_0') <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f2825,plain,(
  sP2('T_0','F_0','F_0') | $spl110),
  inference(resolution,[],[f1207,f485])).
fof(f2821,plain,(
  sP2('T_0','F_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f1207,f485])).
fof(f2822,plain,(
  sP3('F_0','T_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f1207,f486])).
fof(f2823,plain,(
  sP4('F_0','T_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f1207,f487])).
fof(f2618,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl314),
  inference(cnf_transformation,[],[f2618_D])).
fof(f2618_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f2614,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f675,f682,f683,f489])).
fof(f2615,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f675,f682,f683,f489])).
fof(f2608,plain,(
  sP4('F_0','F_0','T_0') | $spl308),
  inference(cnf_transformation,[],[f2608_D])).
fof(f2608_D,plain,(
  sP4('F_0','F_0','T_0') <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f2605,plain,(
  sP4('F_0','F_0','T_0') | $spl106),
  inference(resolution,[],[f1203,f487])).
fof(f2610,plain,(
  sP3('T_0','F_0','F_0') | $spl310),
  inference(cnf_transformation,[],[f2610_D])).
fof(f2610_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f2604,plain,(
  sP3('T_0','F_0','F_0') | $spl106),
  inference(resolution,[],[f1203,f486])).
fof(f2612,plain,(
  sP2('F_0','T_0','F_0') | $spl312),
  inference(cnf_transformation,[],[f2612_D])).
fof(f2612_D,plain,(
  sP2('F_0','T_0','F_0') <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f2603,plain,(
  sP2('F_0','T_0','F_0') | $spl106),
  inference(resolution,[],[f1203,f485])).
fof(f2598,plain,(
  sP2('F_0','T_0','F_0') | $spl106),
  inference(unit_resulting_resolution,[],[f1203,f485])).
fof(f2599,plain,(
  sP3('T_0','F_0','F_0') | $spl106),
  inference(unit_resulting_resolution,[],[f1203,f486])).
fof(f2600,plain,(
  sP4('F_0','F_0','T_0') | $spl106),
  inference(unit_resulting_resolution,[],[f1203,f487])).
fof(f2442,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl304),
  inference(cnf_transformation,[],[f2442_D])).
fof(f2442_D,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f2201,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(resolution,[],[f600,f368])).
fof(f2200,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(resolution,[],[f600,f366])).
fof(f2484,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))) | $spl307),
  inference(cnf_transformation,[],[f2484_D])).
fof(f2484_D,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))) <=> ~$spl307),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl307])])).
fof(f1787,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK24(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f2241,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl211),
  inference(cnf_transformation,[],[f2241_D])).
fof(f2241_D,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f1790,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f361])).
fof(f1791,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f363])).
fof(f1792,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f366])).
fof(f1793,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f368])).
fof(f1822,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f431])).
fof(f1824,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f431])).
fof(f2440,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl303),
  inference(cnf_transformation,[],[f2440_D])).
fof(f2440_D,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f1826,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f437])).
fof(f2436,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl299),
  inference(cnf_transformation,[],[f2436_D])).
fof(f2436_D,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl299),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl299])])).
fof(f1827,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f437])).
fof(f2438,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl301),
  inference(cnf_transformation,[],[f2438_D])).
fof(f2438_D,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl301),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl301])])).
fof(f1828,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f437])).
fof(f1829,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f437])).
fof(f1830,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f441])).
fof(f1831,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f441])).
fof(f1832,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f441])).
fof(f1833,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f441])).
fof(f2428,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl291),
  inference(cnf_transformation,[],[f2428_D])).
fof(f2428_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl291),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl291])])).
fof(f1834,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f452])).
fof(f2426,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl289),
  inference(cnf_transformation,[],[f2426_D])).
fof(f2426_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl289),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl289])])).
fof(f1835,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f452])).
fof(f2430,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl293),
  inference(cnf_transformation,[],[f2430_D])).
fof(f2430_D,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl293),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl293])])).
fof(f1836,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f452])).
fof(f1837,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f452])).
fof(f2432,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl294),
  inference(cnf_transformation,[],[f2432_D])).
fof(f2432_D,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f1838,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f454])).
fof(f1839,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f454])).
fof(f2434,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl296),
  inference(cnf_transformation,[],[f2434_D])).
fof(f2434_D,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f1840,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f457])).
fof(f1841,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f457])).
fof(f1842,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f458])).
fof(f1843,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f458])).
fof(f1844,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f459])).
fof(f1845,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f459])).
fof(f1846,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f461])).
fof(f1847,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f461])).
fof(f1848,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f461])).
fof(f1849,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f461])).
fof(f2424,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl287),
  inference(cnf_transformation,[],[f2424_D])).
fof(f2424_D,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f1850,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f488])).
fof(f1851,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f488])).
fof(f1852,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f488])).
fof(f1853,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f488])).
fof(f2422,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl285),
  inference(cnf_transformation,[],[f2422_D])).
fof(f2422_D,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f1854,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f492])).
fof(f1855,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f492])).
fof(f1856,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f492])).
fof(f1857,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f492])).
fof(f2420,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl283),
  inference(cnf_transformation,[],[f2420_D])).
fof(f2420_D,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f1860,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f496])).
fof(f1861,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f496])).
fof(f1862,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f496])).
fof(f1863,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f496])).
fof(f2418,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl281),
  inference(cnf_transformation,[],[f2418_D])).
fof(f2418_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') <=> ~$spl281),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl281])])).
fof(f1866,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f500])).
fof(f2414,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl277),
  inference(cnf_transformation,[],[f2414_D])).
fof(f2414_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f1867,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f500])).
fof(f2410,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl273),
  inference(cnf_transformation,[],[f2410_D])).
fof(f2410_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f1868,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f500])).
fof(f2406,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl269),
  inference(cnf_transformation,[],[f2406_D])).
fof(f2406_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f1869,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f500])).
fof(f2416,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl279),
  inference(cnf_transformation,[],[f2416_D])).
fof(f2416_D,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl279),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl279])])).
fof(f1870,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f500])).
fof(f1871,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f500])).
fof(f2408,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl271),
  inference(cnf_transformation,[],[f2408_D])).
fof(f2408_D,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl271])])).
fof(f1872,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f500])).
fof(f1873,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f500])).
fof(f2412,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl275),
  inference(cnf_transformation,[],[f2412_D])).
fof(f2412_D,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f1874,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f500])).
fof(f1875,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f500])).
fof(f1876,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f500])).
fof(f1877,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f500])).
fof(f2404,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','F_0') | $spl266),
  inference(cnf_transformation,[],[f2404_D])).
fof(f2404_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','F_0') <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f1878,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','F_0') | ($spl1 | $spl29)),
  inference(unit_resulting_resolution,[],[f1003,f600,f504])).
fof(f2402,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | $spl264),
  inference(cnf_transformation,[],[f2402_D])).
fof(f2402_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f1879,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | ($spl1 | $spl29)),
  inference(unit_resulting_resolution,[],[f1003,f600,f504])).
fof(f2400,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl263),
  inference(cnf_transformation,[],[f2400_D])).
fof(f2400_D,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl263),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl263])])).
fof(f1894,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f515])).
fof(f1895,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f515])).
fof(f1896,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f515])).
fof(f1897,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f515])).
fof(f2398,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl261),
  inference(cnf_transformation,[],[f2398_D])).
fof(f2398_D,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') <=> ~$spl261),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl261])])).
fof(f1912,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f524])).
fof(f1913,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f524])).
fof(f1914,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f524])).
fof(f1915,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f524])).
fof(f2396,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl259),
  inference(cnf_transformation,[],[f2396_D])).
fof(f2396_D,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f1920,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f532])).
fof(f1921,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f532])).
fof(f2392,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl255),
  inference(cnf_transformation,[],[f2392_D])).
fof(f2392_D,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f1922,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f532])).
fof(f1923,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f532])).
fof(f2394,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl257),
  inference(cnf_transformation,[],[f2394_D])).
fof(f2394_D,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f1924,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f532])).
fof(f1925,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f532])).
fof(f1926,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f532])).
fof(f1927,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f532])).
fof(f2390,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl252),
  inference(cnf_transformation,[],[f2390_D])).
fof(f2390_D,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f1929,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f536])).
fof(f2382,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl244),
  inference(cnf_transformation,[],[f2382_D])).
fof(f2382_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f1930,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f536])).
fof(f2386,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl248),
  inference(cnf_transformation,[],[f2386_D])).
fof(f2386_D,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f1931,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f536])).
fof(f2378,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl240),
  inference(cnf_transformation,[],[f2378_D])).
fof(f2378_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f1932,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f536])).
fof(f2388,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl250),
  inference(cnf_transformation,[],[f2388_D])).
fof(f2388_D,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f1933,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f536])).
fof(f2380,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl242),
  inference(cnf_transformation,[],[f2380_D])).
fof(f2380_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f1934,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f536])).
fof(f1935,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f536])).
fof(f1936,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f536])).
fof(f2384,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl246),
  inference(cnf_transformation,[],[f2384_D])).
fof(f2384_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f1937,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f536])).
fof(f1938,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f536])).
fof(f1939,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f536])).
fof(f1940,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f536])).
fof(f2376,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl239),
  inference(cnf_transformation,[],[f2376_D])).
fof(f2376_D,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f1941,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f600,f539])).
fof(f2372,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl235),
  inference(cnf_transformation,[],[f2372_D])).
fof(f2372_D,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f1942,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f600,f539])).
fof(f1943,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f600,f539])).
fof(f1944,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f600,f539])).
fof(f2374,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl237),
  inference(cnf_transformation,[],[f2374_D])).
fof(f2374_D,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f1945,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f600,f539])).
fof(f1946,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f600,f539])).
fof(f1947,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f600,f539])).
fof(f1948,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f600,f539])).
fof(f2370,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl233),
  inference(cnf_transformation,[],[f2370_D])).
fof(f2370_D,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f1959,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f549])).
fof(f1960,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f549])).
fof(f1961,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f549])).
fof(f1962,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f549])).
fof(f2368,plain,(
  sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl230),
  inference(cnf_transformation,[],[f2368_D])).
fof(f2368_D,plain,(
  sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f1965,plain,(
  sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f600,f553])).
fof(f2364,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl226),
  inference(cnf_transformation,[],[f2364_D])).
fof(f2364_D,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f1966,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f600,f553])).
fof(f1967,plain,(
  sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f600,f553])).
fof(f1968,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f600,f553])).
fof(f2366,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | $spl228),
  inference(cnf_transformation,[],[f2366_D])).
fof(f2366_D,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f1969,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f553])).
fof(f1970,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f553])).
fof(f1971,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f553])).
fof(f1972,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f553])).
fof(f2362,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl225),
  inference(cnf_transformation,[],[f2362_D])).
fof(f2362_D,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f1975,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f556])).
fof(f1976,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f556])).
fof(f1977,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f556])).
fof(f1978,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f556])).
fof(f1985,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f600,f560])).
fof(f1986,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f600,f560])).
fof(f1987,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f421,f600,f560])).
fof(f1988,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f421,f600,f560])).
fof(f1989,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f600,f560])).
fof(f1990,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f600,f560])).
fof(f1991,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f422,f600,f560])).
fof(f1992,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f422,f600,f560])).
fof(f2002,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f421,f600,f560])).
fof(f2004,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f421,f600,f560])).
fof(f2014,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f422,f600,f560])).
fof(f2016,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f422,f600,f560])).
fof(f2025,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f421,f600,f562])).
fof(f2026,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f421,f600,f562])).
fof(f2027,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f421,f600,f562])).
fof(f2028,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f421,f600,f562])).
fof(f2029,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f422,f600,f562])).
fof(f2030,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f422,f600,f562])).
fof(f2031,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f422,f600,f562])).
fof(f2032,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f422,f600,f562])).
fof(f2041,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f600,f564])).
fof(f2042,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f600,f564])).
fof(f2043,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f421,f600,f564])).
fof(f2044,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f421,f600,f564])).
fof(f2045,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f600,f564])).
fof(f2046,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f600,f564])).
fof(f2047,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f422,f600,f564])).
fof(f2048,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f422,f600,f564])).
fof(f2058,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f421,f600,f564])).
fof(f2060,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f421,f600,f564])).
fof(f2070,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f422,f600,f564])).
fof(f2072,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f422,f600,f564])).
fof(f2081,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f345,f600,f565])).
fof(f2082,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f345,f600,f565])).
fof(f2083,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f344,f600,f565])).
fof(f2084,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f344,f600,f565])).
fof(f2095,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f345,f600,f565])).
fof(f2096,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f345,f600,f565])).
fof(f2099,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f344,f600,f565])).
fof(f2100,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f344,f600,f565])).
fof(f2117,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f421,f600,f566])).
fof(f2118,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f421,f600,f566])).
fof(f2119,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f421,f600,f566])).
fof(f2120,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f421,f600,f566])).
fof(f2121,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f422,f600,f566])).
fof(f2122,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f422,f600,f566])).
fof(f2123,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f422,f600,f566])).
fof(f2124,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f422,f600,f566])).
fof(f2272,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl223),
  inference(cnf_transformation,[],[f2272_D])).
fof(f2272_D,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f2133,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f567])).
fof(f2134,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f567])).
fof(f2135,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f567])).
fof(f2136,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f567])).
fof(f2270,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl221),
  inference(cnf_transformation,[],[f2270_D])).
fof(f2270_D,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f2139,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f571])).
fof(f2140,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f571])).
fof(f2141,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f571])).
fof(f2142,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f571])).
fof(f2268,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl219),
  inference(cnf_transformation,[],[f2268_D])).
fof(f2268_D,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f2145,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f575])).
fof(f2146,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f575])).
fof(f2147,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f575])).
fof(f2148,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f575])).
fof(f2266,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl217),
  inference(cnf_transformation,[],[f2266_D])).
fof(f2266_D,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f2151,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f579])).
fof(f2152,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f579])).
fof(f2262,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl213),
  inference(cnf_transformation,[],[f2262_D])).
fof(f2262_D,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f2153,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f579])).
fof(f2154,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f579])).
fof(f2264,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl215),
  inference(cnf_transformation,[],[f2264_D])).
fof(f2264_D,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f2155,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f579])).
fof(f2156,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f579])).
fof(f2157,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f579])).
fof(f2158,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f579])).
fof(f2182,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f600,f598])).
fof(f2183,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f600,f598])).
fof(f2190,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f598])).
fof(f2191,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f598])).
fof(f1740,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl208),
  inference(cnf_transformation,[],[f1740_D])).
fof(f1740_D,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f1781,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(global_subsumption,[],[f343,f342,f341,f604,f1519,f1566,f1515,f1514,f1513,f1512,f1511,f1510,f1509,f1508,f1507,f1506,f1505,f1504,f1503,f1502,f1501,f1500,f1499,f1498,f1497,f1496,f1495,f1494,f1493,f1492,f1479,f1478,f1469,f1458,f1447,f1438,f1437,f1436,f1431,f1430,f1421,f1410,f1395,f1394,f1389,f1388,f1387,f1386,f1385,f1384,f1383,f1382,f1381,f1380,f1379,f1376,f1374,f1373,f1372,f1371,f1370,f1369,f1368,f1367,f1366,f1364,f1363,f1362,f1361,f1360,f1357,f1355,f1354,f1351,f1349,f1348,f1347,f1344,f1343,f1342,f1341,f1340,f1339,f1338,f1337,f1336,f1335,f1334,f1333,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1325,f1324,f1323,f1322,f1321,f1320,f1319,f1318,f1317,f1316,f1315,f1314,f1313,f1312,f1311,f1310,f1309,f1308,f1307,f1306,f1305,f1304,f1303,f1298,f1297,f1296,f1295,f1294,f1293,f1292,f1291,f1290,f1289,f1288,f1287,f1286,f1285,f1284,f1283,f1282,f1281,f1280,f1279,f1278,f1277,f1276,f1275,f1274,f1273,f1272,f1271,f1270,f1269,f1268,f1267,f1266,f1265,f1264,f1263,f1262,f1261,f1260,f1259,f1258,f1257,f1256,f1255,f1254,f1253,f1252,f1251,f1250,f1249,f1248,f1244,f1243,f1241,f1240,f1742,f1744,f1746,f1212,f1211,f1210,f1209,f1208,f1522,f1523,f1530,f1535,f1536,f1537,f1540,f1547,f1550,f1551,f1553,f1554,f1555,f1561])).
fof(f1561,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(superposition,[],[f604,f354])).
fof(f1555,plain,(
  ( ! [X39,X40] : (sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X39,X40) | ~p(s(bool,X39))) ) | $spl4),
  inference(resolution,[],[f604,f555])).
fof(f1554,plain,(
  ( ! [X37,X38] : (sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X38,X37) | p(s(bool,X37))) ) | $spl4),
  inference(resolution,[],[f604,f554])).
fof(f1551,plain,(
  ( ! [X31,X32] : (sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X31,X32) | ~p(s(bool,X31))) ) | $spl4),
  inference(resolution,[],[f604,f538])).
fof(f1550,plain,(
  ( ! [X30,X29] : (sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X30,X29) | ~p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f604,f537])).
fof(f1547,plain,(
  ( ! [X24,X23] : (sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X23,X24) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl4),
  inference(resolution,[],[f604,f521])).
fof(f1530,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(resolution,[],[f604,f434])).
fof(f1746,plain,(
  ( ! [X2,X0,X5,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X5)))) ) | $spl4),
  inference(forward_demodulation,[],[f1745,f422])).
fof(f1745,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,X1),s(X0,X5)))) ) | $spl4),
  inference(forward_demodulation,[],[f1237,f422])).
fof(f1237,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X1))),s(X0,X5)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f422,f422,f604,f428])).
fof(f1744,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl4),
  inference(forward_demodulation,[],[f1743,f422])).
fof(f1743,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl4),
  inference(forward_demodulation,[],[f1238,f426])).
fof(f1238,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3))),s(X0,X6)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f426,f422,f604,f428])).
fof(f1742,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl4),
  inference(forward_demodulation,[],[f1741,f422])).
fof(f1741,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl4),
  inference(forward_demodulation,[],[f1239,f591])).
fof(f1239,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X4,X0),happ(s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(X2,happ(s(fun(X1,X2),X6),s(X1,X9))))),s(X3,happ(s(fun(X1,X3),X7),s(X1,X9))))),s(X4,happ(s(fun(X1,X4),X8),s(X1,X9))))),s(X0,X12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f591,f422,f604,f428])).
fof(f1248,plain,(
  ( ! [X0] : (sP0(X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f443])).
fof(f1249,plain,(
  ( ! [X0] : (sP0('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f444])).
fof(f1288,plain,(
  ( ! [X0,X1] : (sP2('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f503])).
fof(f1289,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f418,f604,f505])).
fof(f1290,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f421,f604,f505])).
fof(f1291,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f422,f604,f505])).
fof(f1292,plain,(
  ( ! [X2,X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f426,f604,f505])).
fof(f1293,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(fun(X0,X1),X5),s(fun(X0,X2),X6),s(fun(X0,X3),X7))),s(X0,X8)),happ(s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),happ(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(X1,happ(s(fun(X0,X1),X5),s(X0,X8))))),s(X2,happ(s(fun(X0,X2),X6),s(X0,X8))))),s(X3,happ(s(fun(X0,X3),X7),s(X0,X8)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f591,f604,f505])).
fof(f1294,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f418,f604,f505])).
fof(f1295,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f421,f604,f505])).
fof(f1296,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f422,f604,f505])).
fof(f1297,plain,(
  ( ! [X2,X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f426,f604,f505])).
fof(f1298,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),happ(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,happ(s(fun(X4,X1),X7),s(X4,X6))))),s(X0,happ(s(fun(X4,X0),X8),s(X4,X6)))),happ(s(fun(X4,bool),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7),s(fun(X4,X0),X8))),s(X4,X6)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f591,f604,f505])).
fof(f1337,plain,(
  ( ! [X0] : (~sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f522])).
fof(f1338,plain,(
  ( ! [X0] : (~sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f523])).
fof(f1348,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f533])).
fof(f1349,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f534])).
fof(f1351,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f537])).
fof(f1354,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f537])).
fof(f1355,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f537])).
fof(f1357,plain,(
  ( ! [X0] : (sP9('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f538])).
fof(f1360,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f538])).
fof(f1361,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f538])).
fof(f1363,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f540])).
fof(f1364,plain,(
  ( ! [X0] : (~sP9('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f541])).
fof(f1372,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f551])).
fof(f1374,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f554])).
fof(f1376,plain,(
  ( ! [X0] : (sP11('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f555])).
fof(f1379,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f555])).
fof(f1380,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f555])).
fof(f1387,plain,(
  ( ! [X0] : (~sP11('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f558])).
fof(f1505,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f572])).
fof(f1512,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f576])).
fof(f1514,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f580])).
fof(f1515,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f581])).
fof(f1523,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(resolution,[],[f604,f364])).
fof(f1522,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(resolution,[],[f604,f362])).
fof(f1209,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f362])).
fof(f1210,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f364])).
fof(f1565,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl113),
  inference(cnf_transformation,[],[f1565_D])).
fof(f1565_D,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f1211,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f365])).
fof(f1212,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f367])).
fof(f1240,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f434])).
fof(f1241,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f434])).
fof(f1243,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f434])).
fof(f1244,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f434])).
fof(f1732,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl201),
  inference(cnf_transformation,[],[f1732_D])).
fof(f1732_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f1250,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f453])).
fof(f1251,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f453])).
fof(f1730,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl199),
  inference(cnf_transformation,[],[f1730_D])).
fof(f1730_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f1252,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f453])).
fof(f1734,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl203),
  inference(cnf_transformation,[],[f1734_D])).
fof(f1734_D,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f1253,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f453])).
fof(f1254,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f453])).
fof(f1255,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f453])).
fof(f1736,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl204),
  inference(cnf_transformation,[],[f1736_D])).
fof(f1736_D,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f1256,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f454])).
fof(f1738,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl206),
  inference(cnf_transformation,[],[f1738_D])).
fof(f1738_D,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f1257,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f457])).
fof(f1258,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f458])).
fof(f1259,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f459])).
fof(f1260,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f460])).
fof(f1261,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f460])).
fof(f1262,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f460])).
fof(f1263,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f460])).
fof(f1264,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f460])).
fof(f1265,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f460])).
fof(f1712,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl181),
  inference(cnf_transformation,[],[f1712_D])).
fof(f1712_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f1266,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f488])).
fof(f1267,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f488])).
fof(f1726,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl195),
  inference(cnf_transformation,[],[f1726_D])).
fof(f1726_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f1268,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f488])).
fof(f1728,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl197),
  inference(cnf_transformation,[],[f1728_D])).
fof(f1728_D,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f1269,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f488])).
fof(f1270,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f488])).
fof(f1271,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f488])).
fof(f1724,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl193),
  inference(cnf_transformation,[],[f1724_D])).
fof(f1724_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl193),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl193])])).
fof(f1272,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f492])).
fof(f1273,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f492])).
fof(f1720,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl189),
  inference(cnf_transformation,[],[f1720_D])).
fof(f1720_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f1274,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f492])).
fof(f1722,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl191),
  inference(cnf_transformation,[],[f1722_D])).
fof(f1722_D,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl191),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl191])])).
fof(f1275,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f492])).
fof(f1276,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f492])).
fof(f1277,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f492])).
fof(f1718,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl187),
  inference(cnf_transformation,[],[f1718_D])).
fof(f1718_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f1279,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f496])).
fof(f1280,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f496])).
fof(f1714,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl183),
  inference(cnf_transformation,[],[f1714_D])).
fof(f1714_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f1281,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f496])).
fof(f1716,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl185),
  inference(cnf_transformation,[],[f1716_D])).
fof(f1716_D,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f1282,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f496])).
fof(f1283,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f496])).
fof(f1284,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f496])).
fof(f1299,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | ($spl4 | $spl29)),
  inference(unit_resulting_resolution,[],[f1003,f604,f506])).
fof(f1710,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl179),
  inference(cnf_transformation,[],[f1710_D])).
fof(f1710_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f1300,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | ($spl4 | $spl29)),
  inference(unit_resulting_resolution,[],[f1003,f604,f506])).
fof(f1708,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl177),
  inference(cnf_transformation,[],[f1708_D])).
fof(f1708_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f1303,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f515])).
fof(f1304,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f515])).
fof(f1704,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl173),
  inference(cnf_transformation,[],[f1704_D])).
fof(f1704_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f1305,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f515])).
fof(f1706,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl175),
  inference(cnf_transformation,[],[f1706_D])).
fof(f1706_D,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f1306,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f515])).
fof(f1307,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f515])).
fof(f1308,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f515])).
fof(f1702,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl170),
  inference(cnf_transformation,[],[f1702_D])).
fof(f1702_D,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f1310,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f521])).
fof(f1311,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f521])).
fof(f1698,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl166),
  inference(cnf_transformation,[],[f1698_D])).
fof(f1698_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f1312,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f521])).
fof(f1313,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f521])).
fof(f1314,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f521])).
fof(f1315,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f521])).
fof(f1694,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl162),
  inference(cnf_transformation,[],[f1694_D])).
fof(f1694_D,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1316,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f521])).
fof(f1317,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f521])).
fof(f1690,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl158),
  inference(cnf_transformation,[],[f1690_D])).
fof(f1690_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f1318,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f521])).
fof(f1700,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl168),
  inference(cnf_transformation,[],[f1700_D])).
fof(f1700_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1319,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f521])).
fof(f1320,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f521])).
fof(f1321,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f521])).
fof(f1322,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f521])).
fof(f1323,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f521])).
fof(f1324,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f521])).
fof(f1692,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl160),
  inference(cnf_transformation,[],[f1692_D])).
fof(f1692_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f1325,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f521])).
fof(f1326,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f521])).
fof(f1327,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f521])).
fof(f1696,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl164),
  inference(cnf_transformation,[],[f1696_D])).
fof(f1696_D,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f1328,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f521])).
fof(f1329,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f521])).
fof(f1330,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f521])).
fof(f1331,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f521])).
fof(f1332,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f521])).
fof(f1333,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f521])).
fof(f1334,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f521])).
fof(f1335,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f521])).
fof(f1336,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f521])).
fof(f1688,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl157),
  inference(cnf_transformation,[],[f1688_D])).
fof(f1688_D,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f1339,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f524])).
fof(f1340,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f524])).
fof(f1684,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl153),
  inference(cnf_transformation,[],[f1684_D])).
fof(f1684_D,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f1341,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f524])).
fof(f1686,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl155),
  inference(cnf_transformation,[],[f1686_D])).
fof(f1686_D,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f1342,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f524])).
fof(f1343,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f524])).
fof(f1344,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f524])).
fof(f1682,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl151),
  inference(cnf_transformation,[],[f1682_D])).
fof(f1682_D,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f1347,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f532])).
fof(f1680,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl149),
  inference(cnf_transformation,[],[f1680_D])).
fof(f1680_D,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f1362,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f539])).
fof(f1678,plain,(
  ~sP10('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl147),
  inference(cnf_transformation,[],[f1678_D])).
fof(f1678_D,plain,(
  ~sP10('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f1366,plain,(
  ~sP10('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f549])).
fof(f1367,plain,(
  ~sP10('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f549])).
fof(f1674,plain,(
  ~sP10('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl143),
  inference(cnf_transformation,[],[f1674_D])).
fof(f1674_D,plain,(
  ~sP10('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f1368,plain,(
  ~sP10('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f549])).
fof(f1676,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl145),
  inference(cnf_transformation,[],[f1676_D])).
fof(f1676_D,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f1369,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f549])).
fof(f1370,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f549])).
fof(f1371,plain,(
  ~sP10('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f549])).
fof(f1672,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl140),
  inference(cnf_transformation,[],[f1672_D])).
fof(f1672_D,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f1373,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f553])).
fof(f1670,plain,(
  ~sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl139),
  inference(cnf_transformation,[],[f1670_D])).
fof(f1670_D,plain,(
  ~sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f1381,plain,(
  ~sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f556])).
fof(f1382,plain,(
  ~sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f556])).
fof(f1666,plain,(
  ~sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl135),
  inference(cnf_transformation,[],[f1666_D])).
fof(f1666_D,plain,(
  ~sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f1383,plain,(
  ~sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f556])).
fof(f1668,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl137),
  inference(cnf_transformation,[],[f1668_D])).
fof(f1668_D,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f1384,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f556])).
fof(f1385,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f556])).
fof(f1386,plain,(
  ~sP11('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f556])).
fof(f1388,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f604,f560])).
fof(f1389,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f604,f560])).
fof(f1394,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f604,f562])).
fof(f1395,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f604,f562])).
fof(f1410,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f421,f604,f562])).
fof(f1421,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f422,f604,f562])).
fof(f1430,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f604,f564])).
fof(f1431,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f604,f564])).
fof(f1436,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f604,f565])).
fof(f1437,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f604,f565])).
fof(f1438,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f604,f565])).
fof(f1447,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f604,f565])).
fof(f1458,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f421,f604,f566])).
fof(f1469,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f422,f604,f566])).
fof(f1478,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f604,f566])).
fof(f1479,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f604,f566])).
fof(f1586,plain,(
  ~sP15('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl133),
  inference(cnf_transformation,[],[f1586_D])).
fof(f1586_D,plain,(
  ~sP15('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f1492,plain,(
  ~sP15('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f567])).
fof(f1493,plain,(
  ~sP15('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f567])).
fof(f1582,plain,(
  ~sP15('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl129),
  inference(cnf_transformation,[],[f1582_D])).
fof(f1582_D,plain,(
  ~sP15('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f1494,plain,(
  ~sP15('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f567])).
fof(f1584,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl131),
  inference(cnf_transformation,[],[f1584_D])).
fof(f1584_D,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f1495,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f567])).
fof(f1496,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f567])).
fof(f1497,plain,(
  ~sP15('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f567])).
fof(f1580,plain,(
  ~sP14('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl127),
  inference(cnf_transformation,[],[f1580_D])).
fof(f1580_D,plain,(
  ~sP14('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f1499,plain,(
  ~sP14('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f571])).
fof(f1500,plain,(
  ~sP14('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f571])).
fof(f1576,plain,(
  ~sP14('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl123),
  inference(cnf_transformation,[],[f1576_D])).
fof(f1576_D,plain,(
  ~sP14('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f1501,plain,(
  ~sP14('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f571])).
fof(f1578,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl125),
  inference(cnf_transformation,[],[f1578_D])).
fof(f1578_D,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f1502,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f571])).
fof(f1503,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f571])).
fof(f1504,plain,(
  ~sP14('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f571])).
fof(f1574,plain,(
  ~sP13('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl121),
  inference(cnf_transformation,[],[f1574_D])).
fof(f1574_D,plain,(
  ~sP13('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f1506,plain,(
  ~sP13('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f575])).
fof(f1507,plain,(
  ~sP13('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f575])).
fof(f1570,plain,(
  ~sP13('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl117),
  inference(cnf_transformation,[],[f1570_D])).
fof(f1570_D,plain,(
  ~sP13('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f1508,plain,(
  ~sP13('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f575])).
fof(f1572,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl119),
  inference(cnf_transformation,[],[f1572_D])).
fof(f1572_D,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f1509,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f575])).
fof(f1510,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f575])).
fof(f1511,plain,(
  ~sP13('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f575])).
fof(f1568,plain,(
  ~sP12('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl115),
  inference(cnf_transformation,[],[f1568_D])).
fof(f1568_D,plain,(
  ~sP12('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f1513,plain,(
  ~sP12('const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f579])).
fof(f1566,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(forward_demodulation,[],[f1516,f422])).
fof(f1516,plain,(
  ( ! [X0] : (s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0'))) != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f587])).
fof(f1519,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f598])).
fof(f1207,plain,(
  sP5('F_0','F_0','T_0') | $spl110),
  inference(cnf_transformation,[],[f1207_D])).
fof(f1207_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1185,plain,(
  sP5('F_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f352,f1003,f504])).
fof(f1205,plain,(
  ~sP5('T_0','F_0','T_0') | $spl109),
  inference(cnf_transformation,[],[f1205_D])).
fof(f1205_D,plain,(
  ~sP5('T_0','F_0','T_0') <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f1186,plain,(
  ~sP5('T_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f345,f1003,f506])).
fof(f1187,plain,(
  ~sP5('T_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f344,f1003,f506])).
fof(f1203,plain,(
  sP5('F_0','T_0','F_0') | $spl106),
  inference(cnf_transformation,[],[f1203_D])).
fof(f1203_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1193,plain,(
  sP5('F_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f352,f1003,f504])).
fof(f1201,plain,(
  ~sP5('T_0','T_0','F_0') | $spl105),
  inference(cnf_transformation,[],[f1201_D])).
fof(f1201_D,plain,(
  ~sP5('T_0','T_0','F_0') <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f1194,plain,(
  ~sP5('T_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f345,f1003,f506])).
fof(f1195,plain,(
  ~sP5('T_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f344,f1003,f506])).
fof(f1135,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl102),
  inference(cnf_transformation,[],[f1135_D])).
fof(f1135_D,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1176,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(global_subsumption,[],[f343,f342,f341,f602,f914,f961,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f899,f898,f897,f896,f895,f894,f893,f892,f891,f890,f889,f888,f887,f874,f873,f864,f863,f862,f853,f852,f851,f842,f841,f840,f833,f832,f831,f826,f825,f816,f815,f814,f805,f804,f803,f790,f789,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f771,f769,f768,f767,f766,f765,f764,f763,f762,f761,f759,f758,f757,f756,f755,f752,f750,f749,f746,f744,f743,f742,f739,f738,f737,f736,f735,f734,f733,f732,f731,f730,f729,f728,f727,f726,f725,f724,f723,f722,f721,f720,f719,f718,f717,f716,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f695,f694,f693,f692,f691,f690,f689,f688,f687,f686,f685,f684,f683,f682,f681,f680,f679,f678,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f665,f664,f663,f662,f661,f660,f659,f658,f657,f656,f655,f654,f653,f652,f651,f650,f649,f648,f647,f646,f645,f641,f640,f638,f637,f1137,f1139,f1141,f609,f608,f607,f606,f605,f917,f918,f925,f930,f931,f932,f935,f942,f945,f946,f948,f949,f950,f956])).
fof(f956,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(superposition,[],[f602,f354])).
fof(f950,plain,(
  ( ! [X39,X40] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X39,X40) | ~p(s(bool,X39))) ) | $spl2),
  inference(resolution,[],[f602,f555])).
fof(f949,plain,(
  ( ! [X37,X38] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X38,X37) | p(s(bool,X37))) ) | $spl2),
  inference(resolution,[],[f602,f554])).
fof(f946,plain,(
  ( ! [X31,X32] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X31,X32) | ~p(s(bool,X31))) ) | $spl2),
  inference(resolution,[],[f602,f538])).
fof(f945,plain,(
  ( ! [X30,X29] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X30,X29) | ~p(s(bool,X29))) ) | $spl2),
  inference(resolution,[],[f602,f537])).
fof(f942,plain,(
  ( ! [X24,X23] : (sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X23,X24) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl2),
  inference(resolution,[],[f602,f521])).
fof(f925,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(resolution,[],[f602,f434])).
fof(f1141,plain,(
  ( ! [X2,X0,X5,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,X1),s(X0,X5)))) ) | $spl2),
  inference(forward_demodulation,[],[f1140,f422])).
fof(f1140,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))))),s(X0,X1),s(X0,X5)))) ) | $spl2),
  inference(forward_demodulation,[],[f634,f422])).
fof(f634,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X1))),s(X0,X5)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f422,f422,f602,f428])).
fof(f1139,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl2),
  inference(forward_demodulation,[],[f1138,f422])).
fof(f1138,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl2),
  inference(forward_demodulation,[],[f635,f426])).
fof(f635,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3))),s(X0,X6)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f426,f422,f602,f428])).
fof(f1137,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl2),
  inference(forward_demodulation,[],[f1136,f422])).
fof(f1136,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl2),
  inference(forward_demodulation,[],[f636,f591])).
fof(f636,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))))),s(X0,happ(s(fun(X4,X0),happ(s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(X2,happ(s(fun(X1,X2),X6),s(X1,X9))))),s(X3,happ(s(fun(X1,X3),X7),s(X1,X9))))),s(X4,happ(s(fun(X1,X4),X8),s(X1,X9))))),s(X0,X12)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f591,f422,f602,f428])).
fof(f645,plain,(
  ( ! [X0] : (sP0(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f443])).
fof(f646,plain,(
  ( ! [X0] : (sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f444])).
fof(f685,plain,(
  ( ! [X0,X1] : (sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f503])).
fof(f686,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f418,f602,f505])).
fof(f687,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f421,f602,f505])).
fof(f688,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f422,f602,f505])).
fof(f689,plain,(
  ( ! [X2,X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f426,f602,f505])).
fof(f690,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(X0,bool),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(fun(X0,X1),X5),s(fun(X0,X2),X6),s(fun(X0,X3),X7))),s(X0,X8)),happ(s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),happ(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(X1,happ(s(fun(X0,X1),X5),s(X0,X8))))),s(X2,happ(s(fun(X0,X2),X6),s(X0,X8))))),s(X3,happ(s(fun(X0,X3),X7),s(X0,X8)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f591,f602,f505])).
fof(f691,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f418,f602,f505])).
fof(f692,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f421,f602,f505])).
fof(f693,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f422,f602,f505])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f426,f602,f505])).
fof(f695,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),happ(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,happ(s(fun(X4,X1),X7),s(X4,X6))))),s(X0,happ(s(fun(X4,X0),X8),s(X4,X6)))),happ(s(fun(X4,bool),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7),s(fun(X4,X0),X8))),s(X4,X6)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f591,f602,f505])).
fof(f732,plain,(
  ( ! [X0] : (~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f522])).
fof(f733,plain,(
  ( ! [X0] : (~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f523])).
fof(f743,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f533])).
fof(f744,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f534])).
fof(f746,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f537])).
fof(f749,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f537])).
fof(f750,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f537])).
fof(f752,plain,(
  ( ! [X0] : (sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f538])).
fof(f755,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f538])).
fof(f756,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f538])).
fof(f758,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f540])).
fof(f759,plain,(
  ( ! [X0] : (~sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f541])).
fof(f767,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f551])).
fof(f769,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f554])).
fof(f771,plain,(
  ( ! [X0] : (sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f555])).
fof(f774,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f555])).
fof(f775,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f555])).
fof(f782,plain,(
  ( ! [X0] : (~sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f558])).
fof(f900,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f572])).
fof(f907,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f576])).
fof(f909,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f580])).
fof(f910,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f581])).
fof(f918,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(resolution,[],[f602,f364])).
fof(f917,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(resolution,[],[f602,f362])).
fof(f606,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f362])).
fof(f607,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f364])).
fof(f960,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl7),
  inference(cnf_transformation,[],[f960_D])).
fof(f960_D,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f608,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f365])).
fof(f609,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f367])).
fof(f637,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f434])).
fof(f638,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f434])).
fof(f640,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f434])).
fof(f641,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f434])).
fof(f1127,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl95),
  inference(cnf_transformation,[],[f1127_D])).
fof(f1127_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f647,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f453])).
fof(f648,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f453])).
fof(f1125,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl93),
  inference(cnf_transformation,[],[f1125_D])).
fof(f1125_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f649,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f453])).
fof(f1129,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl97),
  inference(cnf_transformation,[],[f1129_D])).
fof(f1129_D,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f650,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f453])).
fof(f651,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f453])).
fof(f652,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f453])).
fof(f1131,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl98),
  inference(cnf_transformation,[],[f1131_D])).
fof(f1131_D,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f653,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f454])).
fof(f1133,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl100),
  inference(cnf_transformation,[],[f1133_D])).
fof(f1133_D,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f654,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f457])).
fof(f655,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f458])).
fof(f656,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f459])).
fof(f657,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f460])).
fof(f658,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f460])).
fof(f659,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f460])).
fof(f660,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f460])).
fof(f661,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f460])).
fof(f662,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f460])).
fof(f1123,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl91),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f663,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f488])).
fof(f664,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f488])).
fof(f1119,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl87),
  inference(cnf_transformation,[],[f1119_D])).
fof(f1119_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f665,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f488])).
fof(f1121,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl89),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f666,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f488])).
fof(f667,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f488])).
fof(f668,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f488])).
fof(f1117,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl85),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f669,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f492])).
fof(f670,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f492])).
fof(f1113,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl81),
  inference(cnf_transformation,[],[f1113_D])).
fof(f1113_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f671,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f492])).
fof(f1115,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl83),
  inference(cnf_transformation,[],[f1115_D])).
fof(f1115_D,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f672,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f492])).
fof(f673,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f492])).
fof(f674,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f492])).
fof(f1111,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl79),
  inference(cnf_transformation,[],[f1111_D])).
fof(f1111_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f676,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f496])).
fof(f677,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f496])).
fof(f1107,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl75),
  inference(cnf_transformation,[],[f1107_D])).
fof(f1107_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f678,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f496])).
fof(f1109,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl77),
  inference(cnf_transformation,[],[f1109_D])).
fof(f1109_D,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f679,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f496])).
fof(f680,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f496])).
fof(f681,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f496])).
fof(f1105,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl73),
  inference(cnf_transformation,[],[f1105_D])).
fof(f1105_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f698,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f515])).
fof(f699,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f515])).
fof(f1101,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl69),
  inference(cnf_transformation,[],[f1101_D])).
fof(f1101_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f700,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f515])).
fof(f1103,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl71),
  inference(cnf_transformation,[],[f1103_D])).
fof(f1103_D,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f701,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f515])).
fof(f702,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f515])).
fof(f703,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f515])).
fof(f1099,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl66),
  inference(cnf_transformation,[],[f1099_D])).
fof(f1099_D,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f705,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f345,f602,f521])).
fof(f706,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f345,f602,f521])).
fof(f1095,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl62),
  inference(cnf_transformation,[],[f1095_D])).
fof(f1095_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f707,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f345,f602,f521])).
fof(f708,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f344,f602,f521])).
fof(f709,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f344,f602,f521])).
fof(f710,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f344,f602,f521])).
fof(f1091,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl58),
  inference(cnf_transformation,[],[f1091_D])).
fof(f1091_D,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f711,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f602,f521])).
fof(f712,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f602,f521])).
fof(f1087,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl54),
  inference(cnf_transformation,[],[f1087_D])).
fof(f1087_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f713,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f602,f521])).
fof(f1097,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl64),
  inference(cnf_transformation,[],[f1097_D])).
fof(f1097_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f714,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f345,f602,f521])).
fof(f715,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f345,f602,f521])).
fof(f716,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f345,f602,f521])).
fof(f717,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f344,f602,f521])).
fof(f718,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f344,f602,f521])).
fof(f719,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f344,f602,f521])).
fof(f1089,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl56),
  inference(cnf_transformation,[],[f1089_D])).
fof(f1089_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f720,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f602,f521])).
fof(f721,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f602,f521])).
fof(f722,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f602,f521])).
fof(f1093,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl60),
  inference(cnf_transformation,[],[f1093_D])).
fof(f1093_D,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f723,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f345,f602,f521])).
fof(f724,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f345,f602,f521])).
fof(f725,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f345,f602,f521])).
fof(f726,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f344,f602,f521])).
fof(f727,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f344,f602,f521])).
fof(f728,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f344,f602,f521])).
fof(f729,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f602,f521])).
fof(f730,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f602,f521])).
fof(f731,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f602,f521])).
fof(f1085,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl53),
  inference(cnf_transformation,[],[f1085_D])).
fof(f1085_D,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f734,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f524])).
fof(f735,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f524])).
fof(f1081,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl49),
  inference(cnf_transformation,[],[f1081_D])).
fof(f1081_D,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f736,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f524])).
fof(f1083,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl51),
  inference(cnf_transformation,[],[f1083_D])).
fof(f1083_D,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f737,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f524])).
fof(f738,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f524])).
fof(f739,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f524])).
fof(f1079,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl47),
  inference(cnf_transformation,[],[f1079_D])).
fof(f1079_D,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f742,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f532])).
fof(f1077,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl45),
  inference(cnf_transformation,[],[f1077_D])).
fof(f1077_D,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f757,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f539])).
fof(f1075,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl43),
  inference(cnf_transformation,[],[f1075_D])).
fof(f1075_D,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f761,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f549])).
fof(f762,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f549])).
fof(f1071,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl39),
  inference(cnf_transformation,[],[f1071_D])).
fof(f1071_D,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f763,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f549])).
fof(f1073,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl41),
  inference(cnf_transformation,[],[f1073_D])).
fof(f1073_D,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f764,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f549])).
fof(f765,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f549])).
fof(f766,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f549])).
fof(f1069,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl36),
  inference(cnf_transformation,[],[f1069_D])).
fof(f1069_D,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f768,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f553])).
fof(f1067,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl35),
  inference(cnf_transformation,[],[f1067_D])).
fof(f1067_D,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f776,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f556])).
fof(f777,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f556])).
fof(f1063,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl31),
  inference(cnf_transformation,[],[f1063_D])).
fof(f1063_D,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f778,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f556])).
fof(f1065,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl33),
  inference(cnf_transformation,[],[f1065_D])).
fof(f1065_D,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f779,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f556])).
fof(f780,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f556])).
fof(f781,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f556])).
fof(f783,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f602,f560])).
fof(f784,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f602,f560])).
fof(f789,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f602,f562])).
fof(f790,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f602,f562])).
fof(f1003,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl29),
  inference(cnf_transformation,[],[f1003_D])).
fof(f1003_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f803,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f602,f562])).
fof(f804,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f602,f562])).
fof(f805,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f421,f602,f562])).
fof(f814,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f602,f562])).
fof(f815,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f602,f562])).
fof(f816,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f422,f602,f562])).
fof(f825,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f602,f564])).
fof(f826,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f602,f564])).
fof(f831,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f602,f565])).
fof(f832,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f602,f565])).
fof(f833,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f602,f565])).
fof(f840,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f602,f565])).
fof(f841,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f602,f565])).
fof(f842,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f602,f565])).
fof(f851,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f602,f566])).
fof(f852,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f602,f566])).
fof(f853,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f421,f602,f566])).
fof(f862,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f602,f566])).
fof(f863,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f602,f566])).
fof(f864,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f422,f602,f566])).
fof(f873,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f602,f566])).
fof(f874,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f602,f566])).
fof(f981,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl27),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f887,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f567])).
fof(f888,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f567])).
fof(f977,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl23),
  inference(cnf_transformation,[],[f977_D])).
fof(f977_D,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f889,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f567])).
fof(f979,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl25),
  inference(cnf_transformation,[],[f979_D])).
fof(f979_D,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f890,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f567])).
fof(f891,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f567])).
fof(f892,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f567])).
fof(f975,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl21),
  inference(cnf_transformation,[],[f975_D])).
fof(f975_D,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f894,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f571])).
fof(f895,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f571])).
fof(f971,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl17),
  inference(cnf_transformation,[],[f971_D])).
fof(f971_D,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f896,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f571])).
fof(f973,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl19),
  inference(cnf_transformation,[],[f973_D])).
fof(f973_D,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f897,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f571])).
fof(f898,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f571])).
fof(f899,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f571])).
fof(f969,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl15),
  inference(cnf_transformation,[],[f969_D])).
fof(f969_D,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f901,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f575])).
fof(f902,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f575])).
fof(f965,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl11),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f903,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f575])).
fof(f967,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl13),
  inference(cnf_transformation,[],[f967_D])).
fof(f967_D,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f904,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f575])).
fof(f905,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f575])).
fof(f906,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f575])).
fof(f963,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl9),
  inference(cnf_transformation,[],[f963_D])).
fof(f963_D,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f908,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f579])).
fof(f961,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(forward_demodulation,[],[f911,f422])).
fof(f911,plain,(
  ( ! [X0] : (s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0'))) != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f587])).
fof(f914,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f598])).
fof(f604,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(cnf_transformation,[],[f604_D])).
fof(f604_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f341,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17,sK18,sK19,sK20,sK21])],[f157])).
fof(f157,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X1,bool),X3),s(fun(X1,bool),X5),s(fun(X1,bool),X4)))))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4))))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X1,bool),X3),s(fun(X1,bool),X5),s(fun(X1,bool),X4)))))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X1,bool),X3),s(fun(X1,bool),X5),s(fun(X1,bool),X4)))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X1,bool),X3),s(fun(X1,bool),X5),s(fun(X1,bool),X4)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X7,bool),X9),s(fun(X7,bool),X11),s(fun(X7,bool),X10)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : ((p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X7,bool),X9),s(fun(X7,bool),X11),s(fun(X7,bool),X10)))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|7','thm.quantHeuristics.GUESS_RULES_COND|split|7')).
fof(f602,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) | $spl2),
  inference(cnf_transformation,[],[f602_D])).
fof(f602_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f342,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK20))))),
  inference(cnf_transformation,[],[f214])).
fof(f600,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | $spl1),
  inference(cnf_transformation,[],[f600_D])).
fof(f600_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f343,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))),
  inference(cnf_transformation,[],[f214])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 502
% Passive clauses: 1069
% Generated clauses: 21470
% Final active clauses: 496
% Final passive clauses: 708
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 85
% Trivial inequalities: 6
% Fw demodulations: 10516
% Bw demodulations: 31
% Global subsumptions: 7
% 
% Simple tautologies: 288
% Equational tautologies: 12
% Forward subsumptions: 5908
% Backward subsumptions: 17
% Fw demodulations to eq. taut.: 2600
% 
% Binary resolution: 1229
% Unit resulting resolution: 8479
% Forward superposition: 419
% Backward superposition: 413
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 3188
% SAT solver unit clauses: 6
% SAT solver binary clauses: 2558
% 
% Sat splits: 4
% 
% Memory used [KB]: 13560
% Time elapsed: 0.567 s
% ------------------------------
% ------------------------------
% Success in time 18.61 s
