Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (66544)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 172
% Passive clauses: 2662
% Generated clauses: 25411
% Final active clauses: 134
% Final passive clauses: 2026
% Input formulas: 15
% Initial clauses: 35
% 
% Duplicate literals: 1708
% Trivial inequalities: 191
% Fw subsumption resolutions: 2759
% Fw demodulations: 3177
% Bw demodulations: 20
% 
% Simple tautologies: 462
% Equational tautologies: 54
% Forward subsumptions: 9873
% 
% Binary resolution: 227
% Forward superposition: 8021
% Backward superposition: 8482
% Self superposition: 227
% Equality factoring: 147
% Equality resolution: 28
% 
% Split clauses: 889
% Split components: 1831
% SAT solver clauses: 1659
% SAT solver unit clauses: 16
% SAT solver binary clauses: 165
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 4
% 
% Sat splits: 889
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 1630
% TWLsolver calls for satisfiability: 156
% 
% Memory used [KB]: 13560
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (66562)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 131
% Passive clauses: 1039
% Generated clauses: 13941
% Final active clauses: 115
% Final passive clauses: 770
% Input formulas: 15
% Initial clauses: 35
% 
% Duplicate literals: 958
% Trivial inequalities: 110
% Fw subsumption resolutions: 991
% Fw demodulations: 1394
% Bw demodulations: 146
% Condensations: 1240
% 
% Simple tautologies: 225
% Equational tautologies: 48
% Forward subsumptions: 5272
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 321
% Forward superposition: 4374
% Backward superposition: 4610
% Self superposition: 148
% Equality factoring: 112
% Equality resolution: 21
% 
% Split clauses: 113
% Split components: 233
% SAT solver clauses: 1446
% SAT solver unit clauses: 24
% SAT solver binary clauses: 30
% 
% Sat splits: 113
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 1446
% TWLsolver calls for satisfiability: 210
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (66585)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 765
% Passive clauses: 6885
% Generated clauses: 194041
% Final active clauses: 204
% Final passive clauses: 1049
% Input formulas: 15
% Initial clauses: 35
% 
% Duplicate literals: 8851
% Trivial inequalities: 1351
% Fw subsumption resolutions: 16885
% Fw demodulations: 34345
% Bw demodulations: 4631
% Condensations: 17743
% Global subsumptions: 725
% 
% Simple tautologies: 2058
% Equational tautologies: 1238
% Forward subsumptions: 101212
% Fw demodulations to eq. taut.: 9
% Bw demodulations to eq. taut.: 83
% 
% Unit resulting resolution: 9082
% Forward superposition: 52759
% Backward superposition: 39247
% Self superposition: 724
% Equality factoring: 373
% Equality resolution: 191
% 
% Split clauses: 468
% Split components: 950
% SAT solver clauses: 9592
% SAT solver unit clauses: 229
% SAT solver binary clauses: 732
% 
% Sat splits: 468
% Sat splitting refutations: 1043
% 
% Memory used [KB]: 99785
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (66863)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 362
% Passive clauses: 29928
% Generated clauses: 92701
% Final active clauses: 234
% Final passive clauses: 26654
% Input formulas: 15
% Initial clauses: 33
% 
% Selected by SInE selection: 13
% SInE iterations: 1
% 
% Duplicate literals: 2797
% Trivial inequalities: 566
% Fw subsumption resolutions: 11393
% Fw demodulations: 5607
% Bw demodulations: 89
% 
% Simple tautologies: 880
% Equational tautologies: 344
% Forward subsumptions: 31693
% 
% Binary resolution: 4541
% Forward superposition: 32696
% Backward superposition: 32810
% Self superposition: 701
% Equality factoring: 424
% Equality resolution: 70
% 
% Split clauses: 4452
% Split components: 9227
% SAT solver clauses: 4567
% SAT solver unit clauses: 10
% SAT solver binary clauses: 85
% 
% Sat splits: 4452
% Sat splitting refutations: 28
% 
% Memory used [KB]: 50404
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (66987)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 91
% Passive clauses: 2060
% Generated clauses: 11681
% Final active clauses: 64
% Final passive clauses: 1436
% Input formulas: 15
% Initial clauses: 35
% 
% Duplicate literals: 442
% Trivial inequalities: 305
% Fw subsumption resolutions: 1616
% Fw demodulations: 25
% Bw demodulations: 452
% Global subsumptions: 8
% 
% Simple tautologies: 560
% Equational tautologies: 62
% Forward subsumptions: 6396
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 149
% Forward superposition: 4083
% Backward superposition: 3658
% Self superposition: 108
% Equality factoring: 50
% Equality resolution: 15
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 2130
% SAT solver unit clauses: 25
% SAT solver binary clauses: 91
% 
% Sat splits: 14
% Sat splitting refutations: 46
% 
% Memory used [KB]: 8443
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (67001)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 198
% Passive clauses: 4839
% Generated clauses: 18749
% Final active clauses: 98
% Final passive clauses: 2442
% Input formulas: 15
% Initial clauses: 35
% 
% Duplicate literals: 432
% Trivial inequalities: 310
% Fw subsumption resolutions: 944
% Bw subsumption resolutions: 104
% Fw demodulations: 2142
% Bw demodulations: 873
% Condensations: 815
% Global subsumptions: 1
% 
% Simple tautologies: 259
% Equational tautologies: 52
% Forward subsumptions: 8793
% Backward subsumptions: 1310
% 
% Binary resolution: 574
% Forward superposition: 6764
% Backward superposition: 3992
% Self superposition: 73
% Equality factoring: 70
% Equality resolution: 43
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 4923
% SAT solver unit clauses: 6
% SAT solver binary clauses: 217
% 
% Sat splits: 2
% Sat splitting refutations: 79
% 
% Memory used [KB]: 19317
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f1805,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f62,f99,f103,f105,f102,f107,f101,f109,f100,f111,f112,f116,f117,f121,f119,f124,f128,f131,f136,f138,f140,f141,f143,f144,f146,f177,f139,f233,f235,f237,f238,f240,f273,f236,f329,f331,f333,f334,f336,f367,f332,f344,f369,f345,f373,f396,f438,f440,f442,f443,f445,f476,f441,f499,f522,f520,f542,f541,f540,f539,f454,f453,f546,f547,f125,f549,f132,f551,f553,f126,f561,f368,f598,f600,f133,f608,f664,f667,f669,f668,f689,f693,f113,f695,f697,f114,f755,f774,f690,f1092,f1101,f135,f1212,f416,f1213,f418,f1257,f1259,f1265,f1270,f1272,f1275,f1266,f1279,f1283,f1284,f1295,f1285,f1297,f1286,f1299,f1287,f1301,f1290,f1303,f1293,f1305,f1306,f1308,f1372,f1373,f1376,f1544,f1549,f1542,f1546,f1267,f372,f1547,f1551,f1552,f1440,f419,f417,f1603,f519,f1604,f521,f1110,f545,f1648,f1652,f1650,f1654,f1656,f1663,f1668,f371,f1660,f1657,f544,f1776,f1781,f1774,f1778,f1662,f543,f1779,f421,f1783,f1785,f1294,f1786,f1801,f1787,f1296,f1788,f1803,f1804])).
fof(f1804,plain,(
  $false | ($spl62 | $spl98)),
  inference(global_subsumption,[],[f62,f1097,f1094,f1093,f1091,f1098,f1099,f546,f1662,f1789])).
fof(f1789,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | ($spl62 | $spl98)),
  inference(backward_demodulation,[],[f1662,f546])).
fof(f1099,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))))) | $spl62),
  inference(superposition,[],[f546,f66])).
fof(f66,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.DIFF_INSERT',aHL_BOOL_CASES)).
fof(f1098,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl62),
  inference(resolution,[],[f546,f73])).
fof(f73,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3)) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & ~s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7,X9,X14] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X7),s(X6,X14)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ~s(X6,X9) = s(X6,X14)))),
  file('thm.pred_set.DIFF_INSERT','thm.pred_set.IN_DELETE')).
fof(f1091,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl62),
  inference(unit_resulting_resolution,[],[f546,f73])).
fof(f1093,plain,(
  s(sK0,sK3) != s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) | $spl62),
  inference(unit_resulting_resolution,[],[f546,f74])).
fof(f74,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f1094,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f546,f80])).
fof(f80,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.DIFF_INSERT','thm.pred_set.IN_DIFF')).
fof(f1097,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f546,f78])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X9,X14,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X14),s(fun(X6,bool),X7)))))) <=> (s(X6,X9) = s(X6,X14) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))))))),
  file('thm.pred_set.DIFF_INSERT','thm.pred_set.IN_INSERT')).
fof(f1803,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))) | $spl107),
  inference(cnf_transformation,[],[f1803_D])).
fof(f1803_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f1788,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))) | ($spl61 | $spl98)),
  inference(backward_demodulation,[],[f1662,f544])).
fof(f1296,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl79),
  inference(cnf_transformation,[],[f1296_D])).
fof(f1296_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f1787,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | ($spl53 | $spl98)),
  inference(backward_demodulation,[],[f1662,f441])).
fof(f1801,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl104),
  inference(cnf_transformation,[],[f1801_D])).
fof(f1801_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1786,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | ($spl24 | $spl98)),
  inference(backward_demodulation,[],[f1662,f134])).
fof(f134,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl24),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1294,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl76),
  inference(cnf_transformation,[],[f1294_D])).
fof(f1294_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1785,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | ($spl22 | $spl98)),
  inference(backward_demodulation,[],[f1662,f133])).
fof(f1783,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl102),
  inference(cnf_transformation,[],[f1783_D])).
fof(f1783_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f421,plain,(
  p(s(bool,'F_0')) | $spl50),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1779,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl58),
  inference(superposition,[],[f522,f66])).
fof(f543,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl60),
  inference(cnf_transformation,[],[f543_D])).
fof(f543_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1662,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) | $spl98),
  inference(cnf_transformation,[],[f1662_D])).
fof(f1662_D,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1778,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl58),
  inference(resolution,[],[f522,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f1774,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) | ($spl58 | $spl61)),
  inference(unit_resulting_resolution,[],[f544,f522,f76])).
fof(f1781,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl100),
  inference(cnf_transformation,[],[f1781_D])).
fof(f1781_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1776,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | ($spl58 | $spl61)),
  inference(unit_resulting_resolution,[],[f544,f522,f81])).
fof(f81,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f544,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl61),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f1657,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl59),
  inference(unit_resulting_resolution,[],[f521,f78])).
fof(f1660,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl59),
  inference(resolution,[],[f521,f78])).
fof(f371,plain,(
  ~p(s(bool,'T_0')) | $spl43),
  inference(cnf_transformation,[],[f371_D])).
fof(f371_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1668,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) | ($spl63 | $spl96)),
  inference(forward_demodulation,[],[f1666,f1654])).
fof(f1666,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl63),
  inference(resolution,[],[f545,f75])).
fof(f75,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f1663,plain,(
  s(sK0,sK3) != s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) | $spl99),
  inference(cnf_transformation,[],[f1663_D])).
fof(f1663_D,plain,(
  s(sK0,sK3) != s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f1656,plain,(
  s(sK0,sK3) != s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)) | $spl59),
  inference(unit_resulting_resolution,[],[f521,f77])).
fof(f77,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f1654,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl96),
  inference(cnf_transformation,[],[f1654_D])).
fof(f1654_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1650,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl56),
  inference(superposition,[],[f519,f66])).
fof(f1652,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl94),
  inference(cnf_transformation,[],[f1652_D])).
fof(f1652_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1648,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | ($spl56 | $spl61)),
  inference(unit_resulting_resolution,[],[f544,f519,f81])).
fof(f545,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl63),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f1110,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl54),
  inference(superposition,[],[f81,f444])).
fof(f444,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl54),
  inference(cnf_transformation,[],[f444_D])).
fof(f444_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f521,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | $spl59),
  inference(cnf_transformation,[],[f521_D])).
fof(f521_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f1604,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | $spl52),
  inference(superposition,[],[f80,f442])).
fof(f519,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl56),
  inference(cnf_transformation,[],[f519_D])).
fof(f519_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1603,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl52),
  inference(superposition,[],[f79,f442])).
fof(f79,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f417,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f419,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | $spl48),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1440,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl20),
  inference(superposition,[],[f81,f127])).
fof(f127,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl20),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1552,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1))))),
  introduced(theory_axiom,[])).
fof(f1551,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl92),
  inference(cnf_transformation,[],[f1551_D])).
fof(f1551_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1547,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl48),
  inference(superposition,[],[f419,f66])).
fof(f372,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl44),
  inference(cnf_transformation,[],[f372_D])).
fof(f372_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1267,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) | $spl70),
  inference(cnf_transformation,[],[f1267_D])).
fof(f1267_D,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1546,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl48),
  inference(resolution,[],[f419,f76])).
fof(f1542,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) | ($spl45 | $spl48)),
  inference(unit_resulting_resolution,[],[f373,f419,f76])).
fof(f1549,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl90),
  inference(cnf_transformation,[],[f1549_D])).
fof(f1549_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1544,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | ($spl45 | $spl48)),
  inference(unit_resulting_resolution,[],[f373,f419,f81])).
fof(f1376,plain,(
  ~p(s(bool,'T_0')) | ($spl45 | $spl72)),
  inference(forward_demodulation,[],[f373,f1272])).
fof(f1373,plain,(
  ~p(s(bool,'T_0')) | ($spl49 | $spl72)),
  inference(forward_demodulation,[],[f1370,f1272])).
fof(f1370,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl49),
  inference(resolution,[],[f418,f78])).
fof(f1372,plain,(
  ~p(s(bool,'T_0')) | ($spl49 | $spl72)),
  inference(forward_demodulation,[],[f1367,f1272])).
fof(f1367,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl49),
  inference(unit_resulting_resolution,[],[f418,f78])).
fof(f1308,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl89),
  inference(cnf_transformation,[],[f1308_D])).
fof(f1308_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f1306,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | ($spl41 | $spl70)),
  inference(forward_demodulation,[],[f368,f1267])).
fof(f1305,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))) | $spl86),
  inference(cnf_transformation,[],[f1305_D])).
fof(f1305_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1293,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))) | ($spl70 | $spl72)),
  inference(backward_demodulation,[],[f1267,f1272])).
fof(f1303,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | $spl84),
  inference(cnf_transformation,[],[f1303_D])).
fof(f1303_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1290,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | ($spl68 | $spl70)),
  inference(backward_demodulation,[],[f1267,f1259])).
fof(f1301,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | $spl83),
  inference(cnf_transformation,[],[f1301_D])).
fof(f1301_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f1287,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | ($spl49 | $spl70)),
  inference(backward_demodulation,[],[f1267,f418])).
fof(f1299,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl80),
  inference(cnf_transformation,[],[f1299_D])).
fof(f1299_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1286,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | ($spl38 | $spl70)),
  inference(backward_demodulation,[],[f1267,f335])).
fof(f335,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl38),
  inference(cnf_transformation,[],[f335_D])).
fof(f335_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1297,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl78),
  inference(cnf_transformation,[],[f1297_D])).
fof(f1297_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1285,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | ($spl36 | $spl70)),
  inference(backward_demodulation,[],[f1267,f333])).
fof(f1295,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl77),
  inference(cnf_transformation,[],[f1295_D])).
fof(f1295_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f1284,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | ($spl19 | $spl70)),
  inference(backward_demodulation,[],[f1267,f125])).
fof(f1283,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))))) | $spl74),
  inference(cnf_transformation,[],[f1283_D])).
fof(f1283_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1279,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))))) | $spl40),
  inference(superposition,[],[f369,f66])).
fof(f1266,plain,(
  s(sK0,sK3) != s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) | $spl71),
  inference(cnf_transformation,[],[f1266_D])).
fof(f1266_D,plain,(
  s(sK0,sK3) != s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f1275,plain,(
  s(sK0,sK3) != s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) | $spl40),
  inference(unit_resulting_resolution,[],[f369,f74])).
fof(f1272,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2))) | $spl72),
  inference(cnf_transformation,[],[f1272_D])).
fof(f1272_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1270,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(superposition,[],[f372,f66])).
fof(f1265,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) | ($spl41 | $spl68)),
  inference(forward_demodulation,[],[f1263,f1259])).
fof(f1263,plain,(
  s(sK0,sK3) = s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl41),
  inference(resolution,[],[f368,f75])).
fof(f1259,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1))) | $spl68),
  inference(cnf_transformation,[],[f1259_D])).
fof(f1259_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1257,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(superposition,[],[f416,f66])).
fof(f418,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | $spl49),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f1213,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | $spl36),
  inference(superposition,[],[f80,f333])).
fof(f416,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl46),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1212,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl36),
  inference(superposition,[],[f79,f333])).
fof(f135,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl25),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f1101,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | ($spl9 | $spl54)),
  inference(backward_demodulation,[],[f444,f111])).
fof(f1092,plain,(
  $false | ($spl57 | $spl62)),
  inference(unit_resulting_resolution,[],[f520,f546,f73])).
fof(f690,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | $spl66),
  inference(cnf_transformation,[],[f690_D])).
fof(f690_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f774,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | $spl28),
  inference(superposition,[],[f721,f66])).
fof(f721,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) ) | $spl28),
  inference(forward_demodulation,[],[f698,f142])).
fof(f142,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f142_D])).
fof(f142_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f698,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))),s(fun(bool,bool),X0))))))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f142,f77])).
fof(f755,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | $spl54),
  inference(superposition,[],[f627,f66])).
fof(f627,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) ) | $spl54),
  inference(forward_demodulation,[],[f601,f444])).
fof(f601,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))),s(fun(bool,bool),X0))))))) ) | $spl54),
  inference(unit_resulting_resolution,[],[f444,f77])).
fof(f114,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f697,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl29),
  inference(trivial_inequality_removal,[],[f696])).
fof(f696,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl29),
  inference(superposition,[],[f143,f66])).
fof(f695,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl29),
  inference(unit_resulting_resolution,[],[f143,f66])).
fof(f113,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f693,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl3 | $spl26)),
  inference(backward_demodulation,[],[f140,f105])).
fof(f689,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | $spl12),
  inference(superposition,[],[f214,f66])).
fof(f214,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) ) | $spl12),
  inference(forward_demodulation,[],[f196,f115])).
fof(f115,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f196,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f115,f77])).
fof(f668,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f65,f421])).
fof(f65,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.DIFF_INSERT',aHL_FALSITY)).
fof(f669,plain,(
  $false | $spl50),
  inference(resolution,[],[f421,f65])).
fof(f667,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) ) | $spl64),
  inference(cnf_transformation,[],[f667_D])).
fof(f667_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f664,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) ) | $spl10),
  inference(superposition,[],[f174,f66])).
fof(f174,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))))) ) | $spl10),
  inference(forward_demodulation,[],[f153,f114])).
fof(f153,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f114,f77])).
fof(f608,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl54),
  inference(superposition,[],[f81,f444])).
fof(f133,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl22),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f600,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl55),
  inference(trivial_inequality_removal,[],[f599])).
fof(f599,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl55),
  inference(superposition,[],[f445,f66])).
fof(f598,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl55),
  inference(unit_resulting_resolution,[],[f445,f66])).
fof(f368,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl41),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f561,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl38),
  inference(superposition,[],[f81,f335])).
fof(f126,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f553,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl39),
  inference(trivial_inequality_removal,[],[f552])).
fof(f552,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl39),
  inference(superposition,[],[f336,f66])).
fof(f551,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl39),
  inference(unit_resulting_resolution,[],[f336,f66])).
fof(f132,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f549,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | ($spl9 | $spl52)),
  inference(backward_demodulation,[],[f442,f111])).
fof(f125,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl19),
  inference(cnf_transformation,[],[f125_D])).
fof(f125_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f547,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | ($spl5 | $spl36)),
  inference(backward_demodulation,[],[f333,f107])).
fof(f546,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl62),
  inference(cnf_transformation,[],[f546_D])).
fof(f546_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f453,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl22),
  inference(superposition,[],[f79,f133])).
fof(f454,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl22),
  inference(superposition,[],[f80,f133])).
fof(f539,plain,(
  $false | $spl43),
  inference(unit_resulting_resolution,[],[f97,f371])).
fof(f97,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f67])).
fof(f67,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X9] : (s(X6,X9) = s(X6,X9) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.DIFF_INSERT','thm.bool.REFL_CLAUSE')).
fof(f540,plain,(
  $false | $spl43),
  inference(resolution,[],[f371,f97])).
fof(f541,plain,(
  $false | $spl43),
  inference(resolution,[],[f371,f64])).
fof(f64,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.DIFF_INSERT',aHL_TRUTH)).
fof(f542,plain,(
  $false | $spl43),
  inference(resolution,[],[f371,f63])).
fof(f63,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.DIFF_INSERT','thm.bool.TRUTH')).
fof(f520,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl57),
  inference(cnf_transformation,[],[f520_D])).
fof(f520_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f522,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | $spl58),
  inference(cnf_transformation,[],[f522_D])).
fof(f522_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f499,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl24),
  inference(superposition,[],[f81,f134])).
fof(f441,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl53),
  inference(cnf_transformation,[],[f441_D])).
fof(f441_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f476,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl22),
  inference(global_subsumption,[],[f62,f133,f452])).
fof(f452,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl22),
  inference(superposition,[],[f72,f133])).
fof(f72,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X9] : s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))),
  file('thm.pred_set.DIFF_INSERT','thm.pred_set.EXTENSION')).
fof(f445,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl55),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f443,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | ($spl9 | $spl24)),
  inference(backward_demodulation,[],[f134,f111])).
fof(f442,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl52),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f440,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl25),
  inference(trivial_inequality_removal,[],[f439])).
fof(f439,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl25),
  inference(superposition,[],[f135,f66])).
fof(f438,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl25),
  inference(unit_resulting_resolution,[],[f135,f66])).
fof(f396,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl20),
  inference(superposition,[],[f81,f127])).
fof(f373,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl45),
  inference(cnf_transformation,[],[f373_D])).
fof(f373_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f345,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl18),
  inference(superposition,[],[f80,f126])).
fof(f369,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl40),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f344,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3)))))) | $spl18),
  inference(superposition,[],[f79,f126])).
fof(f332,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl37),
  inference(cnf_transformation,[],[f332_D])).
fof(f332_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f367,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl18),
  inference(global_subsumption,[],[f62,f126,f343])).
fof(f343,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(superposition,[],[f72,f126])).
fof(f336,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl39),
  inference(cnf_transformation,[],[f336_D])).
fof(f336_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f334,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | ($spl5 | $spl20)),
  inference(backward_demodulation,[],[f127,f107])).
fof(f333,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl36),
  inference(cnf_transformation,[],[f333_D])).
fof(f333_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f331,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl21),
  inference(trivial_inequality_removal,[],[f330])).
fof(f330,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl21),
  inference(superposition,[],[f128,f66])).
fof(f329,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f128,f66])).
fof(f236,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl33),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f273,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl14),
  inference(global_subsumption,[],[f62,f119,f253])).
fof(f253,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl14),
  inference(superposition,[],[f90,f119])).
fof(f90,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.DIFF_INSERT',aHL_EXT)).
fof(f240,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl35),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f238,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f120,f109])).
fof(f120,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f120_D])).
fof(f120_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f237,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f237_D])).
fof(f237_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f235,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(trivial_inequality_removal,[],[f234])).
fof(f234,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(superposition,[],[f121,f66])).
fof(f233,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(unit_resulting_resolution,[],[f121,f66])).
fof(f139,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl27),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f177,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl10),
  inference(global_subsumption,[],[f62,f114,f173,f174,f157])).
fof(f157,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl10),
  inference(superposition,[],[f90,f114])).
fof(f173,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'T_0'))))))) ) | $spl10),
  inference(forward_demodulation,[],[f154,f114])).
fof(f154,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'T_0'))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f114,f74])).
fof(f146,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl31),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f144,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl10 | $spl29)),
  inference(forward_demodulation,[],[f143,f114])).
fof(f143,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f143_D])).
fof(f143_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f141,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f115,f105])).
fof(f140,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl26),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f138,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl13),
  inference(trivial_inequality_removal,[],[f137])).
fof(f137,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl13),
  inference(superposition,[],[f116,f66])).
fof(f136,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl13),
  inference(unit_resulting_resolution,[],[f116,f66])).
fof(f131,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl9),
  inference(superposition,[],[f111,f66])).
fof(f128,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl21),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f124,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) | $spl5),
  inference(superposition,[],[f107,f66])).
fof(f119,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f121,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f121_D])).
fof(f121_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f117,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl7),
  inference(superposition,[],[f109,f66])).
fof(f116,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f116_D])).
fof(f116_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f112,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl3),
  inference(superposition,[],[f105,f66])).
fof(f111,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl9),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f100,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f99,f72])).
fof(f109,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f101,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f99,f90])).
fof(f107,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl5),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f102,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f99,f72])).
fof(f105,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl3),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f103,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f99,f90])).
fof(f99,plain,(
  s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) | $spl1),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f62,plain,(
  s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,sK3))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f38])).
fof(f38,plain,(
  ? [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X2))))) != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))),s(fun(X0,bool),X2)))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X2))))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X2))))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X8))))) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X7),s(X6,X9))),s(fun(X6,bool),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X8))))) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X7),s(X6,X9))),s(fun(X6,bool),X8)))),
  file('thm.pred_set.DIFF_INSERT','thm.pred_set.DIFF_INSERT')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 103
% Passive clauses: 240
% Generated clauses: 1787
% Final active clauses: 36
% Final passive clauses: 21
% Input formulas: 15
% Initial clauses: 35
% 
% Duplicate literals: 2
% Trivial inequalities: 8
% Fw demodulations: 580
% Bw demodulations: 47
% Global subsumptions: 5
% 
% Simple tautologies: 19
% Equational tautologies: 25
% Forward subsumptions: 813
% Backward subsumptions: 1
% 
% Binary resolution: 34
% Unit resulting resolution: 295
% Forward superposition: 163
% Backward superposition: 465
% 
% Split clauses: 32
% Split components: 70
% SAT solver clauses: 448
% SAT solver unit clauses: 8
% SAT solver binary clauses: 254
% 
% Sat splits: 32
% Sat splitting refutations: 8
% 
% Memory used [KB]: 2174
% Time elapsed: 0.066 s
% ------------------------------
% ------------------------------
% Success in time 18.092 s
