Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (98220)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 117
% Passive clauses: 3875
% Generated clauses: 14053
% Final active clauses: 110
% Final passive clauses: 3474
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 521
% Trivial inequalities: 211
% Fw subsumption resolutions: 1456
% Fw demodulations: 488
% Bw demodulations: 5
% 
% Simple tautologies: 297
% Equational tautologies: 90
% Forward subsumptions: 5808
% 
% Binary resolution: 282
% Factoring: 2
% Forward superposition: 4329
% Backward superposition: 6023
% Self superposition: 293
% Equality factoring: 197
% Equality resolution: 34
% 
% Split clauses: 474
% Split components: 996
% SAT solver clauses: 1054
% SAT solver unit clauses: 42
% SAT solver binary clauses: 379
% 
% Sat splits: 474
% 
% TWLsolver clauses: 1028
% TWLsolver calls for satisfiability: 139
% 
% Memory used [KB]: 11897
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (98241)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 134
% Passive clauses: 1239
% Generated clauses: 11221
% Final active clauses: 115
% Final passive clauses: 1003
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 172
% Trivial inequalities: 57
% Fw subsumption resolutions: 114
% Fw demodulations: 843
% Bw demodulations: 118
% Condensations: 357
% Global subsumptions: 12
% 
% Simple tautologies: 91
% Equational tautologies: 40
% Forward subsumptions: 6262
% Fw demodulations to eq. taut.: 5
% 
% Binary resolution: 117
% Factoring: 1
% Forward superposition: 3859
% Backward superposition: 3999
% Self superposition: 191
% Equality factoring: 50
% Equality resolution: 24
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 2160
% SAT solver unit clauses: 158
% SAT solver binary clauses: 71
% SAT solver learnt clauses: 24
% 
% Sat splits: 23
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2112
% TWLsolver calls for satisfiability: 426
% 
% Memory used [KB]: 23794
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f111739,plain,(
  $false),
  inference(subsumption_resolution,[],[f111738,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|2','thm.ringNorm.polynom_distinct|split|1')).
fof(f111738,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(backward_demodulation,[],[f111737,f109930])).
fof(f109930,plain,(
  ( ! [X59] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(sK0)),X59),s('type.quote.index',sK3)))) )),
  inference(superposition,[],[f35095,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X6] : ! [X12,X13] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X12),s('type.ringNorm.polynom'(X6),X13))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X6),X12))),s('type.ringNorm.polynom'(X6),X13)))),
  file('thm.ringNorm.polynom_distinct|split|2','ARITY_0_2_const.ringNorm.Pplus')).
fof(f35095,plain,(
  ( ! [X6,X10,X7] : (s(X6,happ(s(fun('type.quote.index',X6),X7),s('type.quote.index',sK3))) = s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X6)),X10),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(superposition,[],[f12895,f344])).
fof(f344,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6,X14] : ! [X9,X15,X16,X17,X18,X19] : s(X14,happ(s(fun('type.quote.index',X14),X15),s('type.quote.index',X9))) = s(X14,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X9))),s(fun('type.quote.index',X14),X15),s(fun(X6,X14),X16),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X17),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X18),s(fun('type.ringNorm.polynom'(X6),X14),X19)))),
  file('thm.ringNorm.polynom_distinct|split|2','thm.ringNorm.polynom_case_def|split|0')).
fof(f12895,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s(fun('type.ringNorm.polynom'(sK0),X0),X5)))) )),
  inference(superposition,[],[f347,f245])).
fof(f245,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|2','thm.ringNorm.polynom_distinct|split|2')).
fof(f347,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14,X6] : ! [X8,X7,X15,X16,X17,X18,X19] : s(X14,happ(s(fun('type.ringNorm.polynom'(X6),X14),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X18),s('type.ringNorm.polynom'(X6),X8))),s('type.ringNorm.polynom'(X6),X7))) = s(X14,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7))),s(fun('type.quote.index',X14),X15),s(fun(X6,X14),X16),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X17),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X14)),X18),s(fun('type.ringNorm.polynom'(X6),X14),X19)))),
  file('thm.ringNorm.polynom_distinct|split|2','thm.ringNorm.polynom_case_def|split|3')).
fof(f111737,plain,(
  ( ! [X60] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK3))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(sK0)),X60),s('type.quote.index',sK3)))) )),
  inference(forward_demodulation,[],[f109931,f245])).
fof(f109931,plain,(
  ( ! [X60] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.quote.index','type.ringNorm.polynom'(sK0)),X60),s('type.quote.index',sK3)))) )),
  inference(superposition,[],[f35095,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X12,X13] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X12),s('type.ringNorm.polynom'(X6),X13))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X6),X12))),s('type.ringNorm.polynom'(X6),X13)))),
  file('thm.ringNorm.polynom_distinct|split|2','ARITY_0_2_const.ringNorm.Pmult')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 488
% Passive clauses: 7372
% Generated clauses: 115455
% Final active clauses: 340
% Final passive clauses: 1918
% Input formulas: 69
% Initial clauses: 125
% Discarded non-redundant clauses: 5956
% 
% Duplicate literals: 1501
% Trivial inequalities: 672
% Fw subsumption resolutions: 3892
% Fw demodulations: 25296
% Bw demodulations: 2143
% Condensations: 2522
% Global subsumptions: 278
% 
% Simple tautologies: 447
% Equational tautologies: 591
% Forward subsumptions: 64025
% Fw demodulations to eq. taut.: 5552
% Bw demodulations to eq. taut.: 568
% 
% Unit resulting resolution: 23518
% Factoring: 5
% Forward superposition: 26468
% Backward superposition: 23094
% Self superposition: 395
% Equality factoring: 192
% Equality resolution: 74
% 
% Split clauses: 119
% Split components: 244
% SAT solver clauses: 8209
% SAT solver unit clauses: 525
% SAT solver binary clauses: 2993
% 
% Sat splits: 119
% Sat splitting refutations: 319
% 
% Memory used [KB]: 101704
% Time elapsed: 6.428 s
% ------------------------------
% ------------------------------
% Success in time 7.76 s
