Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (50042)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 2907
% Generated clauses: 19870
% Final active clauses: 121
% Final passive clauses: 2660
% Input formulas: 40
% Initial clauses: 49
% 
% Duplicate literals: 902
% Trivial inequalities: 223
% Fw subsumption resolutions: 2431
% Fw demodulations: 3494
% Bw demodulations: 1
% 
% Simple tautologies: 369
% Equational tautologies: 27
% Forward subsumptions: 8757
% 
% Binary resolution: 33
% Forward superposition: 8064
% Backward superposition: 4439
% Self superposition: 229
% Equality factoring: 88
% Equality resolution: 46
% 
% Split clauses: 529
% Split components: 1058
% SAT solver clauses: 880
% SAT solver unit clauses: 62
% SAT solver binary clauses: 243
% 
% Sat splits: 529
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 880
% TWLsolver calls for satisfiability: 115
% 
% Memory used [KB]: 14072
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (50066)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 86
% Passive clauses: 750
% Generated clauses: 11715
% Final active clauses: 72
% Final passive clauses: 577
% Input formulas: 40
% Initial clauses: 49
% 
% Duplicate literals: 618
% Trivial inequalities: 176
% Fw subsumption resolutions: 870
% Fw demodulations: 1465
% Bw demodulations: 96
% Condensations: 1352
% 
% Simple tautologies: 149
% Equational tautologies: 23
% Forward subsumptions: 5074
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 16
% Forward superposition: 3910
% Backward superposition: 2433
% Self superposition: 188
% Equality factoring: 72
% Equality resolution: 19
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 1276
% SAT solver unit clauses: 64
% SAT solver binary clauses: 39
% 
% Sat splits: 16
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1276
% TWLsolver calls for satisfiability: 290
% 
% Memory used [KB]: 13944
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f37743,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f134,f183,f181,f199,f202,f200,f715,f2193,f2138,f2196,f2209,f13488,f13494,f13511,f13513,f13523,f13343,f13490,f13576,f16832,f16838,f16839,f16872,f16878,f16675,f16834,f16931,f20065,f20069,f20067,f21047,f21048,f21049,f21050,f35449,f35455,f35456,f35494,f35507,f35508,f35301,f35451,f35570,f36721,f36727,f36728,f36767,f36780,f36781,f2194,f36569,f36723,f36843,f37583,f37603,f37584,f37585,f37586,f37587,f37604,f37588,f37589,f37590,f37611,f37591,f37613,f37592,f37614,f37593,f37615,f37594,f37607,f37617,f37595,f37618,f37600,f37596,f37606,f37619,f37597,f37598,f37610,f37601,f37609,f37739,f37740,f37742])).
fof(f37742,plain,(
  $false | ($spl7 | $spl19 | $spl25 | $spl29)),
  inference(subsumption_resolution,[],[f37741,f202])).
fof(f37741,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl19 | $spl25 | $spl29)),
  inference(forward_demodulation,[],[f37731,f19879])).
fof(f19879,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f19414,f193])).
fof(f193,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f135,f139])).
fof(f139,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_SUB_ADD2',aHL_BOOL_CASES)).
fof(f135,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_SUB_ADD2',aHL_FALSITY)).
fof(f19414,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f19413,f170])).
fof(f170,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X39,X37,X40] : ! [X41,X38,X42] : s(X39,happ(s(fun(X37,X39),'_dst_x1x0_2'(s(fun(X37,fun(X40,X39)),X41),s(X40,X38))),s(X37,X42))) = s(X39,happ(s(fun(X40,X39),happ(s(fun(X37,fun(X40,X39)),X41),s(X37,X42))),s(X40,X38)))),
  file('thm.integer.INT_SUB_ADD2','_dst_x1x0_ax')).
fof(f19413,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f19412,f184])).
fof(f184,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f142,f140])).
fof(f140,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : ! [X6] : s(X8,X6) = s(X8,'const.combin.I_1'(s(X8,X6)))),
  file('thm.integer.INT_SUB_ADD2','thm.combin.I_THM')).
fof(f142,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X8] : ! [X35] : s(X8,'const.combin.I_1'(s(X8,X35))) = s(X8,happ(s(fun(X8,X8),'const.combin.I_0'),s(X8,X35)))),
  file('thm.integer.INT_SUB_ADD2','ARITY_0_1_const.combin.I')).
fof(f19412,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))))) )),
  inference(forward_demodulation,[],[f19152,f170])).
fof(f19152,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f184,f9310])).
fof(f9310,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f9309,f184])).
fof(f9309,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f9308,f184])).
fof(f9308,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2))))))) )),
  inference(subsumption_resolution,[],[f9280,f132])).
fof(f132,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_SUB_ADD2',aHL_TRUTH)).
fof(f9280,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2))))))) )),
  inference(superposition,[],[f156,f262])).
fof(f262,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0')))) )),
  inference(unit_resulting_resolution,[],[f136,f193])).
fof(f136,plain,(
  ( ! [X0] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0'))))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0'))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),'|dollar|equals_0'),s(fun(X8,X8),'const.combin.I_0'),s(fun(X8,X8),'const.combin.I_0'))))),
  file('thm.integer.INT_SUB_ADD2','thm.quotient.IDENTITY_QUOTIENT')).
fof(f156,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | s(X1,X5) != s(X1,X6) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ! [X5,X6] : ((s(X1,X5) != s(X1,X6) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))) | s(X1,X5) = s(X1,X6))))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8,X20] : ! [X31,X32,X33] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X31),s(fun(X8,X20),X32),s(fun(X20,X8),X33)))) => ! [X6,X7] : (s(X20,X6) = s(X20,X7) <=> p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X31),s(X8,happ(s(fun(X20,X8),X33),s(X20,X6))))),s(X8,happ(s(fun(X20,X8),X33),s(X20,X7))))))))),
  file('thm.integer.INT_SUB_ADD2','thm.quotient.EQUALS_PRS')).
fof(f37731,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK0))) | ($spl19 | $spl25 | $spl29)),
  inference(backward_demodulation,[],[f37622,f37263])).
fof(f37263,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))),s('type.integer.int',sK0))) | ($spl19 | $spl25 | $spl29)),
  inference(forward_demodulation,[],[f37013,f34271])).
fof(f34271,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | $spl19),
  inference(forward_demodulation,[],[f34244,f184])).
fof(f34244,plain,(
  s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f16834,f33696,f346])).
fof(f346,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f132,f190])).
fof(f190,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f139,f139])).
fof(f33696,plain,(
  ~p(s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))))),
  inference(unit_resulting_resolution,[],[f131,f9780])).
fof(f9780,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f9779,f184])).
fof(f9779,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f9778,f184])).
fof(f9778,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(subsumption_resolution,[],[f9750,f132])).
fof(f9750,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(superposition,[],[f157,f262])).
fof(f157,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))) | s(X1,X5) = s(X1,X6)) )),
  inference(cnf_transformation,[],[f125])).
fof(f131,plain,(
  s('type.integer.int',sK0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))))),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  s('type.integer.int',sK0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f102])).
fof(f102,plain,(
  ? [X0,X1] : s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ~! [X0,X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X6),s('type.integer.int',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X6),s('type.integer.int',X7)))))),
  file('thm.integer.INT_SUB_ADD2','thm.integer.INT_SUB_ADD2')).
fof(f37013,plain,(
  s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))),s('type.integer.int',sK0))) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | ($spl25 | $spl29)),
  inference(unit_resulting_resolution,[],[f35451,f170,f36723,f3803])).
fof(f3803,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f3562,f139])).
fof(f3562,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f3466,f140])).
fof(f3466,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f3362])).
fof(f3362,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1091,f139])).
fof(f1091,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f277])).
fof(f277,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f190,f140])).
fof(f37622,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) )),
  inference(superposition,[],[f36874,f146])).
fof(f146,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X6),s('type.integer.int',X7))),s('type.integer.int',X7)))),
  file('thm.integer.INT_SUB_ADD2','thm.integer.INT_SUB_ADD')).
fof(f36874,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X3),s('type.integer.int',X2)))) )),
  inference(superposition,[],[f1765,f148])).
fof(f148,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X0))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X1)))))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X0))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X1)))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X18,X19] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X18),s('type.integer.int',X19))) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X18))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X19)))))))),
  file('thm.integer.INT_SUB_ADD2','thm.integer.int_add')).
fof(f1765,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X1))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X0)))))))) )),
  inference(superposition,[],[f148,f147])).
fof(f147,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X0)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X7))) = s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X7),s('type.pair.prod'('type.num.num','type.num.num'),X6)))),
  file('thm.integer.INT_SUB_ADD2','thm.integer.TINT_ADD_SYM')).
fof(f37740,plain,(
  $false | $spl25),
  inference(subsumption_resolution,[],[f37719,f19879])).
fof(f37719,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK0))) | $spl25),
  inference(backward_demodulation,[],[f37622,f35451])).
fof(f37739,plain,(
  $false | ($spl7 | $spl19)),
  inference(subsumption_resolution,[],[f37738,f202])).
fof(f37738,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl19),
  inference(forward_demodulation,[],[f37702,f19879])).
fof(f37702,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK0))) | $spl19),
  inference(backward_demodulation,[],[f37622,f34271])).
fof(f37609,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0'))))) | $spl44),
  inference(cnf_transformation,[],[f37609_D])).
fof(f37609_D,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0'))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f37601,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool))) ) | $spl34),
  inference(cnf_transformation,[],[f37601_D])).
fof(f37601_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f37610,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) | $spl46),
  inference(cnf_transformation,[],[f37610_D])).
fof(f37610_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f37598,plain,(
  ( ! [X50,X48,X49] : (s(bool,X48) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,'T_0') = s(bool,sK4(X49,X50,bool,bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X48))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0')))))) )),
  inference(superposition,[],[f1355,f1225])).
fof(f1225,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1)))) )),
  inference(condensation,[],[f1224])).
fof(f1224,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X20))) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(duplicate_literal_removal,[],[f876])).
fof(f876,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X20))) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f277,f277])).
fof(f1355,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,X2,bool))) )),
  inference(superposition,[],[f152,f139])).
fof(f152,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.integer.INT_SUB_ADD2',aHL_EXT)).
fof(f37597,plain,(
  ( ! [X47,X45,X46] : (s(bool,X45) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') = s(bool,sK4(X46,X47,bool,bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X45)))) )),
  inference(superposition,[],[f1355,f1225])).
fof(f37619,plain,(
  ( ! [X40] : (s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40)))) ) | $spl60),
  inference(cnf_transformation,[],[f37619_D])).
fof(f37619_D,plain,(
  ( ! [X40] : (s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f37606,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl41),
  inference(cnf_transformation,[],[f37606_D])).
fof(f37606_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f37596,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,sK4(X38,X39,bool,bool)) | s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40)))) )),
  inference(superposition,[],[f1355,f279])).
fof(f279,plain,(
  ( ! [X26,X27] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X26)))) )),
  inference(superposition,[],[f190,f184])).
fof(f37600,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl32),
  inference(cnf_transformation,[],[f37600_D])).
fof(f37600_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f37618,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X35)))) ) | $spl58),
  inference(cnf_transformation,[],[f37618_D])).
fof(f37618_D,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X35)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f37595,plain,(
  ( ! [X37,X35,X36] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'T_0') = s(bool,sK4(X36,X37,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X35)))) )),
  inference(superposition,[],[f1355,f279])).
fof(f37617,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,g(s(bool,'F_0'))))) | $spl56),
  inference(cnf_transformation,[],[f37617_D])).
fof(f37617_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,g(s(bool,'F_0'))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f37607,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl42),
  inference(cnf_transformation,[],[f37607_D])).
fof(f37607_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f37594,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X33) = s(fun(bool,bool),X34) | s(bool,'T_0') = s(bool,sK4(X33,X34,bool,bool)) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,g(s(bool,'F_0')))))) )),
  inference(superposition,[],[f1355,f279])).
fof(f37615,plain,(
  ( ! [X31] : (s(bool,X31) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31)))) ) | $spl54),
  inference(cnf_transformation,[],[f37615_D])).
fof(f37615_D,plain,(
  ( ! [X31] : (s(bool,X31) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f37593,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30) | s(bool,'T_0') = s(bool,sK4(X29,X30,bool,bool)) | s(bool,X31) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31)))) )),
  inference(superposition,[],[f1355,f278])).
fof(f278,plain,(
  ( ! [X24,X25] : (s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24)))) )),
  inference(superposition,[],[f190,f141])).
fof(f141,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X37] : ! [X38] : s(X37,X38) = s(X37,happ(s(fun(X37,X37),'_i_0_0'),s(X37,X38)))),
  file('thm.integer.INT_SUB_ADD2','_i0_ax')).
fof(f37614,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26)))) ) | $spl52),
  inference(cnf_transformation,[],[f37614_D])).
fof(f37614_D,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f37592,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X27) = s(fun(bool,bool),X28) | s(bool,'T_0') = s(bool,sK4(X27,X28,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26)))) )),
  inference(superposition,[],[f1355,f278])).
fof(f37613,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,g(s(bool,'F_0'))))) | $spl50),
  inference(cnf_transformation,[],[f37613_D])).
fof(f37613_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,g(s(bool,'F_0'))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f37591,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,'T_0') = s(bool,sK4(X24,X25,bool,bool)) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,g(s(bool,'F_0')))))) )),
  inference(superposition,[],[f1355,f278])).
fof(f37611,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) ) | $spl48),
  inference(cnf_transformation,[],[f37611_D])).
fof(f37611_D,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f37590,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,'T_0') = s(bool,sK4(X20,X21,bool,bool)) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f1355,f277])).
fof(f37589,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK4(X18,X19,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) )),
  inference(superposition,[],[f1355,f277])).
fof(f37588,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') = s(bool,sK4(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0')))))) )),
  inference(superposition,[],[f1355,f277])).
fof(f37604,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl38),
  inference(cnf_transformation,[],[f37604_D])).
fof(f37604_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f37587,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1355,f190])).
fof(f37586,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK4(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1355,f190])).
fof(f37585,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1355,f190])).
fof(f37584,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK4(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1355,f190])).
fof(f37603,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl37),
  inference(cnf_transformation,[],[f37603_D])).
fof(f37603_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f37583,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1355,f139])).
fof(f36843,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))))))) ) | $spl30),
  inference(cnf_transformation,[],[f36843_D])).
fof(f36843_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f36723,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | $spl29),
  inference(cnf_transformation,[],[f36723_D])).
fof(f36723_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f36569,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))))))) )),
  inference(equality_factoring,[],[f34337])).
fof(f34337,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f34308,f132])).
fof(f34308,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f33923,f190])).
fof(f33923,plain,(
  ~p(s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))))),
  inference(forward_demodulation,[],[f33701,f170])).
fof(f33701,plain,(
  ~p(s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))),s('type.integer.int',sK0))))),
  inference(unit_resulting_resolution,[],[f131,f9780])).
fof(f2194,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) | $spl8),
  inference(cnf_transformation,[],[f2194_D])).
fof(f2194_D,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f36781,plain,(
  ( ! [X249] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X249)) )),
  inference(subsumption_resolution,[],[f36429,f16219])).
fof(f16219,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f16014,f140])).
fof(f16014,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(condensation,[],[f16013])).
fof(f16013,plain,(
  ( ! [X92,X93] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X93)) )),
  inference(duplicate_literal_removal,[],[f15763])).
fof(f15763,plain,(
  ( ! [X92,X93] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f1103,f304])).
fof(f304,plain,(
  ( ! [X19,X20] : (s(bool,X20) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X19))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f190,f184])).
fof(f1103,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f1085])).
fof(f1085,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f277])).
fof(f36429,plain,(
  ( ! [X249] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X249) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X249)))) )),
  inference(superposition,[],[f16219,f34337])).
fof(f36780,plain,(
  ( ! [X242] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X242)) )),
  inference(subsumption_resolution,[],[f36424,f12717])).
fof(f12717,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f12550,f140])).
fof(f12550,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(condensation,[],[f12549])).
fof(f12549,plain,(
  ( ! [X87,X86] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87)) )),
  inference(duplicate_literal_removal,[],[f12326])).
fof(f12326,plain,(
  ( ! [X87,X86] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f1103,f303])).
fof(f303,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f190,f141])).
fof(f36424,plain,(
  ( ! [X242] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X242) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X242)))) )),
  inference(superposition,[],[f12717,f34337])).
fof(f36767,plain,(
  ( ! [X144] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X144)) )),
  inference(subsumption_resolution,[],[f36399,f1996])).
fof(f1996,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f1924,f140])).
fof(f1924,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(condensation,[],[f1923])).
fof(f1923,plain,(
  ( ! [X41,X42] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42)) )),
  inference(duplicate_literal_removal,[],[f1851])).
fof(f1851,plain,(
  ( ! [X41,X42] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f1103,f302])).
fof(f302,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f190,f140])).
fof(f36399,plain,(
  ( ! [X144] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X144)))) )),
  inference(superposition,[],[f1996,f34337])).
fof(f36728,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f36355,f16219])).
fof(f36355,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f1336,f34337])).
fof(f1336,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X30)))) )),
  inference(forward_demodulation,[],[f1293,f140])).
fof(f1293,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X30)))) )),
  inference(superposition,[],[f1103,f184])).
fof(f36727,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f36354,f12717])).
fof(f36354,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f1335,f34337])).
fof(f1335,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29)))) )),
  inference(forward_demodulation,[],[f1292,f140])).
fof(f1292,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29)))) )),
  inference(superposition,[],[f1103,f141])).
fof(f36721,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(subsumption_resolution,[],[f36352,f1996])).
fof(f36352,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(superposition,[],[f1262,f34337])).
fof(f1262,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X37)))) ) | $spl7),
  inference(subsumption_resolution,[],[f1012,f202])).
fof(f1012,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X37)))) ) | $spl7),
  inference(superposition,[],[f202,f277])).
fof(f35570,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))))))) ) | $spl26),
  inference(cnf_transformation,[],[f35570_D])).
fof(f35570_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f35451,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | $spl25),
  inference(cnf_transformation,[],[f35451_D])).
fof(f35451_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f35301,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1)))))))) )),
  inference(equality_factoring,[],[f34275])).
fof(f34275,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f34246,f132])).
fof(f34246,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f33696,f190])).
fof(f35508,plain,(
  ( ! [X248] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X248)) )),
  inference(subsumption_resolution,[],[f35160,f16219])).
fof(f35160,plain,(
  ( ! [X248] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X248) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X248)))) )),
  inference(superposition,[],[f16219,f34275])).
fof(f35507,plain,(
  ( ! [X241] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X241)) )),
  inference(subsumption_resolution,[],[f35155,f12717])).
fof(f35155,plain,(
  ( ! [X241] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X241) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X241)))) )),
  inference(superposition,[],[f12717,f34275])).
fof(f35494,plain,(
  ( ! [X143] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f35130,f1996])).
fof(f35130,plain,(
  ( ! [X143] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X143) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X143)))) )),
  inference(superposition,[],[f1996,f34275])).
fof(f35456,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f35087,f16219])).
fof(f35087,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f1336,f34275])).
fof(f35455,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f35086,f12717])).
fof(f35086,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f1335,f34275])).
fof(f35449,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(subsumption_resolution,[],[f35084,f1996])).
fof(f35084,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK0),s('type.integer.int',sK1))))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(superposition,[],[f1262,f34275])).
fof(f21050,plain,(
  $false | ($spl8 | $spl19)),
  inference(subsumption_resolution,[],[f20306,f2194])).
fof(f20306,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | ($spl8 | $spl19)),
  inference(backward_demodulation,[],[f2194,f16834])).
fof(f21049,plain,(
  $false | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f20286,f2194])).
fof(f20286,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | ($spl8 | $spl15)),
  inference(backward_demodulation,[],[f2194,f13490])).
fof(f21048,plain,(
  $false | ($spl8 | $spl11)),
  inference(subsumption_resolution,[],[f20172,f2194])).
fof(f20172,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ($spl8 | $spl11)),
  inference(backward_demodulation,[],[f2194,f2196])).
fof(f21047,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f20155,f132])).
fof(f20155,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f2194,f135])).
fof(f20067,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | $spl22),
  inference(cnf_transformation,[],[f20067_D])).
fof(f20067_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f20069,plain,(
  ( ! [X90] : (s(bool,'T_0') = s(bool,X90) | p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f20068,f140])).
fof(f20068,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X90)))) )),
  inference(forward_demodulation,[],[f19914,f178])).
fof(f178,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X39,X40,X43,X37] : ! [X41,X38,X44,X42] : s(X39,happ(s(fun(X37,X39),'_dst_x0x1x1_3'(s(fun(X40,fun(X43,X39)),X41),s(fun(X37,X40),X38),s(fun(X37,X43),X44))),s(X37,X42))) = s(X39,happ(s(fun(X43,X39),happ(s(fun(X40,fun(X43,X39)),X41),s(X40,happ(s(fun(X37,X40),X38),s(X37,X42))))),s(X43,happ(s(fun(X37,X43),X44),s(X37,X42)))))),
  file('thm.integer.INT_SUB_ADD2','_dst_x0x1x1_ax')).
fof(f19914,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))))),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X90)))) )),
  inference(superposition,[],[f19414,f16757])).
fof(f16757,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) )),
  inference(condensation,[],[f16756])).
fof(f16756,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(duplicate_literal_removal,[],[f16349])).
fof(f16349,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f16083,f277])).
fof(f16083,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f16082])).
fof(f16082,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f15633])).
fof(f15633,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f304,f139])).
fof(f20065,plain,(
  ( ! [X89] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(forward_demodulation,[],[f19913,f178])).
fof(f19913,plain,(
  ( ! [X89] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))))),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f19414,f16083])).
fof(f16931,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) | $spl20),
  inference(cnf_transformation,[],[f16931_D])).
fof(f16931_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f16834,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | $spl19),
  inference(cnf_transformation,[],[f16834_D])).
fof(f16834_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f16675,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) )),
  inference(equality_factoring,[],[f16083])).
fof(f16878,plain,(
  ( ! [X186] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X186)) )),
  inference(subsumption_resolution,[],[f16546,f12717])).
fof(f16546,plain,(
  ( ! [X186] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X186) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X186)))) )),
  inference(superposition,[],[f12717,f16083])).
fof(f16872,plain,(
  ( ! [X137] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X137)) )),
  inference(subsumption_resolution,[],[f16531,f1996])).
fof(f16531,plain,(
  ( ! [X137] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X137) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X137)))) )),
  inference(superposition,[],[f1996,f16083])).
fof(f16839,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f16493,f16219])).
fof(f16493,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f1336,f16083])).
fof(f16838,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f16492,f12717])).
fof(f16492,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f1335,f16083])).
fof(f16832,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(subsumption_resolution,[],[f16490,f1996])).
fof(f16490,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | $spl7),
  inference(superposition,[],[f1262,f16083])).
fof(f13576,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | $spl16),
  inference(cnf_transformation,[],[f13576_D])).
fof(f13576_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f13490,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl15),
  inference(cnf_transformation,[],[f13490_D])).
fof(f13490_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f13343,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) )),
  inference(equality_factoring,[],[f12611])).
fof(f12611,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f12610])).
fof(f12610,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f12209])).
fof(f12209,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f303,f139])).
fof(f13523,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f13215,f1996])).
fof(f13215,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X129)))) )),
  inference(superposition,[],[f1996,f12611])).
fof(f13513,plain,(
  ( ! [X108] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(forward_demodulation,[],[f13512,f184])).
fof(f13512,plain,(
  ( ! [X108] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f13203,f1996])).
fof(f13203,plain,(
  ( ! [X108] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X108))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f11363,f12611])).
fof(f11363,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(condensation,[],[f11362])).
fof(f11362,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(duplicate_literal_removal,[],[f11257])).
fof(f11257,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f10590,f277])).
fof(f10590,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f10572])).
fof(f10572,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f279])).
fof(f13511,plain,(
  ( ! [X107] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X107)) )),
  inference(forward_demodulation,[],[f13357,f184])).
fof(f13357,plain,(
  ( ! [X107] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X107)) )),
  inference(duplicate_literal_removal,[],[f13202])).
fof(f13202,plain,(
  ( ! [X107] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f10590,f12611])).
fof(f13494,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f13180,f12717])).
fof(f13180,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X70))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1335,f12611])).
fof(f13488,plain,(
  ( ! [X68] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X68)) ) | $spl7),
  inference(subsumption_resolution,[],[f13178,f1996])).
fof(f13178,plain,(
  ( ! [X68] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X68))) | s(bool,'T_0') = s(bool,X68)) ) | $spl7),
  inference(superposition,[],[f1262,f12611])).
fof(f2209,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl12),
  inference(cnf_transformation,[],[f2209_D])).
fof(f2209_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2196,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl11),
  inference(cnf_transformation,[],[f2196_D])).
fof(f2196_D,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f2138,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) )),
  inference(equality_factoring,[],[f1960])).
fof(f1960,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f1959])).
fof(f1959,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f1796])).
fof(f1796,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f302,f139])).
fof(f2193,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30)) ) | $spl7),
  inference(subsumption_resolution,[],[f2085,f1996])).
fof(f2085,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) ) | $spl7),
  inference(superposition,[],[f1262,f1960])).
fof(f715,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f714,f132])).
fof(f714,plain,(
  ~p(s(bool,'T_0')) | $spl1),
  inference(backward_demodulation,[],[f703,f188])).
fof(f188,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK2),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK2),s('type.pair.prod'('type.num.num','type.num.num'),sK3)))))) | $spl1),
  inference(forward_demodulation,[],[f187,f147])).
fof(f187,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3),s('type.pair.prod'('type.num.num','type.num.num'),sK2))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK2),s('type.pair.prod'('type.num.num','type.num.num'),sK3)))))) | $spl1),
  inference(forward_demodulation,[],[f186,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X35,X36] : s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X35),s('type.pair.prod'('type.num.num','type.num.num'),X36))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X35))),s('type.pair.prod'('type.num.num','type.num.num'),X36)))),
  file('thm.integer.INT_SUB_ADD2','ARITY_0_2_const.integer.tint_add')).
fof(f186,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK2))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK2),s('type.pair.prod'('type.num.num','type.num.num'),sK3)))))) | $spl1),
  inference(forward_demodulation,[],[f185,f150])).
fof(f185,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK2))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK2))),s('type.pair.prod'('type.num.num','type.num.num'),sK3)))))) | $spl1),
  inference(forward_demodulation,[],[f181,f149])).
fof(f149,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X35,X36] : s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X35),s('type.pair.prod'('type.num.num','type.num.num'),X36))) = s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),X35))),s('type.pair.prod'('type.num.num','type.num.num'),X36)))),
  file('thm.integer.INT_SUB_ADD2','ARITY_0_2_const.integer.tint_eq')).
fof(f703,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X0)))) )),
  inference(unit_resulting_resolution,[],[f683,f193])).
fof(f683,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X0))))) )),
  inference(forward_demodulation,[],[f669,f184])).
fof(f669,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'const.combin.I_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X0))))) )),
  inference(unit_resulting_resolution,[],[f184,f151])).
fof(f151,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | s('type.pair.prod'('type.num.num','type.num.num'),X0) != s('type.pair.prod'('type.num.num','type.num.num'),X1)) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),X0) != s('type.pair.prod'('type.num.num','type.num.num'),X1) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),X0) = s('type.pair.prod'('type.num.num','type.num.num'),X1) => p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : (s('type.pair.prod'('type.num.num','type.num.num'),X12) = s('type.pair.prod'('type.num.num','type.num.num'),X13) => p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X12),s('type.pair.prod'('type.num.num','type.num.num'),X13)))))),
  file('thm.integer.INT_SUB_ADD2','thm.integer.TINT_EQ_AP')).
fof(f200,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f202,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(cnf_transformation,[],[f202_D])).
fof(f202_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f199,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f139])).
fof(f181,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK2))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK2))),s('type.pair.prod'('type.num.num','type.num.num'),sK3)))))) | $spl1),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK2))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK2))),s('type.pair.prod'('type.num.num','type.num.num'),sK3)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f183,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0')))))))) | $spl2),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0')))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f134,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0')))))))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK2))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK2))),s('type.pair.prod'('type.num.num','type.num.num'),sK3))))))),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK2))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK2))),s('type.pair.prod'('type.num.num','type.num.num'),sK3)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f103])).
fof(f103,plain,(
  ? [X0,X1] : ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X1))),s('type.pair.prod'('type.num.num','type.num.num'),X0))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'))))))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X1))),s('type.pair.prod'('type.num.num','type.num.num'),X0))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))))) => p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'))))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7,X6] : p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X6))),s('type.pair.prod'('type.num.num','type.num.num'),X7))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X7))),s('type.pair.prod'('type.num.num','type.num.num'),X6)))))) => p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'))))))))),
  file('thm.integer.INT_SUB_ADD2','thm.integer.integer_66_prove_ltb')).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 156
% Passive clauses: 2121
% Generated clauses: 38095
% Final active clauses: 99
% Final passive clauses: 1128
% Input formulas: 40
% Initial clauses: 49
% 
% Duplicate literals: 1635
% Trivial inequalities: 302
% Fw subsumption resolutions: 3057
% Fw demodulations: 7159
% Bw demodulations: 864
% Condensations: 3624
% 
% Simple tautologies: 315
% Equational tautologies: 241
% Forward subsumptions: 19540
% Fw demodulations to eq. taut.: 52
% Bw demodulations to eq. taut.: 26
% 
% Unit resulting resolution: 1710
% Forward superposition: 10738
% Backward superposition: 7715
% Self superposition: 290
% Equality factoring: 103
% Equality resolution: 24
% 
% Split clauses: 48
% Split components: 112
% SAT solver clauses: 2227
% SAT solver unit clauses: 237
% SAT solver binary clauses: 806
% 
% Sat splits: 48
% Sat splitting refutations: 8
% 
% Memory used [KB]: 24050
% Time elapsed: 2.154 s
% ------------------------------
% ------------------------------
% Success in time 3.463 s
