Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (99418)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 647
% Passive clauses: 10632
% Generated clauses: 18738
% Final active clauses: 639
% Final passive clauses: 9573
% Input formulas: 41
% Initial clauses: 166
% 
% Duplicate literals: 277
% Trivial inequalities: 59
% Fw subsumption resolutions: 1480
% Fw demodulations: 181
% Bw demodulations: 5
% 
% Simple tautologies: 206
% Equational tautologies: 56
% Forward subsumptions: 5597
% 
% Binary resolution: 11956
% Factoring: 1020
% Forward superposition: 1591
% Backward superposition: 1790
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 20
% 
% Split clauses: 314
% Split components: 630
% SAT solver clauses: 668
% SAT solver unit clauses: 20
% SAT solver binary clauses: 184
% 
% Sat splits: 314
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 668
% TWLsolver calls for satisfiability: 409
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (99431)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 404
% Passive clauses: 3054
% Generated clauses: 11217
% Final active clauses: 396
% Final passive clauses: 2625
% Input formulas: 41
% Initial clauses: 166
% 
% Duplicate literals: 167
% Trivial inequalities: 20
% Fw subsumption resolutions: 940
% Fw demodulations: 309
% Bw demodulations: 29
% Condensations: 254
% Global subsumptions: 13
% 
% Simple tautologies: 209
% Equational tautologies: 16
% Forward subsumptions: 2829
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 3589
% Factoring: 197
% Forward superposition: 1434
% Backward superposition: 981
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 58
% Split components: 121
% SAT solver clauses: 3586
% SAT solver unit clauses: 175
% SAT solver binary clauses: 82
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 58
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 3534
% TWLsolver calls for satisfiability: 1568
% 
% Memory used [KB]: 13944
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% (99472)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1683
% Passive clauses: 6067
% Generated clauses: 67070
% Final active clauses: 723
% Final passive clauses: 245
% Input formulas: 41
% Initial clauses: 166
% 
% Duplicate literals: 3519
% Trivial inequalities: 187
% Fw subsumption resolutions: 3031
% Fw demodulations: 5568
% Bw demodulations: 2652
% Condensations: 5191
% Global subsumptions: 3661
% 
% Simple tautologies: 180
% Equational tautologies: 750
% Forward subsumptions: 33841
% Bw demodulations to eq. taut.: 151
% 
% Unit resulting resolution: 15974
% Factoring: 1
% Forward superposition: 12706
% Backward superposition: 7974
% Self superposition: 167
% Equality factoring: 140
% Equality resolution: 5
% 
% Split clauses: 299
% Split components: 605
% SAT solver clauses: 14570
% SAT solver unit clauses: 425
% SAT solver binary clauses: 5418
% 
% Sat splits: 299
% Sat splitting refutations: 4241
% 
% Memory used [KB]: 97738
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quantHeuristics
% (99880)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2246
% Passive clauses: 39584
% Generated clauses: 208982
% Final active clauses: 1969
% Final passive clauses: 33928
% Input formulas: 41
% Initial clauses: 151
% 
% Selected by SInE selection: 34
% SInE iterations: 1
% 
% Duplicate literals: 8858
% Trivial inequalities: 636
% Fw subsumption resolutions: 26319
% Fw demodulations: 4924
% Bw demodulations: 126
% 
% Simple tautologies: 9208
% Equational tautologies: 968
% Forward subsumptions: 110117
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 151901
% Factoring: 278
% Forward superposition: 11306
% Backward superposition: 3125
% Self superposition: 11
% Equality resolution: 76
% 
% Split clauses: 10899
% Split components: 24987
% SAT solver clauses: 11826
% SAT solver unit clauses: 52
% SAT solver binary clauses: 4255
% 
% Sat splits: 10899
% Sat splitting refutations: 60
% 
% Memory used [KB]: 52579
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.quantHeuristics
% (100018)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 347
% Passive clauses: 3273
% Generated clauses: 17109
% Final active clauses: 339
% Final passive clauses: 2827
% Input formulas: 41
% Initial clauses: 166
% 
% Duplicate literals: 548
% Trivial inequalities: 134
% Fw subsumption resolutions: 1646
% Fw demodulations: 231
% Bw demodulations: 48
% Global subsumptions: 13
% 
% Simple tautologies: 399
% Equational tautologies: 63
% Forward subsumptions: 10692
% 
% Unit resulting resolution: 2245
% Factoring: 14
% Forward superposition: 5138
% Backward superposition: 6732
% Self superposition: 110
% Equality factoring: 6
% Equality resolution: 13
% 
% Split clauses: 18
% Split components: 37
% SAT solver clauses: 3376
% SAT solver unit clauses: 210
% SAT solver binary clauses: 187
% 
% Sat splits: 18
% Sat splitting refutations: 69
% 
% Memory used [KB]: 11769
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.quantHeuristics
% (100045)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1007
% Passive clauses: 5859
% Generated clauses: 26519
% Final active clauses: 472
% Final passive clauses: 3267
% Input formulas: 41
% Initial clauses: 166
% 
% Duplicate literals: 1703
% Trivial inequalities: 175
% Fw subsumption resolutions: 2212
% Bw subsumption resolutions: 429
% Fw demodulations: 2054
% Bw demodulations: 856
% Condensations: 2244
% Global subsumptions: 5
% 
% Simple tautologies: 2259
% Equational tautologies: 60
% Forward subsumptions: 9773
% Backward subsumptions: 940
% 
% Binary resolution: 9282
% Factoring: 25
% Forward superposition: 2500
% Backward superposition: 3008
% Self superposition: 44
% Equality factoring: 12
% Equality resolution: 20
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 6081
% SAT solver unit clauses: 133
% SAT solver binary clauses: 970
% 
% Sat splits: 3
% Sat splitting refutations: 211
% 
% Memory used [KB]: 12792
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f3774,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f236,f405,f235,f407,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f585,f548,f547,f546,f545,f538,f537,f536,f535,f587,f532,f531,f530,f529,f591,f528,f527,f526,f589,f525,f593,f522,f521,f520,f519,f595,f508,f507,f506,f505,f599,f504,f503,f502,f597,f501,f601,f500,f499,f498,f497,f609,f496,f495,f494,f605,f493,f613,f492,f603,f491,f611,f490,f607,f489,f615,f487,f486,f485,f484,f619,f483,f482,f617,f481,f480,f621,f473,f472,f471,f470,f629,f469,f468,f625,f467,f466,f633,f465,f623,f464,f627,f463,f631,f462,f635,f459,f458,f457,f456,f637,f453,f452,f451,f450,f639,f449,f448,f447,f446,f641,f445,f444,f443,f442,f441,f440,f439,f438,f437,f436,f651,f435,f434,f649,f433,f432,f647,f431,f643,f430,f645,f429,f428,f427,f426,f425,f424,f655,f423,f653,f422,f657,f420,f418,f413,f412,f411,f410,f583,f408,f661,f563,f564,f659,f677,f676,f683,f675,f685,f669,f668,f687,f667,f689,f799,f798,f795,f794,f793,f791,f790,f789,f835,f788,f833,f787,f786,f837,f778,f839,f776,f775,f774,f843,f773,f841,f772,f771,f845,f767,f847,f752,f849,f749,f851,f748,f741,f740,f739,f857,f738,f855,f737,f736,f859,f734,f733,f732,f863,f731,f861,f730,f729,f865,f728,f727,f726,f869,f725,f867,f724,f723,f853,f722,f721,f720,f719,f718,f717,f716,f715,f714,f879,f713,f877,f712,f711,f710,f875,f709,f871,f708,f707,f873,f701,f700,f698,f697,f693,f692,f831,f691,f690,f801,f802,f882,f881,f910,f915,f913,f1090,f1089,f1085,f1084,f1082,f1081,f1078,f1077,f1072,f1071,f1070,f1069,f1067,f1065,f1062,f1061,f1059,f1057,f1054,f1053,f1046,f1045,f1044,f1043,f1123,f1040,f1039,f1038,f1037,f1127,f1036,f1035,f1034,f1125,f1033,f1129,f1030,f1029,f1028,f1027,f1131,f1016,f1015,f1014,f1013,f1135,f1012,f1011,f1010,f1133,f1009,f1137,f1008,f1007,f1006,f1005,f1145,f1004,f1003,f1002,f1141,f1001,f1149,f1000,f1139,f999,f1147,f998,f1143,f997,f1151,f995,f994,f993,f992,f1155,f991,f990,f1153,f989,f988,f1157,f981,f1159,f980,f1161,f979,f978,f977,f976,f1169,f975,f974,f1165,f973,f972,f1173,f971,f1163,f970,f1167,f969,f1171,f968,f1175,f965,f964,f963,f962,f1177,f959,f958,f957,f956,f1179,f955,f954,f953,f952,f1181,f951,f950,f949,f948,f947,f946,f945,f944,f943,f942,f1191,f941,f940,f1189,f939,f938,f1187,f937,f1183,f936,f1185,f935,f934,f933,f932,f931,f930,f1195,f929,f1193,f928,f1197,f926,f924,f919,f918,f917,f916,f1113,f1092,f1093,f1199,f1294,f1293,f1292,f1297,f1306,f1298,f1304,f1299,f1302,f1314,f1315,f1445,f1451,f1453,f1469,f1468,f1467,f1471,f1481,f1472,f1479,f1473,f1477,f1509,f1515,f1517,f1526,f1532,f1534,f1554,f1556,f1566,f1568,f1573,f1575,f1580,f1584,f1589,f1593,f1618,f1617,f1621,f1616,f1615,f1623,f1632,f1631,f1635,f1850,f1856,f1858,f1895,f1899,f1905,f1909,f1917,f1921,f2150,f2149,f2146,f2145,f2144,f2143,f2140,f2139,f2136,f2135,f2134,f2133,f2123,f2122,f2121,f2120,f2175,f2115,f2114,f2113,f2112,f2180,f2110,f2109,f2108,f2178,f2107,f2183,f2102,f2101,f2100,f2099,f2187,f2096,f2095,f2190,f2077,f2076,f2075,f2074,f2197,f2072,f2071,f2070,f2195,f2069,f2200,f2068,f2067,f2066,f2065,f2208,f2063,f2062,f2061,f2204,f2060,f2213,f2058,f2202,f2057,f2211,f2056,f2206,f2055,f2216,f2051,f2050,f2049,f2048,f2222,f2045,f2044,f2220,f2043,f2042,f2226,f2039,f2038,f2231,f2035,f2229,f2034,f2234,f2019,f2238,f2018,f2240,f2015,f2014,f2013,f2012,f2250,f2009,f2008,f2246,f2007,f2006,f2256,f2003,f2244,f2002,f2248,f2001,f2254,f2000,f2260,f1995,f1994,f1993,f1992,f2264,f1987,f1986,f1985,f1984,f2268,f1983,f1982,f2270,f1981,f1980,f2272,f1979,f1978,f1977,f1976,f2274,f1975,f1974,f1973,f1972,f1971,f1970,f1969,f1968,f1967,f1966,f2284,f1965,f1964,f2282,f1963,f1962,f2280,f1961,f2276,f1960,f2278,f1958,f1957,f1955,f1954,f1952,f1951,f2289,f1949,f2287,f1948,f2292,f1946,f1944,f1939,f1938,f1937,f1936,f2171,f1935,f2298,f2151,f2152,f2296,f2308,f2316,f2304,f2319,f2332,f2331,f2424,f2325,f2324,f2428,f2494,f2493,f2590,f2478,f2477,f2601,f2681,f2689,f2680,f2691,f2679,f2693,f2678,f2695,f2837,f2838,f2835,f2840,f2841,f3143,f3149,f3152,f3171,f3173,f3179,f3181,f3193,f3195,f3260,f3264,f3272,f3276,f3729,f2297,f3769,f3770,f3771,f1317,f3773])).
fof(f3773,plain,(
  s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21('const.min.@_1'(s(fun(sK11,bool),sK13)),sK13,sK12,sK10,sK11)))) | $spl408),
  inference(cnf_transformation,[],[f3773_D])).
fof(f3773_D,plain,(
  s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21('const.min.@_1'(s(fun(sK11,bool),sK13)),sK13,sK12,sK10,sK11)))) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f1317,plain,(
  p(s(bool,'F_0')) | $spl236),
  inference(cnf_transformation,[],[f1317_D])).
fof(f1317_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f3771,plain,(
  p(s(bool,'F_0')) | s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21('const.min.@_1'(s(fun(sK11,bool),sK13)),sK13,sK12,sK10,sK11)))) | ($spl2 | $spl226)),
  inference(forward_demodulation,[],[f3766,f1199])).
fof(f3766,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13)))) | s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21('const.min.@_1'(s(fun(sK11,bool),sK13)),sK13,sK12,sK10,sK11)))) | $spl2),
  inference(superposition,[],[f800,f277])).
fof(f277,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X7,bool),X5))) = s(bool,happ(s(fun(X7,bool),X5),s(X7,'const.min.@_1'(s(fun(X7,bool),X5)))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.EXISTS_DEF')).
fof(f800,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0)))) | s(sK11,X0) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21(X0,sK13,sK12,sK10,sK11))))) ) | $spl2),
  inference(resolution,[],[f407,f321])).
fof(f321,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK21(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK21(X6,X3,X2,X1,X0)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK22(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK22(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X7] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X7,X6),X21),s(fun(X6,bool),X22)))) <=> ! [X19] : (~p(s(bool,happ(s(fun(X6,bool),X22),s(X6,X19)))) => ? [X20] : s(X6,X19) = s(X6,happ(s(fun(X7,X6),X21),s(X7,X20)))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.quantHeuristics.GUESS_REWRITES|split|5')).
fof(f3770,plain,(
  ( ! [X1] : (s(sK11,X1) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21(X1,sK13,sK12,sK10,sK11))))) ) | $spl406),
  inference(cnf_transformation,[],[f3770_D])).
fof(f3770_D,plain,(
  ( ! [X1] : (s(sK11,X1) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21(X1,sK13,sK12,sK10,sK11))))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f3769,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(sK11,X1) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21(X1,sK13,sK12,sK10,sK11))))) ) | ($spl2 | $spl226)),
  inference(forward_demodulation,[],[f3768,f1199])).
fof(f3768,plain,(
  ( ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13)))) | s(sK11,X1) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21(X1,sK13,sK12,sK10,sK11))))) ) | $spl2),
  inference(forward_demodulation,[],[f3731,f277])).
fof(f3731,plain,(
  ( ! [X1] : (s(sK11,X1) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK21(X1,sK13,sK12,sK10,sK11)))) | p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))))) ) | $spl2),
  inference(resolution,[],[f800,f278])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ! [X9,X11] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X11)))) => p(s(bool,happ(s(fun(X7,bool),X9),s(X7,'const.min.@_1'(s(fun(X7,bool),X9)))))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.SELECT_AX')).
fof(f2297,plain,(
  p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK22(sK14,sK12,sK10,sK11))))) | $spl366),
  inference(cnf_transformation,[],[f2297_D])).
fof(f2297_D,plain,(
  p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK22(sK14,sK12,sK10,sK11))))) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f3729,plain,(
  p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK22(sK14,sK12,sK10,sK11))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f409,f800])).
fof(f409,plain,(
  ( ! [X0] : (s(sK11,sK22(sK14,sK12,sK10,sK11)) != s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f323])).
fof(f323,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,sK22(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f3276,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl405),
  inference(cnf_transformation,[],[f3276_D])).
fof(f3276_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl405),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl405])])).
fof(f3272,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl123),
  inference(unit_resulting_resolution,[],[f861,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f205])).
fof(f205,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f3264,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl403),
  inference(cnf_transformation,[],[f3264_D])).
fof(f3264_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl403),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl403])])).
fof(f3260,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl117),
  inference(unit_resulting_resolution,[],[f855,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f3195,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl401),
  inference(cnf_transformation,[],[f3195_D])).
fof(f3195_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl401),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl401])])).
fof(f3193,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl53),
  inference(unit_resulting_resolution,[],[f631,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f3181,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl399),
  inference(cnf_transformation,[],[f3181_D])).
fof(f3181_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl399),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl399])])).
fof(f3179,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl49),
  inference(unit_resulting_resolution,[],[f627,f335])).
fof(f3173,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl397),
  inference(cnf_transformation,[],[f3173_D])).
fof(f3173_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl397),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl397])])).
fof(f3171,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl47),
  inference(unit_resulting_resolution,[],[f625,f335])).
fof(f3152,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl394),
  inference(cnf_transformation,[],[f3152_D])).
fof(f3152_D,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f3149,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl32),
  inference(resolution,[],[f611,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f217])).
fof(f217,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f3143,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl32),
  inference(unit_resulting_resolution,[],[f611,f360])).
fof(f2841,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK14)))),X1)) ) | $spl392),
  inference(cnf_transformation,[],[f2841_D])).
fof(f2841_D,plain,(
  ( ! [X0,X1] : (sP4(X0,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK14)))),X1)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f2840,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK14)))) | $spl390),
  inference(cnf_transformation,[],[f2840_D])).
fof(f2840_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK14)))) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f2835,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK14)))) | sP4(X0,happ(s(fun(sK11,bool),sK13),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK14)))),X1)) )),
  inference(superposition,[],[f899,f277])).
fof(f899,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,X24)))) | sP4(X25,happ(s(fun(sK11,bool),sK13),s(sK11,X24)),X26)) )),
  inference(resolution,[],[f234,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f208])).
fof(f208,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f234,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X5)))) | p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,X5))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X5] : (~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X5)))) | p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,X5))))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13,sK14])],[f107])).
fof(f107,plain,(
  ? [X0,X1,X2,X3,X4] : (! [X5] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X5))))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ? [X0,X1,X2,X3,X4] : (! [X5] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X5))))) & (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4))))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1,X2,X3,X4] : (! [X5] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X5))))) => (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (! [X5] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X5))))) => (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (! [X11] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X11)))) => p(s(bool,happ(s(fun(X7,bool),X10),s(X7,X11))))) => (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (! [X11] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X11)))) => p(s(bool,happ(s(fun(X7,bool),X10),s(X7,X11))))) => (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10))))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP')).
fof(f2838,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,happ(s(fun(sK11,bool),sK13),s(sK11,X1)),X2)) ) | $spl388),
  inference(cnf_transformation,[],[f2838_D])).
fof(f2838_D,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,happ(s(fun(sK11,bool),sK13),s(sK11,X1)),X2)) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f2837,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK14)))) | sP4(X0,happ(s(fun(sK11,bool),sK13),s(sK11,X1)),X2)) )),
  inference(forward_demodulation,[],[f2796,f277])).
fof(f2796,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,happ(s(fun(sK11,bool),sK13),s(sK11,X1)),X2) | p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK14))))))) )),
  inference(resolution,[],[f899,f278])).
fof(f2695,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','F_0') | $spl386),
  inference(cnf_transformation,[],[f2695_D])).
fof(f2695_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','F_0') <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f2678,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','F_0') | ($spl1 | $spl7)),
  inference(unit_resulting_resolution,[],[f585,f574])).
fof(f574,plain,(
  ( ! [X10,X11] : (sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X10,X11) | s(bool,X10) = s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f405,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f100,f145,f144,f143,f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f143,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> s(bool,X23) = s(bool,X24)) <=> ((p(s(bool,X25)) | p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X25)) | ~p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X25))) & (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.sat.sat_2|49|_prove_ltb')).
fof(f2693,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl384),
  inference(cnf_transformation,[],[f2693_D])).
fof(f2693_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f2679,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl5)),
  inference(unit_resulting_resolution,[],[f583,f574])).
fof(f2691,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','T_0') | $spl382),
  inference(cnf_transformation,[],[f2691_D])).
fof(f2691_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','T_0') <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f2680,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','T_0') | ($spl1 | $spl7)),
  inference(unit_resulting_resolution,[],[f585,f574])).
fof(f2689,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl380),
  inference(cnf_transformation,[],[f2689_D])).
fof(f2689_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f2681,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | ($spl1 | $spl5)),
  inference(unit_resulting_resolution,[],[f583,f574])).
fof(f2601,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl378),
  inference(cnf_transformation,[],[f2601_D])).
fof(f2601_D,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f2477,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl142),
  inference(unit_resulting_resolution,[],[f238,f881,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f215])).
fof(f238,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.TRUTH')).
fof(f2478,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl142),
  inference(unit_resulting_resolution,[],[f237,f881,f355])).
fof(f237,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP',aHL_TRUTH)).
fof(f2590,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl376),
  inference(cnf_transformation,[],[f2590_D])).
fof(f2590_D,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f2493,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl142),
  inference(unit_resulting_resolution,[],[f238,f881,f355])).
fof(f2494,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl142),
  inference(unit_resulting_resolution,[],[f237,f881,f355])).
fof(f2428,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl374),
  inference(cnf_transformation,[],[f2428_D])).
fof(f2428_D,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f2324,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl80),
  inference(unit_resulting_resolution,[],[f238,f659,f355])).
fof(f2325,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl80),
  inference(unit_resulting_resolution,[],[f237,f659,f355])).
fof(f2424,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl372),
  inference(cnf_transformation,[],[f2424_D])).
fof(f2424_D,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f2331,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl80),
  inference(unit_resulting_resolution,[],[f238,f659,f355])).
fof(f2332,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl80),
  inference(unit_resulting_resolution,[],[f237,f659,f355])).
fof(f2319,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl370),
  inference(cnf_transformation,[],[f2319_D])).
fof(f2319_D,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f2304,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f243,f583,f354])).
fof(f243,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP',aHL_FALSITY)).
fof(f2316,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl368),
  inference(cnf_transformation,[],[f2316_D])).
fof(f2316_D,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f2308,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl5),
  inference(unit_resulting_resolution,[],[f243,f583,f354])).
fof(f2296,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl364),
  inference(cnf_transformation,[],[f2296_D])).
fof(f2296_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f2152,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(resolution,[],[f661,f257])).
fof(f257,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.EQ_CLAUSES|split|2')).
fof(f2151,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(resolution,[],[f661,f255])).
fof(f255,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.EQ_CLAUSES|split|3')).
fof(f2298,plain,(
  ~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK22(sK14,sK12,sK10,sK11))))) | $spl367),
  inference(cnf_transformation,[],[f2298_D])).
fof(f2298_D,plain,(
  ~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK22(sK14,sK12,sK10,sK11))))) <=> ~$spl367),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl367])])).
fof(f1935,plain,(
  ~p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK22(sK14,sK12,sK10,sK11))))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f234])).
fof(f2171,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl275),
  inference(cnf_transformation,[],[f2171_D])).
fof(f2171_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f1936,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f250])).
fof(f250,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.EQ_CLAUSES|split|1')).
fof(f1937,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f252])).
fof(f252,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.EQ_CLAUSES|split|0')).
fof(f1938,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f255])).
fof(f1939,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f257])).
fof(f1944,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f281])).
fof(f281,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.IMP_ANTISYM_AX')).
fof(f1946,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f281])).
fof(f2292,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl363),
  inference(cnf_transformation,[],[f2292_D])).
fof(f2292_D,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl363),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl363])])).
fof(f1948,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f287])).
fof(f287,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2287,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl359),
  inference(cnf_transformation,[],[f2287_D])).
fof(f2287_D,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl359),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl359])])).
fof(f1949,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f287])).
fof(f2289,plain,(
  ~sP0('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl361),
  inference(cnf_transformation,[],[f2289_D])).
fof(f2289_D,plain,(
  ~sP0('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl361),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl361])])).
fof(f1951,plain,(
  ~sP0('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f287])).
fof(f1952,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f287])).
fof(f1954,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f291])).
fof(f291,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f125,f138])).
fof(f125,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6,X7] : ((~(p(s(bool,X7)) | p(s(bool,X6))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X7)) => p(s(bool,'F_0'))) => (~p(s(bool,X6)) => p(s(bool,'F_0')))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.sat.sat_9_prove_ltb')).
fof(f1955,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f291])).
fof(f1957,plain,(
  ~sP0('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f291])).
fof(f1958,plain,(
  ~sP0(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f291])).
fof(f2278,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl351),
  inference(cnf_transformation,[],[f2278_D])).
fof(f2278_D,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl351),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl351])])).
fof(f1960,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f302])).
fof(f302,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f180])).
fof(f180,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2276,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl349),
  inference(cnf_transformation,[],[f2276_D])).
fof(f2276_D,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f1961,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f302])).
fof(f2280,plain,(
  ~sP1('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl353),
  inference(cnf_transformation,[],[f2280_D])).
fof(f2280_D,plain,(
  ~sP1('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl353),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl353])])).
fof(f1962,plain,(
  ~sP1('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f302])).
fof(f1963,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f302])).
fof(f2282,plain,(
  sP1('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl354),
  inference(cnf_transformation,[],[f2282_D])).
fof(f2282_D,plain,(
  sP1('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f1964,plain,(
  sP1('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f304])).
fof(f304,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f1965,plain,(
  sP1('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f304])).
fof(f2284,plain,(
  sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl356),
  inference(cnf_transformation,[],[f2284_D])).
fof(f2284_D,plain,(
  sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f1966,plain,(
  sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f307])).
fof(f307,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f1967,plain,(
  sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f307])).
fof(f1968,plain,(
  sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f308])).
fof(f308,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f83,f140])).
fof(f83,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X24,X25] : ((p(s(bool,X25)) <=> ~p(s(bool,X24))) <=> ((p(s(bool,X25)) | p(s(bool,X24))) & (~p(s(bool,X24)) | ~p(s(bool,X25)))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.sat.sat_29_prove_ltb')).
fof(f1969,plain,(
  sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f308])).
fof(f1970,plain,(
  sP1('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f309])).
fof(f309,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f1971,plain,(
  sP1('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f309])).
fof(f1972,plain,(
  ~sP1('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f311])).
fof(f311,plain,(
  ( ! [X0,X1] : (~sP1(X1,X0) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f182])).
fof(f1973,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f311])).
fof(f1974,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f311])).
fof(f1975,plain,(
  ~sP1(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f311])).
fof(f2274,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl347),
  inference(cnf_transformation,[],[f2274_D])).
fof(f2274_D,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl347),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl347])])).
fof(f1976,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f238,f661,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f206])).
fof(f1977,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f238,f661,f338])).
fof(f1978,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f237,f661,f338])).
fof(f1979,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f237,f661,f338])).
fof(f2272,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl344),
  inference(cnf_transformation,[],[f2272_D])).
fof(f2272_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f1980,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl2 | $spl83)),
  inference(unit_resulting_resolution,[],[f743,f460,f454,f661,f340])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f454,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f344])).
fof(f460,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f211])).
fof(f211,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f143])).
fof(f743,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f142])).
fof(f1981,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl2 | $spl83)),
  inference(unit_resulting_resolution,[],[f743,f742,f454,f661,f340])).
fof(f742,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f2270,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl342),
  inference(cnf_transformation,[],[f2270_D])).
fof(f2270_D,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f1982,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl2 | $spl83)),
  inference(unit_resulting_resolution,[],[f744,f460,f454,f661,f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f744,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f1983,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl2 | $spl83)),
  inference(unit_resulting_resolution,[],[f744,f460,f735,f661,f341])).
fof(f735,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f209])).
fof(f2268,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl341),
  inference(cnf_transformation,[],[f2268_D])).
fof(f2268_D,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl341),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl341])])).
fof(f1984,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f238,f661,f342])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f1985,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f238,f661,f342])).
fof(f1986,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f237,f661,f342])).
fof(f1987,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f237,f661,f342])).
fof(f2264,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl339),
  inference(cnf_transformation,[],[f2264_D])).
fof(f2264_D,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f1992,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f238,f661,f346])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f212])).
fof(f1993,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f238,f661,f346])).
fof(f1994,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f237,f661,f346])).
fof(f1995,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f237,f661,f346])).
fof(f2260,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','F_0') | $spl337),
  inference(cnf_transformation,[],[f2260_D])).
fof(f2260_D,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','F_0') <=> ~$spl337),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl337])])).
fof(f2000,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f2254,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl333),
  inference(cnf_transformation,[],[f2254_D])).
fof(f2254_D,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl333),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl333])])).
fof(f2001,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f350])).
fof(f2248,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl329),
  inference(cnf_transformation,[],[f2248_D])).
fof(f2248_D,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f2002,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f661,f350])).
fof(f2244,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl325),
  inference(cnf_transformation,[],[f2244_D])).
fof(f2244_D,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl325),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl325])])).
fof(f2003,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f661,f350])).
fof(f2256,plain,(
  ~sP2('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl335),
  inference(cnf_transformation,[],[f2256_D])).
fof(f2256_D,plain,(
  ~sP2('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f2006,plain,(
  ~sP2('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f350])).
fof(f2007,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f350])).
fof(f2246,plain,(
  ~sP2('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl327),
  inference(cnf_transformation,[],[f2246_D])).
fof(f2246_D,plain,(
  ~sP2('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f2008,plain,(
  ~sP2('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f661,f350])).
fof(f2009,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f661,f350])).
fof(f2250,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl331),
  inference(cnf_transformation,[],[f2250_D])).
fof(f2250_D,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl331),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl331])])).
fof(f2012,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f350])).
fof(f2013,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f350])).
fof(f2014,plain,(
  ~sP2('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f661,f350])).
fof(f2015,plain,(
  ~sP2(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f661,f350])).
fof(f2240,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0','F_0') | $spl322),
  inference(cnf_transformation,[],[f2240_D])).
fof(f2240_D,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0','F_0') <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f2018,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0','F_0') | ($spl7 | $spl83)),
  inference(unit_resulting_resolution,[],[f585,f661,f354])).
fof(f2238,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','T_0') | $spl320),
  inference(cnf_transformation,[],[f2238_D])).
fof(f2238_D,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','T_0') <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f2019,plain,(
  sP5(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','T_0') | ($spl7 | $spl83)),
  inference(unit_resulting_resolution,[],[f585,f661,f354])).
fof(f2234,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl318),
  inference(cnf_transformation,[],[f2234_D])).
fof(f2234_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f2034,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | ($spl1 | $spl83)),
  inference(unit_resulting_resolution,[],[f243,f488,f661,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP7(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f488,plain,(
  ( ! [X0,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f2229,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl314),
  inference(cnf_transformation,[],[f2229_D])).
fof(f2229_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f2035,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | ($spl1 | $spl83)),
  inference(unit_resulting_resolution,[],[f661,f488,f661,f364])).
fof(f2231,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl316),
  inference(cnf_transformation,[],[f2231_D])).
fof(f2231_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f2038,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | ($spl1 | $spl83)),
  inference(unit_resulting_resolution,[],[f243,f488,f661,f364])).
fof(f2039,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | ($spl1 | $spl83)),
  inference(unit_resulting_resolution,[],[f661,f488,f661,f364])).
fof(f2226,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl313),
  inference(cnf_transformation,[],[f2226_D])).
fof(f2226_D,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl313),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl313])])).
fof(f2042,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f243,f661,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f2043,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f243,f661,f365])).
fof(f2220,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl309),
  inference(cnf_transformation,[],[f2220_D])).
fof(f2220_D,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl309),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl309])])).
fof(f2044,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f661,f365])).
fof(f2045,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f661,f365])).
fof(f2222,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl311),
  inference(cnf_transformation,[],[f2222_D])).
fof(f2222_D,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl311),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl311])])).
fof(f2048,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f243,f661,f365])).
fof(f2049,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f243,f661,f365])).
fof(f2050,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f661,f365])).
fof(f2051,plain,(
  ~sP6('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f661,f365])).
fof(f2216,plain,(
  sP7('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl306),
  inference(cnf_transformation,[],[f2216_D])).
fof(f2216_D,plain,(
  sP7('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f2055,plain,(
  sP7('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f102,f148,f147])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> (p(s(bool,X24)) | p(s(bool,X23)))) <=> ((p(s(bool,X25)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | p(s(bool,X23)) | ~p(s(bool,X25)))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.sat.sat_25_prove_ltb')).
fof(f2206,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl298),
  inference(cnf_transformation,[],[f2206_D])).
fof(f2206_D,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f2056,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f369])).
fof(f2211,plain,(
  sP7('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl302),
  inference(cnf_transformation,[],[f2211_D])).
fof(f2211_D,plain,(
  sP7('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f2057,plain,(
  sP7('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f661,f369])).
fof(f2202,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl294),
  inference(cnf_transformation,[],[f2202_D])).
fof(f2202_D,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f2058,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f661,f369])).
fof(f2213,plain,(
  sP7('F_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl304),
  inference(cnf_transformation,[],[f2213_D])).
fof(f2213_D,plain,(
  sP7('F_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f2060,plain,(
  sP7('F_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f369])).
fof(f2204,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl296),
  inference(cnf_transformation,[],[f2204_D])).
fof(f2204_D,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f2061,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f369])).
fof(f2062,plain,(
  sP7('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f661,f369])).
fof(f2063,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f661,f369])).
fof(f2208,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','F_0') | $spl300),
  inference(cnf_transformation,[],[f2208_D])).
fof(f2208_D,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','F_0') <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f2065,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f243,f661,f369])).
fof(f2066,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f661,f243,f661,f369])).
fof(f2067,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f661,f661,f369])).
fof(f2068,plain,(
  sP7(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f661,f661,f369])).
fof(f2200,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl293),
  inference(cnf_transformation,[],[f2200_D])).
fof(f2200_D,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') <=> ~$spl293),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl293])])).
fof(f2069,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f238,f661,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f2195,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl289),
  inference(cnf_transformation,[],[f2195_D])).
fof(f2195_D,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl289),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl289])])).
fof(f2070,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f238,f661,f372])).
fof(f2071,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f237,f661,f372])).
fof(f2072,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f237,f661,f372])).
fof(f2197,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl291),
  inference(cnf_transformation,[],[f2197_D])).
fof(f2197_D,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) <=> ~$spl291),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl291])])).
fof(f2074,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f238,f661,f372])).
fof(f2075,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f238,f661,f372])).
fof(f2076,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f237,f661,f372])).
fof(f2077,plain,(
  ~sP7('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f661,f237,f661,f372])).
fof(f2190,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl286),
  inference(cnf_transformation,[],[f2190_D])).
fof(f2190_D,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f2095,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | ($spl1 | $spl83)),
  inference(unit_resulting_resolution,[],[f238,f524,f661,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f225])).
fof(f225,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f524,plain,(
  ( ! [X0,X1] : (sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f2096,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | ($spl1 | $spl83)),
  inference(unit_resulting_resolution,[],[f237,f524,f661,f381])).
fof(f2187,plain,(
  ~sP8('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl285),
  inference(cnf_transformation,[],[f2187_D])).
fof(f2187_D,plain,(
  ~sP8('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f2099,plain,(
  ~sP8('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f238,f661,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f2100,plain,(
  ~sP8('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f238,f661,f382])).
fof(f2101,plain,(
  ~sP8('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f237,f661,f382])).
fof(f2102,plain,(
  ~sP8('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f237,f661,f382])).
fof(f2183,plain,(
  sP9('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl282),
  inference(cnf_transformation,[],[f2183_D])).
fof(f2183_D,plain,(
  sP9('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f2107,plain,(
  sP9('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f238,f661,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP9(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f135,f151,f150])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> (p(s(bool,X24)) => p(s(bool,X23)))) <=> ((p(s(bool,X25)) | p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X23))) & (~p(s(bool,X24)) | p(s(bool,X23)) | ~p(s(bool,X25)))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.sat.sat_27_prove_ltb')).
fof(f2178,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl278),
  inference(cnf_transformation,[],[f2178_D])).
fof(f2178_D,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f2108,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f661,f238,f661,f386])).
fof(f2109,plain,(
  sP9('F_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f243,f237,f661,f386])).
fof(f2110,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f661,f237,f661,f386])).
fof(f2180,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','T_0') | $spl280),
  inference(cnf_transformation,[],[f2180_D])).
fof(f2180_D,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','T_0') <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f2112,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f243,f661,f386])).
fof(f2113,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'F_0','T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f243,f661,f386])).
fof(f2114,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f661,f386])).
fof(f2115,plain,(
  sP9(happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f661,f386])).
fof(f2175,plain,(
  ~sP9('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl277),
  inference(cnf_transformation,[],[f2175_D])).
fof(f2175_D,plain,(
  ~sP9('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f2120,plain,(
  ~sP9('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f238,f661,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f2121,plain,(
  ~sP9('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f238,f661,f389])).
fof(f2122,plain,(
  ~sP9('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f238,f237,f661,f389])).
fof(f2123,plain,(
  ~sP9('T_0',happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))),'T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f237,f237,f661,f389])).
fof(f2133,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f238,f238,f661,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11,X17,X15,X18] : ((s(bool,X11) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X11)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP','thm.bool.IMP_CONG')).
fof(f2134,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f237,f238,f661,f398])).
fof(f2135,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f238,f237,f661,f398])).
fof(f2136,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f237,f237,f661,f398])).
fof(f2139,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f238,f661,f398])).
fof(f2140,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f238,f661,f398])).
fof(f2143,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f237,f661,f398])).
fof(f2144,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f237,f661,f398])).
fof(f2145,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f238,f661,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f2146,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f243,f237,f661,f403])).
fof(f2149,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f238,f661,f661,f403])).
fof(f2150,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11)))) | $spl83),
  inference(unit_resulting_resolution,[],[f237,f661,f661,f403])).
fof(f1921,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl273),
  inference(cnf_transformation,[],[f1921_D])).
fof(f1921_D,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f1917,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl125),
  inference(unit_resulting_resolution,[],[f863,f337])).
fof(f1909,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl271),
  inference(cnf_transformation,[],[f1909_D])).
fof(f1909_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl271])])).
fof(f1905,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl121),
  inference(unit_resulting_resolution,[],[f859,f336])).
fof(f1899,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl269),
  inference(cnf_transformation,[],[f1899_D])).
fof(f1899_D,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f1895,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl119),
  inference(unit_resulting_resolution,[],[f857,f336])).
fof(f1858,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl266),
  inference(cnf_transformation,[],[f1858_D])).
fof(f1858_D,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f1856,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl100),
  inference(resolution,[],[f839,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f226])).
fof(f1850,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl100),
  inference(unit_resulting_resolution,[],[f839,f377])).
fof(f1635,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl264),
  inference(cnf_transformation,[],[f1635_D])).
fof(f1635_D,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f1631,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f243,f735,f460,f744,f341])).
fof(f1632,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f243,f454,f460,f744,f341])).
fof(f1623,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl262),
  inference(cnf_transformation,[],[f1623_D])).
fof(f1623_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f1615,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f735,f742,f743,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f1616,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f735,f742,f743,f339])).
fof(f1621,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl260),
  inference(cnf_transformation,[],[f1621_D])).
fof(f1621_D,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f1617,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f243,f454,f742,f743,f340])).
fof(f1618,plain,(
  sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f243,f454,f460,f743,f340])).
fof(f1593,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl259),
  inference(cnf_transformation,[],[f1593_D])).
fof(f1593_D,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f1589,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl61),
  inference(unit_resulting_resolution,[],[f639,f337])).
fof(f1584,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','T_0') | $spl257),
  inference(cnf_transformation,[],[f1584_D])).
fof(f1584_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','T_0') <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f1580,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0','T_0') | $spl59),
  inference(unit_resulting_resolution,[],[f637,f336])).
fof(f1575,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl255),
  inference(cnf_transformation,[],[f1575_D])).
fof(f1575_D,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f1573,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl57),
  inference(unit_resulting_resolution,[],[f635,f335])).
fof(f1568,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl253),
  inference(cnf_transformation,[],[f1568_D])).
fof(f1568_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f1566,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl55),
  inference(unit_resulting_resolution,[],[f633,f335])).
fof(f1556,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') | $spl251),
  inference(cnf_transformation,[],[f1556_D])).
fof(f1556_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f1554,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') | $spl51),
  inference(unit_resulting_resolution,[],[f629,f335])).
fof(f1534,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl248),
  inference(cnf_transformation,[],[f1534_D])).
fof(f1534_D,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f1532,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl36),
  inference(resolution,[],[f615,f360])).
fof(f1526,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f615,f360])).
fof(f1517,plain,(
  sP6('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl246),
  inference(cnf_transformation,[],[f1517_D])).
fof(f1517_D,plain,(
  sP6('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f1515,plain,(
  sP6('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl34),
  inference(resolution,[],[f613,f360])).
fof(f1509,plain,(
  sP6('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl34),
  inference(unit_resulting_resolution,[],[f613,f360])).
fof(f1477,plain,(
  sP4('F_0','T_0','F_0') | $spl240),
  inference(cnf_transformation,[],[f1477_D])).
fof(f1477_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f1473,plain,(
  sP4('F_0','T_0','F_0') | $spl90),
  inference(resolution,[],[f689,f337])).
fof(f1479,plain,(
  sP3('F_0','T_0','F_0') | $spl242),
  inference(cnf_transformation,[],[f1479_D])).
fof(f1479_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f1472,plain,(
  sP3('F_0','T_0','F_0') | $spl90),
  inference(resolution,[],[f689,f336])).
fof(f1481,plain,(
  sP2('T_0','F_0','F_0') | $spl244),
  inference(cnf_transformation,[],[f1481_D])).
fof(f1481_D,plain,(
  sP2('T_0','F_0','F_0') <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f1471,plain,(
  sP2('T_0','F_0','F_0') | $spl90),
  inference(resolution,[],[f689,f335])).
fof(f1467,plain,(
  sP2('T_0','F_0','F_0') | $spl90),
  inference(unit_resulting_resolution,[],[f689,f335])).
fof(f1468,plain,(
  sP3('F_0','T_0','F_0') | $spl90),
  inference(unit_resulting_resolution,[],[f689,f336])).
fof(f1469,plain,(
  sP4('F_0','T_0','F_0') | $spl90),
  inference(unit_resulting_resolution,[],[f689,f337])).
fof(f1453,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl238),
  inference(cnf_transformation,[],[f1453_D])).
fof(f1453_D,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f1451,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl14),
  inference(resolution,[],[f593,f377])).
fof(f1445,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f593,f377])).
fof(f1315,plain,(
  ( ! [X6,X7] : (sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X6,X7) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)))) ) | $spl234),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  ( ! [X6,X7] : (sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X6,X7) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f1314,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X6,X7) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)))) ) | ($spl1 | $spl80)),
  inference(forward_demodulation,[],[f1311,f659])).
fof(f1311,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)))) | sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X6,X7) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f455,f341])).
fof(f455,plain,(
  ( ! [X0,X1] : (sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f1302,plain,(
  sP4('F_0','F_0','T_0') | $spl228),
  inference(cnf_transformation,[],[f1302_D])).
fof(f1302_D,plain,(
  sP4('F_0','F_0','T_0') <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f1299,plain,(
  sP4('F_0','F_0','T_0') | $spl86),
  inference(resolution,[],[f685,f337])).
fof(f1304,plain,(
  sP3('T_0','F_0','F_0') | $spl230),
  inference(cnf_transformation,[],[f1304_D])).
fof(f1304_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f1298,plain,(
  sP3('T_0','F_0','F_0') | $spl86),
  inference(resolution,[],[f685,f336])).
fof(f1306,plain,(
  sP2('F_0','T_0','F_0') | $spl232),
  inference(cnf_transformation,[],[f1306_D])).
fof(f1306_D,plain,(
  sP2('F_0','T_0','F_0') <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f1297,plain,(
  sP2('F_0','T_0','F_0') | $spl86),
  inference(resolution,[],[f685,f335])).
fof(f1292,plain,(
  sP2('F_0','T_0','F_0') | $spl86),
  inference(unit_resulting_resolution,[],[f685,f335])).
fof(f1293,plain,(
  sP3('T_0','F_0','F_0') | $spl86),
  inference(unit_resulting_resolution,[],[f685,f336])).
fof(f1294,plain,(
  sP4('F_0','F_0','T_0') | $spl86),
  inference(unit_resulting_resolution,[],[f685,f337])).
fof(f1199,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl226),
  inference(cnf_transformation,[],[f1199_D])).
fof(f1199_D,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f1093,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(resolution,[],[f915,f257])).
fof(f1092,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(resolution,[],[f915,f255])).
fof(f1113,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl149),
  inference(cnf_transformation,[],[f1113_D])).
fof(f1113_D,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f916,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f250])).
fof(f917,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f252])).
fof(f918,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f255])).
fof(f919,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f257])).
fof(f924,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f281])).
fof(f926,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f281])).
fof(f1197,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl225),
  inference(cnf_transformation,[],[f1197_D])).
fof(f1197_D,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f928,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f287])).
fof(f1193,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl221),
  inference(cnf_transformation,[],[f1193_D])).
fof(f1193_D,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f929,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f287])).
fof(f1195,plain,(
  ~sP0('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl223),
  inference(cnf_transformation,[],[f1195_D])).
fof(f1195_D,plain,(
  ~sP0('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f930,plain,(
  ~sP0('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f287])).
fof(f931,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f287])).
fof(f932,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f291])).
fof(f933,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f291])).
fof(f934,plain,(
  ~sP0('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f291])).
fof(f935,plain,(
  ~sP0('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f291])).
fof(f1185,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl213),
  inference(cnf_transformation,[],[f1185_D])).
fof(f1185_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f936,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f302])).
fof(f1183,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl211),
  inference(cnf_transformation,[],[f1183_D])).
fof(f1183_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f937,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f302])).
fof(f1187,plain,(
  ~sP1('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl215),
  inference(cnf_transformation,[],[f1187_D])).
fof(f1187_D,plain,(
  ~sP1('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f938,plain,(
  ~sP1('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f302])).
fof(f939,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f302])).
fof(f1189,plain,(
  sP1('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl216),
  inference(cnf_transformation,[],[f1189_D])).
fof(f1189_D,plain,(
  sP1('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f940,plain,(
  sP1('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f304])).
fof(f941,plain,(
  sP1('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f304])).
fof(f1191,plain,(
  sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl218),
  inference(cnf_transformation,[],[f1191_D])).
fof(f1191_D,plain,(
  sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f942,plain,(
  sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f307])).
fof(f943,plain,(
  sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f307])).
fof(f944,plain,(
  sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f308])).
fof(f945,plain,(
  sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f308])).
fof(f946,plain,(
  sP1('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f309])).
fof(f947,plain,(
  sP1('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f309])).
fof(f948,plain,(
  ~sP1('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f311])).
fof(f949,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f311])).
fof(f950,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f311])).
fof(f951,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f311])).
fof(f1181,plain,(
  ~sP5('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl209),
  inference(cnf_transformation,[],[f1181_D])).
fof(f1181_D,plain,(
  ~sP5('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl209),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl209])])).
fof(f952,plain,(
  ~sP5('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f238,f915,f338])).
fof(f953,plain,(
  ~sP5('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f238,f915,f338])).
fof(f954,plain,(
  ~sP5('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f237,f915,f338])).
fof(f955,plain,(
  ~sP5('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f237,f915,f338])).
fof(f1179,plain,(
  ~sP4('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl207),
  inference(cnf_transformation,[],[f1179_D])).
fof(f1179_D,plain,(
  ~sP4('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl207),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl207])])).
fof(f956,plain,(
  ~sP4('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f238,f915,f342])).
fof(f957,plain,(
  ~sP4('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f238,f915,f342])).
fof(f958,plain,(
  ~sP4('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f237,f915,f342])).
fof(f959,plain,(
  ~sP4('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f237,f915,f342])).
fof(f1177,plain,(
  ~sP3('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl205),
  inference(cnf_transformation,[],[f1177_D])).
fof(f1177_D,plain,(
  ~sP3('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl205),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl205])])).
fof(f962,plain,(
  ~sP3('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f238,f915,f346])).
fof(f963,plain,(
  ~sP3('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f238,f915,f346])).
fof(f964,plain,(
  ~sP3('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f237,f915,f346])).
fof(f965,plain,(
  ~sP3('T_0','T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f237,f915,f346])).
fof(f1175,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','F_0') | $spl203),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','F_0') <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f968,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f350])).
fof(f1171,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl199),
  inference(cnf_transformation,[],[f1171_D])).
fof(f1171_D,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f969,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f350])).
fof(f1167,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl195),
  inference(cnf_transformation,[],[f1167_D])).
fof(f1167_D,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f970,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f915,f350])).
fof(f1163,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl191),
  inference(cnf_transformation,[],[f1163_D])).
fof(f1163_D,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl191),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl191])])).
fof(f971,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f915,f350])).
fof(f1173,plain,(
  ~sP2('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl201),
  inference(cnf_transformation,[],[f1173_D])).
fof(f1173_D,plain,(
  ~sP2('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f972,plain,(
  ~sP2('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f350])).
fof(f973,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f350])).
fof(f1165,plain,(
  ~sP2('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl193),
  inference(cnf_transformation,[],[f1165_D])).
fof(f1165_D,plain,(
  ~sP2('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl193),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl193])])).
fof(f974,plain,(
  ~sP2('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f915,f350])).
fof(f975,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f915,f350])).
fof(f1169,plain,(
  ~sP2('F_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl197),
  inference(cnf_transformation,[],[f1169_D])).
fof(f1169_D,plain,(
  ~sP2('F_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f976,plain,(
  ~sP2('F_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f350])).
fof(f977,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f350])).
fof(f978,plain,(
  ~sP2('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f915,f350])).
fof(f979,plain,(
  ~sP2('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f915,f350])).
fof(f1161,plain,(
  sP5('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl188),
  inference(cnf_transformation,[],[f1161_D])).
fof(f1161_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0','F_0') <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f980,plain,(
  sP5('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0','F_0') | ($spl7 | $spl147)),
  inference(unit_resulting_resolution,[],[f585,f915,f354])).
fof(f1159,plain,(
  sP5('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl186),
  inference(cnf_transformation,[],[f1159_D])).
fof(f1159_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','T_0') <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f981,plain,(
  sP5('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','T_0') | ($spl7 | $spl147)),
  inference(unit_resulting_resolution,[],[f585,f915,f354])).
fof(f1157,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl185),
  inference(cnf_transformation,[],[f1157_D])).
fof(f1157_D,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f988,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f243,f915,f365])).
fof(f989,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f243,f915,f365])).
fof(f1153,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl181),
  inference(cnf_transformation,[],[f1153_D])).
fof(f1153_D,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f990,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f915,f365])).
fof(f991,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f915,f365])).
fof(f1155,plain,(
  ~sP6('T_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl183),
  inference(cnf_transformation,[],[f1155_D])).
fof(f1155_D,plain,(
  ~sP6('T_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f992,plain,(
  ~sP6('T_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f243,f915,f365])).
fof(f993,plain,(
  ~sP6('T_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f243,f915,f365])).
fof(f994,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f915,f365])).
fof(f995,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f915,f365])).
fof(f1151,plain,(
  sP7('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl178),
  inference(cnf_transformation,[],[f1151_D])).
fof(f1151_D,plain,(
  sP7('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f997,plain,(
  sP7('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f369])).
fof(f1143,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl170),
  inference(cnf_transformation,[],[f1143_D])).
fof(f1143_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f998,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f369])).
fof(f1147,plain,(
  sP7('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl174),
  inference(cnf_transformation,[],[f1147_D])).
fof(f1147_D,plain,(
  sP7('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f999,plain,(
  sP7('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f915,f369])).
fof(f1139,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl166),
  inference(cnf_transformation,[],[f1139_D])).
fof(f1139_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f1000,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f915,f369])).
fof(f1149,plain,(
  sP7('F_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl176),
  inference(cnf_transformation,[],[f1149_D])).
fof(f1149_D,plain,(
  sP7('F_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f1001,plain,(
  sP7('F_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f369])).
fof(f1141,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl168),
  inference(cnf_transformation,[],[f1141_D])).
fof(f1141_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1002,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f369])).
fof(f1003,plain,(
  sP7('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f915,f369])).
fof(f1004,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f915,f369])).
fof(f1145,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','F_0') | $spl172),
  inference(cnf_transformation,[],[f1145_D])).
fof(f1145_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','F_0') <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f1005,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f243,f915,f369])).
fof(f1006,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f915,f243,f915,f369])).
fof(f1007,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f915,f915,f369])).
fof(f1008,plain,(
  sP7('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f915,f915,f369])).
fof(f1137,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl165),
  inference(cnf_transformation,[],[f1137_D])).
fof(f1137_D,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f1009,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f238,f915,f372])).
fof(f1133,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl161),
  inference(cnf_transformation,[],[f1133_D])).
fof(f1133_D,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f1010,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f238,f915,f372])).
fof(f1011,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f237,f915,f372])).
fof(f1012,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f237,f915,f372])).
fof(f1135,plain,(
  ~sP7('T_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl163),
  inference(cnf_transformation,[],[f1135_D])).
fof(f1135_D,plain,(
  ~sP7('T_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f1013,plain,(
  ~sP7('T_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f238,f915,f372])).
fof(f1014,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f238,f915,f372])).
fof(f1015,plain,(
  ~sP7('T_0','F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f237,f915,f372])).
fof(f1016,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f237,f915,f372])).
fof(f1131,plain,(
  ~sP8('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl159),
  inference(cnf_transformation,[],[f1131_D])).
fof(f1131_D,plain,(
  ~sP8('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f1027,plain,(
  ~sP8('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f238,f915,f382])).
fof(f1028,plain,(
  ~sP8('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f238,f915,f382])).
fof(f1029,plain,(
  ~sP8('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f237,f915,f382])).
fof(f1030,plain,(
  ~sP8('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f237,f915,f382])).
fof(f1129,plain,(
  sP9('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl156),
  inference(cnf_transformation,[],[f1129_D])).
fof(f1129_D,plain,(
  sP9('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f1033,plain,(
  sP9('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f238,f915,f386])).
fof(f1125,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl152),
  inference(cnf_transformation,[],[f1125_D])).
fof(f1125_D,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f1034,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f915,f238,f915,f386])).
fof(f1035,plain,(
  sP9('F_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f243,f237,f915,f386])).
fof(f1036,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f915,f237,f915,f386])).
fof(f1127,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl154),
  inference(cnf_transformation,[],[f1127_D])).
fof(f1127_D,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','T_0') <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f1037,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f243,f915,f386])).
fof(f1038,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f243,f915,f386])).
fof(f1039,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f915,f386])).
fof(f1040,plain,(
  sP9('|dollar|exists_1'(s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f915,f386])).
fof(f1123,plain,(
  ~sP9('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl151),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  ~sP9('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f1043,plain,(
  ~sP9('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f238,f915,f389])).
fof(f1044,plain,(
  ~sP9('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f238,f915,f389])).
fof(f1045,plain,(
  ~sP9('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f238,f237,f915,f389])).
fof(f1046,plain,(
  ~sP9('T_0','|dollar|exists_1'(s(fun(sK11,bool),sK13)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f237,f237,f915,f389])).
fof(f1053,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f277,f915,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f1054,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f277,f915,f393])).
fof(f1057,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f238,f277,f915,f393])).
fof(f1059,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f237,f277,f915,f393])).
fof(f1061,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f277,f915,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f1062,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f277,f915,f397])).
fof(f1065,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f238,f277,f915,f397])).
fof(f1067,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f915,f237,f277,f915,f397])).
fof(f1069,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f238,f238,f915,f398])).
fof(f1070,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f237,f238,f915,f398])).
fof(f1071,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f238,f237,f915,f398])).
fof(f1072,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f237,f237,f915,f398])).
fof(f1077,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f238,f915,f398])).
fof(f1078,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f238,f915,f398])).
fof(f1081,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f237,f915,f398])).
fof(f1082,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f237,f915,f398])).
fof(f1084,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f238,f915,f403])).
fof(f1085,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f243,f237,f915,f403])).
fof(f1089,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f238,f915,f915,f403])).
fof(f1090,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f237,f915,f915,f403])).
fof(f913,plain,(
  p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) | $spl144),
  inference(cnf_transformation,[],[f913_D])).
fof(f913_D,plain,(
  p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13)))))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f915,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13)))) | $spl147),
  inference(cnf_transformation,[],[f915_D])).
fof(f915_D,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13)))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f910,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK11,bool),sK13)))) | p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,'const.min.@_1'(s(fun(sK11,bool),sK13))))))),
  inference(superposition,[],[f234,f277])).
fof(f881,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl142),
  inference(cnf_transformation,[],[f881_D])).
fof(f881_D,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f882,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(global_subsumption,[],[f236,f235,f234,f407,f799,f798,f795,f794,f793,f792,f791,f790,f789,f788,f787,f786,f785,f784,f781,f779,f778,f777,f776,f775,f774,f773,f772,f771,f769,f768,f767,f766,f765,f762,f760,f759,f756,f754,f753,f752,f747,f746,f745,f744,f743,f742,f741,f740,f739,f738,f737,f736,f735,f734,f733,f732,f731,f730,f729,f728,f727,f726,f725,f724,f723,f722,f721,f720,f719,f718,f717,f716,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f701,f700,f698,f697,f693,f692,f691,f690,f800,f801,f802,f805,f810,f811,f812,f815,f823,f824,f826,f827,f828,f829])).
fof(f829,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(superposition,[],[f407,f244])).
fof(f244,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_STRENGTHEN_FORALL_GAP',aHL_BOOL_CASES)).
fof(f828,plain,(
  ( ! [X37,X36] : (sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X36,X37) | ~p(s(bool,X36))) ) | $spl2),
  inference(resolution,[],[f407,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f827,plain,(
  ( ! [X35,X34] : (sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X35,X34) | p(s(bool,X34))) ) | $spl2),
  inference(resolution,[],[f407,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f826,plain,(
  ( ! [X33,X32] : (sP9(X33,X32,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X32)) | p(s(bool,X33))) ) | $spl2),
  inference(resolution,[],[f407,f386])).
fof(f824,plain,(
  ( ! [X28,X29] : (sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X28,X29) | ~p(s(bool,X28))) ) | $spl2),
  inference(resolution,[],[f407,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f223])).
fof(f823,plain,(
  ( ! [X26,X27] : (sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X27,X26) | ~p(s(bool,X26))) ) | $spl2),
  inference(resolution,[],[f407,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f223])).
fof(f815,plain,(
  ( ! [X11] : (~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X11) | ~p(s(bool,X11))) ) | $spl2),
  inference(resolution,[],[f407,f310])).
fof(f310,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f812,plain,(
  ( ! [X8] : (sP1(X8,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | p(s(bool,X8))) ) | $spl2),
  inference(resolution,[],[f407,f307])).
fof(f811,plain,(
  ( ! [X7] : (sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X7) | p(s(bool,X7))) ) | $spl2),
  inference(resolution,[],[f407,f304])).
fof(f810,plain,(
  ( ! [X6] : (~sP1(X6,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | ~p(s(bool,X6))) ) | $spl2),
  inference(resolution,[],[f407,f303])).
fof(f303,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f805,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(resolution,[],[f407,f284])).
fof(f284,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f119])).
fof(f705,plain,(
  ( ! [X0] : (sP0(X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f293])).
fof(f293,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f176])).
fof(f706,plain,(
  ( ! [X0] : (sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f294])).
fof(f294,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f176])).
fof(f745,plain,(
  ( ! [X0,X1] : (sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f746,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f277,f407,f355])).
fof(f747,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f277,f407,f355])).
fof(f753,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f221])).
fof(f754,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f221])).
fof(f756,plain,(
  ( ! [X0] : (sP7('T_0',X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f370])).
fof(f759,plain,(
  ( ! [X0] : (sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f370])).
fof(f760,plain,(
  ( ! [X0] : (sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f407,f370])).
fof(f762,plain,(
  ( ! [X0] : (sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f371])).
fof(f765,plain,(
  ( ! [X0] : (sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f371])).
fof(f766,plain,(
  ( ! [X0] : (sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f407,f371])).
fof(f768,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f769,plain,(
  ( ! [X0] : (~sP7('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f777,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f229])).
fof(f779,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f387])).
fof(f781,plain,(
  ( ! [X0] : (sP9('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f388])).
fof(f784,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f388])).
fof(f785,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f407,f388])).
fof(f792,plain,(
  ( ! [X0] : (~sP9('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f802,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(resolution,[],[f407,f253])).
fof(f253,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f157])).
fof(f801,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(resolution,[],[f407,f251])).
fof(f251,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f156])).
fof(f690,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f251])).
fof(f691,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f253])).
fof(f831,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl93),
  inference(cnf_transformation,[],[f831_D])).
fof(f831_D,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f692,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f254])).
fof(f254,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f158])).
fof(f693,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f256])).
fof(f256,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f159])).
fof(f697,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f238,f407,f284])).
fof(f698,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f284])).
fof(f700,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f238,f407,f284])).
fof(f701,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f284])).
fof(f873,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl135),
  inference(cnf_transformation,[],[f873_D])).
fof(f873_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f707,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f407,f303])).
fof(f708,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f303])).
fof(f871,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl133),
  inference(cnf_transformation,[],[f871_D])).
fof(f871_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f709,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f407,f303])).
fof(f875,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl137),
  inference(cnf_transformation,[],[f875_D])).
fof(f875_D,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f710,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f238,f407,f303])).
fof(f711,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f303])).
fof(f712,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f407,f303])).
fof(f877,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl138),
  inference(cnf_transformation,[],[f877_D])).
fof(f877_D,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f713,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f304])).
fof(f879,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl140),
  inference(cnf_transformation,[],[f879_D])).
fof(f879_D,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f714,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f307])).
fof(f715,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f308])).
fof(f716,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f309])).
fof(f717,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f238,f407,f310])).
fof(f718,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f310])).
fof(f719,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f407,f310])).
fof(f720,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f407,f310])).
fof(f721,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f407,f310])).
fof(f722,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f407,f310])).
fof(f853,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl115),
  inference(cnf_transformation,[],[f853_D])).
fof(f853_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f723,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f243,f407,f338])).
fof(f724,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f243,f407,f338])).
fof(f867,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl129),
  inference(cnf_transformation,[],[f867_D])).
fof(f867_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f725,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f407,f243,f407,f338])).
fof(f869,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl131),
  inference(cnf_transformation,[],[f869_D])).
fof(f869_D,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f726,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f243,f407,f338])).
fof(f727,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f243,f407,f338])).
fof(f728,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f407,f243,f407,f338])).
fof(f865,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl127),
  inference(cnf_transformation,[],[f865_D])).
fof(f865_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f729,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f243,f407,f342])).
fof(f730,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f243,f407,f342])).
fof(f861,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl123),
  inference(cnf_transformation,[],[f861_D])).
fof(f861_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f731,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f407,f243,f407,f342])).
fof(f863,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl125),
  inference(cnf_transformation,[],[f863_D])).
fof(f863_D,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f732,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f243,f407,f342])).
fof(f733,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f243,f407,f342])).
fof(f734,plain,(
  ~sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f407,f243,f407,f342])).
fof(f859,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl121),
  inference(cnf_transformation,[],[f859_D])).
fof(f859_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f736,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f243,f407,f346])).
fof(f737,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f243,f407,f346])).
fof(f855,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl117),
  inference(cnf_transformation,[],[f855_D])).
fof(f855_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f738,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f407,f243,f407,f346])).
fof(f857,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl119),
  inference(cnf_transformation,[],[f857_D])).
fof(f857_D,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f739,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f238,f243,f407,f346])).
fof(f740,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f237,f243,f407,f346])).
fof(f741,plain,(
  ~sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f407,f243,f407,f346])).
fof(f748,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'T_0','F_0') | ($spl2 | $spl7)),
  inference(unit_resulting_resolution,[],[f585,f407,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f215])).
fof(f851,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl113),
  inference(cnf_transformation,[],[f851_D])).
fof(f851_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f749,plain,(
  ~sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') | ($spl2 | $spl7)),
  inference(unit_resulting_resolution,[],[f585,f407,f356])).
fof(f849,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','F_0') | $spl111),
  inference(cnf_transformation,[],[f849_D])).
fof(f849_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','F_0') <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f752,plain,(
  ~sP6('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f243,f243,f407,f365])).
fof(f847,plain,(
  ~sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','F_0') | $spl109),
  inference(cnf_transformation,[],[f847_D])).
fof(f847_D,plain,(
  ~sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','F_0') <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f767,plain,(
  ~sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f243,f243,f407,f372])).
fof(f845,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl107),
  inference(cnf_transformation,[],[f845_D])).
fof(f845_D,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f771,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f243,f238,f407,f382])).
fof(f772,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f243,f237,f407,f382])).
fof(f841,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl103),
  inference(cnf_transformation,[],[f841_D])).
fof(f841_D,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f773,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f407,f382])).
fof(f843,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl105),
  inference(cnf_transformation,[],[f843_D])).
fof(f843_D,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f774,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f238,f243,f407,f382])).
fof(f775,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f237,f243,f407,f382])).
fof(f776,plain,(
  ~sP8('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f407,f243,f407,f382])).
fof(f839,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl100),
  inference(cnf_transformation,[],[f839_D])).
fof(f839_D,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f778,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f243,f407,f386])).
fof(f837,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl99),
  inference(cnf_transformation,[],[f837_D])).
fof(f837_D,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f786,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f238,f407,f389])).
fof(f787,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f237,f407,f389])).
fof(f833,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl95),
  inference(cnf_transformation,[],[f833_D])).
fof(f833_D,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f788,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f407,f389])).
fof(f835,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl97),
  inference(cnf_transformation,[],[f835_D])).
fof(f835_D,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f789,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f243,f238,f407,f389])).
fof(f790,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f243,f237,f407,f389])).
fof(f791,plain,(
  ~sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f407,f407,f389])).
fof(f793,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f243,f238,f407,f398])).
fof(f794,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f243,f237,f407,f398])).
fof(f795,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f243,f407,f407,f398])).
fof(f798,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f243,f407,f407,f398])).
fof(f799,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f243,f243,f407,f403])).
fof(f689,plain,(
  sP5('F_0','F_0','T_0') | $spl90),
  inference(cnf_transformation,[],[f689_D])).
fof(f689_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f667,plain,(
  sP5('F_0','F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f243,f585,f354])).
fof(f687,plain,(
  ~sP5('T_0','F_0','T_0') | $spl89),
  inference(cnf_transformation,[],[f687_D])).
fof(f687_D,plain,(
  ~sP5('T_0','F_0','T_0') <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f668,plain,(
  ~sP5('T_0','F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f238,f585,f356])).
fof(f669,plain,(
  ~sP5('T_0','F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f237,f585,f356])).
fof(f685,plain,(
  sP5('F_0','T_0','F_0') | $spl86),
  inference(cnf_transformation,[],[f685_D])).
fof(f685_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f675,plain,(
  sP5('F_0','T_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f243,f585,f354])).
fof(f683,plain,(
  ~sP5('T_0','T_0','F_0') | $spl85),
  inference(cnf_transformation,[],[f683_D])).
fof(f683_D,plain,(
  ~sP5('T_0','T_0','F_0') <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f676,plain,(
  ~sP5('T_0','T_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f238,f585,f356])).
fof(f677,plain,(
  ~sP5('T_0','T_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f237,f585,f356])).
fof(f659,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl80),
  inference(cnf_transformation,[],[f659_D])).
fof(f659_D,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f564,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(resolution,[],[f405,f257])).
fof(f563,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(resolution,[],[f405,f255])).
fof(f661,plain,(
  ~p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))))) | $spl83),
  inference(cnf_transformation,[],[f661_D])).
fof(f661_D,plain,(
  ~p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f408,plain,(
  ~p(s(bool,happ(s(fun(sK11,bool),sK14),s(sK11,sK22(sK14,sK12,sK10,sK11))))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f322])).
fof(f322,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK22(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f583,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl5),
  inference(cnf_transformation,[],[f583_D])).
fof(f583_D,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f410,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f250])).
fof(f411,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f252])).
fof(f412,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f255])).
fof(f413,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f257])).
fof(f418,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f281])).
fof(f420,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f281])).
fof(f657,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl79),
  inference(cnf_transformation,[],[f657_D])).
fof(f657_D,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f422,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f287])).
fof(f653,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl75),
  inference(cnf_transformation,[],[f653_D])).
fof(f653_D,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f423,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f287])).
fof(f655,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl77),
  inference(cnf_transformation,[],[f655_D])).
fof(f655_D,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f424,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f287])).
fof(f425,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f287])).
fof(f426,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f291])).
fof(f427,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f291])).
fof(f428,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f291])).
fof(f429,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f291])).
fof(f645,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl67),
  inference(cnf_transformation,[],[f645_D])).
fof(f645_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f430,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f302])).
fof(f643,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl65),
  inference(cnf_transformation,[],[f643_D])).
fof(f643_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f431,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f302])).
fof(f647,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl69),
  inference(cnf_transformation,[],[f647_D])).
fof(f647_D,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f432,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f302])).
fof(f433,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f302])).
fof(f649,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl70),
  inference(cnf_transformation,[],[f649_D])).
fof(f649_D,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f434,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f304])).
fof(f435,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f304])).
fof(f651,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl72),
  inference(cnf_transformation,[],[f651_D])).
fof(f651_D,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f436,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f307])).
fof(f437,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f307])).
fof(f438,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f308])).
fof(f439,plain,(
  sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f308])).
fof(f440,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f309])).
fof(f441,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f309])).
fof(f442,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f311])).
fof(f443,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f311])).
fof(f444,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f311])).
fof(f445,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f311])).
fof(f641,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl63),
  inference(cnf_transformation,[],[f641_D])).
fof(f641_D,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f446,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f238,f405,f338])).
fof(f447,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f238,f405,f338])).
fof(f448,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f237,f405,f338])).
fof(f449,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f237,f405,f338])).
fof(f639,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl61),
  inference(cnf_transformation,[],[f639_D])).
fof(f639_D,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f450,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f238,f405,f342])).
fof(f451,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f238,f405,f342])).
fof(f452,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f237,f405,f342])).
fof(f453,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f237,f405,f342])).
fof(f637,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl59),
  inference(cnf_transformation,[],[f637_D])).
fof(f637_D,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f456,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f238,f405,f346])).
fof(f457,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f238,f405,f346])).
fof(f458,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f237,f405,f346])).
fof(f459,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f237,f405,f346])).
fof(f635,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') | $spl57),
  inference(cnf_transformation,[],[f635_D])).
fof(f635_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f462,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f350])).
fof(f631,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl53),
  inference(cnf_transformation,[],[f631_D])).
fof(f631_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f463,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f350])).
fof(f627,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl49),
  inference(cnf_transformation,[],[f627_D])).
fof(f627_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f464,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f405,f350])).
fof(f623,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl45),
  inference(cnf_transformation,[],[f623_D])).
fof(f623_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f465,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f405,f350])).
fof(f633,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl55),
  inference(cnf_transformation,[],[f633_D])).
fof(f633_D,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f466,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f350])).
fof(f467,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f350])).
fof(f625,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl47),
  inference(cnf_transformation,[],[f625_D])).
fof(f625_D,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f468,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f405,f350])).
fof(f469,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f405,f350])).
fof(f629,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl51),
  inference(cnf_transformation,[],[f629_D])).
fof(f629_D,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f470,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f350])).
fof(f471,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f350])).
fof(f472,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f405,f350])).
fof(f473,plain,(
  ~sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f405,f350])).
fof(f621,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl43),
  inference(cnf_transformation,[],[f621_D])).
fof(f621_D,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f480,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f243,f405,f365])).
fof(f481,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f243,f405,f365])).
fof(f617,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl39),
  inference(cnf_transformation,[],[f617_D])).
fof(f617_D,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f482,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f405,f365])).
fof(f483,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f405,f365])).
fof(f619,plain,(
  ~sP6('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl41),
  inference(cnf_transformation,[],[f619_D])).
fof(f619_D,plain,(
  ~sP6('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f484,plain,(
  ~sP6('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f243,f405,f365])).
fof(f485,plain,(
  ~sP6('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f243,f405,f365])).
fof(f486,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f405,f365])).
fof(f487,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f405,f365])).
fof(f615,plain,(
  sP7('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl36),
  inference(cnf_transformation,[],[f615_D])).
fof(f615_D,plain,(
  sP7('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f489,plain,(
  sP7('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f369])).
fof(f607,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl28),
  inference(cnf_transformation,[],[f607_D])).
fof(f607_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f490,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f369])).
fof(f611,plain,(
  sP7('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl32),
  inference(cnf_transformation,[],[f611_D])).
fof(f611_D,plain,(
  sP7('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f491,plain,(
  sP7('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f405,f369])).
fof(f603,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl24),
  inference(cnf_transformation,[],[f603_D])).
fof(f603_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f492,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f405,f369])).
fof(f613,plain,(
  sP7('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl34),
  inference(cnf_transformation,[],[f613_D])).
fof(f613_D,plain,(
  sP7('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f493,plain,(
  sP7('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f369])).
fof(f605,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl26),
  inference(cnf_transformation,[],[f605_D])).
fof(f605_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f494,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f369])).
fof(f495,plain,(
  sP7('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f405,f369])).
fof(f496,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f405,f369])).
fof(f609,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') | $spl30),
  inference(cnf_transformation,[],[f609_D])).
fof(f609_D,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f497,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f243,f405,f369])).
fof(f498,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f405,f243,f405,f369])).
fof(f499,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f405,f405,f369])).
fof(f500,plain,(
  sP7('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f405,f405,f369])).
fof(f601,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl23),
  inference(cnf_transformation,[],[f601_D])).
fof(f601_D,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f501,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f238,f405,f372])).
fof(f597,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl19),
  inference(cnf_transformation,[],[f597_D])).
fof(f597_D,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f502,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f238,f405,f372])).
fof(f503,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f237,f405,f372])).
fof(f504,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f237,f405,f372])).
fof(f599,plain,(
  ~sP7('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl21),
  inference(cnf_transformation,[],[f599_D])).
fof(f599_D,plain,(
  ~sP7('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f505,plain,(
  ~sP7('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f238,f405,f372])).
fof(f506,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f238,f405,f372])).
fof(f507,plain,(
  ~sP7('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f237,f405,f372])).
fof(f508,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f405,f237,f405,f372])).
fof(f595,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl17),
  inference(cnf_transformation,[],[f595_D])).
fof(f595_D,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f519,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f238,f405,f382])).
fof(f520,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f238,f405,f382])).
fof(f521,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f237,f405,f382])).
fof(f522,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f237,f405,f382])).
fof(f593,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl14),
  inference(cnf_transformation,[],[f593_D])).
fof(f593_D,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f525,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f238,f405,f386])).
fof(f589,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl10),
  inference(cnf_transformation,[],[f589_D])).
fof(f589_D,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f526,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f405,f238,f405,f386])).
fof(f527,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f243,f237,f405,f386])).
fof(f528,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f405,f237,f405,f386])).
fof(f591,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','T_0') | $spl12),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','T_0') <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f529,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f243,f405,f386])).
fof(f530,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f243,f405,f386])).
fof(f531,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f405,f386])).
fof(f532,plain,(
  sP9('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f405,f386])).
fof(f587,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl9),
  inference(cnf_transformation,[],[f587_D])).
fof(f587_D,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f535,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f238,f405,f389])).
fof(f536,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f238,f405,f389])).
fof(f537,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f237,f405,f389])).
fof(f538,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f237,f405,f389])).
fof(f545,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f238,f238,f405,f398])).
fof(f546,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f237,f238,f405,f398])).
fof(f547,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f238,f237,f405,f398])).
fof(f548,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f237,f237,f405,f398])).
fof(f585,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(cnf_transformation,[],[f585_D])).
fof(f585_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f549,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f243,f238,f405,f398])).
fof(f550,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f243,f238,f405,f398])).
fof(f551,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f238,f405,f398])).
fof(f552,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f238,f405,f398])).
fof(f553,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f243,f237,f405,f398])).
fof(f554,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f243,f237,f405,f398])).
fof(f555,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f237,f405,f398])).
fof(f556,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f237,f405,f398])).
fof(f557,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f238,f405,f403])).
fof(f558,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f243,f237,f405,f403])).
fof(f559,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f238,f243,f405,f403])).
fof(f560,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f237,f243,f405,f403])).
fof(f561,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f238,f405,f405,f403])).
fof(f562,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f237,f405,f405,f403])).
fof(f407,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | $spl2),
  inference(cnf_transformation,[],[f407_D])).
fof(f407_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f235,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))))),
  inference(cnf_transformation,[],[f153])).
fof(f405,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)))) | $spl1),
  inference(cnf_transformation,[],[f405_D])).
fof(f405_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f236,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK14))))),
  inference(cnf_transformation,[],[f153])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 276
% Passive clauses: 817
% Generated clauses: 3166
% Final active clauses: 272
% Final passive clauses: 659
% Input formulas: 41
% Initial clauses: 166
% 
% Duplicate literals: 58
% Trivial inequalities: 4
% Fw demodulations: 423
% Bw demodulations: 8
% Global subsumptions: 3
% 
% Simple tautologies: 129
% Equational tautologies: 10
% Forward subsumptions: 1264
% Backward subsumptions: 2
% 
% Binary resolution: 1107
% Unit resulting resolution: 1161
% Factoring: 7
% Forward superposition: 49
% Backward superposition: 150
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 1828
% SAT solver unit clauses: 4
% SAT solver binary clauses: 1388
% 
% Sat splits: 6
% 
% Memory used [KB]: 4477
% Time elapsed: 0.076 s
% ------------------------------
% ------------------------------
% Success in time 18.098 s
