Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f19622,plain,(
  $false),
  inference(subsumption_resolution,[],[f19393,f4793])).
fof(f4793,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f4692,f1397])).
fof(f1397,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(cnf_transformation,[],[f276])).
fof(f276,axiom,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.rat.RAT_|49|','thm.numeral.numeral_suc|split|0')).
fof(f4692,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(backward_demodulation,[],[f4612,f1398])).
fof(f1398,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f24])).
fof(f24,axiom,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_|49|','thm.rat.rat_of_num_def|split|1')).
fof(f4612,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  inference(superposition,[],[f1439,f1426])).
fof(f1426,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f528])).
fof(f528,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X38] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.|48|_0')))),
  file('thm.rat.RAT_|49|','thm.arithmetic.MULT_|48|')).
fof(f1439,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f374])).
fof(f374,axiom,(
  ! [X12] : s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X12)))),
  file('thm.rat.RAT_|49|','thm.numeral.numeral_mult|split|0')).
fof(f19393,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(backward_demodulation,[],[f7390,f1385])).
fof(f1385,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f519])).
fof(f519,plain,(
  s('type.rat.rat','const.rat.rat_|49|_0') != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  s('type.rat.rat','const.rat.rat_|49|_0') = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.rat.RAT_|49|','thm.rat.RAT_|49|')).
fof(f7390,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f7389,f1397])).
fof(f7389,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f1402,f4612])).
fof(f1402,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f71])).
fof(f71,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.rat.RAT_|49|','thm.numeral.numeral_distrib|split|13')).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 786
% Passive clauses: 3560
% Generated clauses: 18220
% Final active clauses: 731
% Final passive clauses: 2546
% Input formulas: 517
% Initial clauses: 794
% 
% Duplicate literals: 206
% Trivial inequalities: 25
% Fw subsumption resolutions: 1304
% Fw demodulations: 966
% Bw demodulations: 50
% 
% Simple tautologies: 237
% Equational tautologies: 25
% Forward subsumptions: 11509
% 
% Binary resolution: 4322
% Factoring: 932
% Forward superposition: 5955
% Backward superposition: 3511
% Self superposition: 33
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 262
% Split components: 526
% SAT solver clauses: 389
% SAT solver unit clauses: 19
% SAT solver binary clauses: 81
% 
% Sat splits: 262
% 
% TWLsolver clauses: 389
% TWLsolver calls for satisfiability: 304
% 
% Memory used [KB]: 8187
% Time elapsed: 0.281 s
% ------------------------------
% ------------------------------
% Success in time 0.32 s
