Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (107667)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 154
% Passive clauses: 3963
% Generated clauses: 19121
% Final active clauses: 146
% Final passive clauses: 3736
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 539
% Trivial inequalities: 419
% Fw subsumption resolutions: 1359
% Fw demodulations: 1423
% Bw demodulations: 4
% 
% Simple tautologies: 328
% Equational tautologies: 21
% Forward subsumptions: 10295
% Fw demodulations to eq. taut.: 12
% 
% Binary resolution: 41
% Forward superposition: 8619
% Backward superposition: 6194
% Self superposition: 201
% Equality factoring: 53
% Equality resolution: 30
% 
% Split clauses: 368
% Split components: 736
% SAT solver clauses: 736
% SAT solver unit clauses: 5
% SAT solver binary clauses: 115
% 
% Sat splits: 368
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 736
% TWLsolver calls for satisfiability: 91
% 
% Memory used [KB]: 18933
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (107684)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 130
% Passive clauses: 1321
% Generated clauses: 15707
% Final active clauses: 118
% Final passive clauses: 1047
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 551
% Trivial inequalities: 205
% Fw subsumption resolutions: 543
% Fw demodulations: 985
% Bw demodulations: 121
% Condensations: 867
% 
% Simple tautologies: 164
% Equational tautologies: 24
% Forward subsumptions: 8918
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 6084
% Backward superposition: 4538
% Self superposition: 209
% Equality factoring: 72
% Equality resolution: 19
% 
% Split clauses: 106
% Split components: 212
% SAT solver clauses: 2074
% SAT solver unit clauses: 235
% SAT solver binary clauses: 50
% 
% Sat splits: 106
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2074
% TWLsolver calls for satisfiability: 263
% 
% Memory used [KB]: 22515
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f102605,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f102469])).
fof(f102469,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',sK0))))) != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',sK0)))))),
  inference(superposition,[],[f8534,f20405])).
fof(f20405,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(forward_demodulation,[],[f20372,f329])).
fof(f329,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0'),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0'),s('type.canonical.spolynom'('type.num.num'),X0)))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X45] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X45))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0'),s('type.canonical.spolynom'('type.num.num'),X45)))),
  file('thm.numRing.num_ring_thms|split|50','ARITY_0_1_const.numRing.num_spolynom_normalize')).
fof(f20372,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0'),s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(superposition,[],[f344,f311])).
fof(f311,plain,(
  s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0') = s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f61])).
fof(f61,axiom,(
  s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0') = s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|50','thm.numRing.num_spolynom_normalize_def')).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X7] : ! [X45,X46] : s('type.canonical.canonical_sum'(X7),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X7),X45),s('type.canonical.spolynom'(X7),X46))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.canonical.spolynom'(X7),'type.canonical.canonical_sum'(X7)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.canonical.spolynom'(X7),X46)))),
  file('thm.numRing.num_ring_thms|split|50','ARITY_1_2_const.canonical.spolynom_normalize')).
fof(f8534,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',sK0))))) != s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'('type.num.num'),X0),s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',sK0)))))) )),
  inference(superposition,[],[f307,f346])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X7] : ! [X12,X6] : s('type.canonical.canonical_sum'(X7),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X7),X12),s('type.canonical.spolynom'(X7),'const.canonical.SPvar_1'(s('type.quote.index',X6))))) = s('type.canonical.canonical_sum'(X7),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|50','thm.canonical.spolynom_normalize_def|split|0')).
fof(f307,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',sK0))))) != s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',sK0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',sK0))))) != s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',sK0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f298])).
fof(f298,plain,(
  ? [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',X0))))) != s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',X0))))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',X6))))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),'const.canonical.SPvar_1'(s('type.quote.index',X6))))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|50','thm.numRing.num_ring_thms|split|50')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 494
% Passive clauses: 8235
% Generated clauses: 105716
% Final active clauses: 248
% Final passive clauses: 2129
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 2701
% 
% Duplicate literals: 3141
% Trivial inequalities: 1147
% Fw subsumption resolutions: 2859
% Fw demodulations: 12210
% Bw demodulations: 1791
% Condensations: 4392
% Global subsumptions: 180
% 
% Simple tautologies: 567
% Equational tautologies: 251
% Forward subsumptions: 71985
% Fw demodulations to eq. taut.: 20
% Bw demodulations to eq. taut.: 116
% 
% Unit resulting resolution: 2581
% Forward superposition: 40319
% Backward superposition: 31788
% Self superposition: 644
% Equality factoring: 202
% Equality resolution: 47
% 
% Split clauses: 259
% Split components: 523
% SAT solver clauses: 9128
% SAT solver unit clauses: 984
% SAT solver binary clauses: 2237
% 
% Sat splits: 259
% Sat splitting refutations: 193
% 
% Memory used [KB]: 142513
% Time elapsed: 5.740 s
% ------------------------------
% ------------------------------
% Success in time 7.069 s
