Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (84261)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 54
% Passive clauses: 2055
% Generated clauses: 9904
% Final active clauses: 52
% Final passive clauses: 1945
% Input formulas: 30
% Initial clauses: 31
% 
% Duplicate literals: 601
% Trivial inequalities: 100
% Fw subsumption resolutions: 1021
% Fw demodulations: 104
% 
% Simple tautologies: 192
% Equational tautologies: 109
% Forward subsumptions: 4528
% 
% Binary resolution: 25
% Forward superposition: 3918
% Backward superposition: 3334
% Self superposition: 401
% Equality factoring: 326
% Equality resolution: 10
% 
% Split clauses: 375
% Split components: 754
% SAT solver clauses: 1186
% SAT solver unit clauses: 251
% SAT solver binary clauses: 401
% 
% Sat splits: 375
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1182
% TWLsolver calls for satisfiability: 34
% 
% Memory used [KB]: 7036
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (84286)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 39
% Passive clauses: 242
% Generated clauses: 5971
% Final active clauses: 36
% Final passive clauses: 201
% Input formulas: 30
% Initial clauses: 31
% 
% Duplicate literals: 835
% Trivial inequalities: 53
% Fw subsumption resolutions: 152
% Fw demodulations: 49
% Bw demodulations: 1
% Condensations: 1679
% 
% Simple tautologies: 96
% Equational tautologies: 57
% Forward subsumptions: 2223
% 
% Binary resolution: 16
% Forward superposition: 1578
% Backward superposition: 1004
% Self superposition: 210
% Equality factoring: 174
% Equality resolution: 5
% 
% Split clauses: 29
% Split components: 58
% SAT solver clauses: 657
% SAT solver unit clauses: 56
% SAT solver binary clauses: 66
% 
% Sat splits: 29
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 657
% TWLsolver calls for satisfiability: 62
% 
% Memory used [KB]: 5756
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f11843,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f11842])).
fof(f11842,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1)))),
  inference(superposition,[],[f11362,f11461])).
fof(f11461,plain,(
  ( ! [X6,X7] : (s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))) )),
  inference(superposition,[],[f112,f100])).
fof(f100,plain,(
  ( ! [X0,X1] : (s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))),s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_cs_def|split|0','thm.ring.semi_ring_of_def')).
fof(f112,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X15,X16,X17,X18] : s(X6,X15) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X15),s(X6,X16),s(fun(X6,fun(X6,X6)),X17),s(fun(X6,fun(X6,X6)),X18)))))),
  file('thm.ringNorm.interp_cs_def|split|0','thm.semi_ring.semi_ring_SR|48|')).
fof(f11362,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(sK0),sK1)))))),
  inference(superposition,[],[f92,f10890])).
fof(f10890,plain,(
  ( ! [X185,X186,X184] : (s(X184,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X184),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X184),X185))))) = s(X184,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X184),X185),s('type.quote.varmap'(X184),X186),s('type.canonical.canonical_sum'(X184),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f10631,f101])).
fof(f101,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11,X8] : s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X11))) = s(X6,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X6),X11),s('type.quote.varmap'(X6),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.interp_cs_def|split|0','thm.canonical.interp_cs_def|split|0')).
fof(f10631,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f10479,f111])).
fof(f111,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X19,X20,X21] : s(X6,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X6),X19),s('type.quote.varmap'(X6),X20),s('type.canonical.canonical_sum'(X6),X21))) = s(X6,happ(s(fun('type.canonical.canonical_sum'(X6),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.canonical.canonical_sum'(X6),X6)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X6),X19))),s('type.quote.varmap'(X6),X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.interp_cs_def|split|0','ARITY_1_3_const.ringNorm.r_interp_cs')).
fof(f10479,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f110,f97])).
fof(f97,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.canonical.canonical_sum'(X6),X6)),'const.ringNorm.r_interp_cs_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.canonical.canonical_sum'(X6),X6)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_cs_def|split|0','thm.ringNorm.r_interp_cs_def')).
fof(f110,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X19,X20,X21] : s(X6,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X6),X19),s('type.quote.varmap'(X6),X20),s('type.canonical.canonical_sum'(X6),X21))) = s(X6,happ(s(fun('type.canonical.canonical_sum'(X6),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.canonical.canonical_sum'(X6),X6)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X6),X19))),s('type.quote.varmap'(X6),X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.interp_cs_def|split|0','ARITY_1_3_const.canonical.interp_cs')).
fof(f92,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f87])).
fof(f87,plain,(
  ? [X0,X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ~! [X0,X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ~! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ringNorm.r_interp_cs_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.interp_cs_def|split|0','thm.ringNorm.interp_cs_def|split|0')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 55
% Passive clauses: 704
% Generated clauses: 11507
% Final active clauses: 51
% Final passive clauses: 535
% Input formulas: 30
% Initial clauses: 31
% 
% Duplicate literals: 1509
% Trivial inequalities: 65
% Fw subsumption resolutions: 323
% Fw demodulations: 143
% Bw demodulations: 110
% Condensations: 2738
% 
% Simple tautologies: 147
% Equational tautologies: 135
% Forward subsumptions: 5850
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 81
% Forward superposition: 2943
% Backward superposition: 3187
% Self superposition: 324
% Equality factoring: 246
% Equality resolution: 6
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 759
% SAT solver unit clauses: 40
% SAT solver binary clauses: 55
% 
% Sat splits: 26
% Sat splitting refutations: 1
% 
% Memory used [KB]: 8059
% Time elapsed: 1.709 s
% ------------------------------
% ------------------------------
% Success in time 3.027 s
