Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.prelim
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.prelim
% SZS output start Proof for thm.prelim
fof(f14780,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14747,f7277])).
fof(f7277,plain,(
  $false | $spl29),
  inference(subsumption_resolution,[],[f7272,f998])).
fof(f998,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl29),
  inference(cnf_transformation,[],[f998_D])).
fof(f998_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f7272,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(trivial_inequality_removal,[],[f7197])).
fof(f7197,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(superposition,[],[f190,f7126])).
fof(f7126,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f7125,f714])).
fof(f714,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(equality_resolution,[],[f191])).
fof(f191,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X13))),
  file('thm.prelim.ordering_case_def|split|0','thm.numeral.numeral_distrib|split|18')).
fof(f7125,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(subsumption_resolution,[],[f7124,f149])).
fof(f149,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.prelim.ordering_case_def|split|0',aHL_TRUTH)).
fof(f7124,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f7081,f164])).
fof(f164,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X13] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))),
  file('thm.prelim.ordering_case_def|split|0','thm.numeral.numeral_distrib|split|26')).
fof(f7081,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f6876,f714])).
fof(f6876,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f194,f6121])).
fof(f6121,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f203,f714])).
fof(f203,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X13,X14] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X13))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|0','thm.numeral.numeral_distrib|split|31')).
fof(f194,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X13] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X13))),
  file('thm.prelim.ordering_case_def|split|0','thm.numeral.numeral_distrib|split|30')).
fof(f190,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f134])).
fof(f14747,plain,(
  $false | $spl28),
  inference(trivial_inequality_removal,[],[f14746])).
fof(f14746,plain,(
  s(sK0,sK1) != s(sK0,sK1) | $spl28),
  inference(superposition,[],[f148,f14283])).
fof(f14283,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3)))) ) | $spl28),
  inference(forward_demodulation,[],[f14282,f198])).
fof(f198,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X11),s(X6,X12)))),
  file('thm.prelim.ordering_case_def|split|0','thm.bool.COND_CLAUSES|split|0')).
fof(f14282,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,'const.bool.COND_3'(s(bool,sK4(X3,X2,X1,'const.prelim.LESS_0',X0)),s(X0,X2),s(X0,X3)))))) ) | $spl28),
  inference(forward_demodulation,[],[f13862,f3429])).
fof(f3429,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f3379,f172])).
fof(f172,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X13] : s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|0','thm.numeral.numeral_lt|split|0')).
fof(f3379,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f179,f155])).
fof(f155,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f23])).
fof(f23,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.prelim.ordering_case_def|split|0','thm.numeral.numeral_distrib|split|13')).
fof(f179,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X13))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|0','thm.numeral.numeral_distrib|split|21')).
fof(f13862,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(X0,X1),s(X0,'const.bool.COND_3'(s(bool,sK4(X3,X2,X1,'const.prelim.LESS_0',X0)),s(X0,X2),s(X0,X3)))))) ) | $spl28),
  inference(superposition,[],[f12940,f152])).
fof(f152,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  inference(cnf_transformation,[],[f61])).
fof(f61,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  file('thm.prelim.ordering_case_def|split|0','thm.prelim.ordering2num_thm|split|0')).
fof(f12940,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) ) | $spl28),
  inference(forward_demodulation,[],[f12939,f7320])).
fof(f7320,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl28),
  inference(backward_demodulation,[],[f997,f155])).
fof(f997,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(cnf_transformation,[],[f997_D])).
fof(f997_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f12939,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) ) | $spl28),
  inference(forward_demodulation,[],[f202,f997])).
fof(f202,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,sK4(X4,X3,X2,X1,X0))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,sK4(X4,X3,X2,X1,X0)))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((~p(s(bool,X5)) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,X5))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : (((~p(s(bool,X5)) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,X5)))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((p(s(bool,X5)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ? [X5] : ((p(s(bool,X5)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X10,X7,X8,X9] : ? [X15] : ((p(s(bool,X15)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X10),s(X6,X7),s(X6,X8),s(X6,X9))) = s(X6,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X6,X7),s(X6,'const.bool.COND_3'(s(bool,X15),s(X6,X8),s(X6,X9))))))),
  file('thm.prelim.ordering_case_def|split|0','thm.prelim.ordering_CASE')).
fof(f148,plain,(
  s(sK0,sK1) != s(sK0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(sK0,sK1),s(sK0,sK2),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  s(sK0,sK1) != s(sK0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(sK0,sK1),s(sK0,sK2),s(sK0,sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f127])).
fof(f127,plain,(
  ? [X0,X1,X2,X3] : s(X0,X1) != s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ~! [X0,X1,X2,X3] : s(X0,X1) = s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,X1) = s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,X7) = s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X6,X7),s(X6,X8),s(X6,X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,X7) = s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X6,X7),s(X6,X8),s(X6,X9)))),
  file('thm.prelim.ordering_case_def|split|0','thm.prelim.ordering_case_def|split|0')).
% SZS output end Proof for thm.prelim
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 294
% Passive clauses: 2187
% Generated clauses: 14629
% Final active clauses: 132
% Final passive clauses: 1533
% Input formulas: 63
% Initial clauses: 79
% 
% Duplicate literals: 287
% Trivial inequalities: 138
% Fw subsumption resolutions: 335
% Fw demodulations: 2532
% Bw demodulations: 126
% 
% Simple tautologies: 40
% Equational tautologies: 26
% Forward subsumptions: 9151
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 25
% Forward superposition: 5052
% Backward superposition: 5778
% Self superposition: 83
% Equality factoring: 36
% Equality resolution: 12
% 
% Split clauses: 55
% Split components: 110
% SAT solver clauses: 97
% SAT solver unit clauses: 9
% SAT solver binary clauses: 39
% SAT solver learnt clauses: 1
% 
% Sat splits: 55
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 96
% TWLsolver calls for satisfiability: 44
% 
% Memory used [KB]: 8315
% Time elapsed: 0.265 s
% ------------------------------
% ------------------------------
% Success in time 0.269 s
