Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (108484)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 150
% Passive clauses: 3303
% Generated clauses: 25107
% Final active clauses: 142
% Final passive clauses: 3010
% Input formulas: 69
% Initial clauses: 151
% 
% Duplicate literals: 825
% Trivial inequalities: 450
% Fw subsumption resolutions: 3332
% Fw demodulations: 4577
% Bw demodulations: 7
% 
% Simple tautologies: 501
% Equational tautologies: 58
% Forward subsumptions: 10888
% 
% Binary resolution: 212
% Factoring: 1
% Forward superposition: 6926
% Backward superposition: 8167
% Self superposition: 235
% Equality factoring: 118
% Equality resolution: 34
% 
% Split clauses: 222
% Split components: 445
% SAT solver clauses: 500
% SAT solver unit clauses: 35
% SAT solver binary clauses: 235
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 2
% 
% Sat splits: 222
% 
% TWLsolver clauses: 499
% TWLsolver calls for satisfiability: 73
% 
% Memory used [KB]: 12665
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f5997,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5975,f519,f525,f281,f441,f336,f456,f5751,f5761,f5654,f5655,f359,f473,f539,f541,f274,f435,f282,f440,f5582,f5608,f344,f460,f5576,f5578,f264,f425,f273,f434,f380,f493,f570,f579,f263,f423,f276,f437,f5726,f2151,f5728,f1961,f1986,f286,f442,f1870,f1871,f363,f474,f1710,f1747,f277,f438,f1690,f1702,f1655,f1656,f817,f818,f547,f577,f285,f443,f2464,f2162,f2166,f2168,f2479,f2264,f2195,f2167,f2152,f5652,f5653,f528,f532])).
fof(f532,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl136),
  inference(cnf_transformation,[],[f532_D])).
fof(f532_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f528,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl28)),
  inference(resolution,[],[f442,f456])).
fof(f5653,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)))))) ) | $spl1010),
  inference(cnf_transformation,[],[f5653_D])).
fof(f5653_D,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)))))) ) <=> ~$spl1010),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1010])])).
fof(f5652,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)))))) ) | ($spl62 | $spl136 | $spl994)),
  inference(forward_demodulation,[],[f5623,f532])).
fof(f5623,plain,(
  ( ! [X6] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)))))) ) | ($spl62 | $spl994)),
  inference(superposition,[],[f473,f5608])).
fof(f2152,plain,(
  ( ! [X2,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) ) | $spl462),
  inference(cnf_transformation,[],[f2152_D])).
fof(f2152_D,plain,(
  ( ! [X2,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f2167,plain,(
  ( ! [X12] : (~p(s(bool,X12))) ) | $spl488),
  inference(cnf_transformation,[],[f2167_D])).
fof(f2167_D,plain,(
  ( ! [X12] : (~p(s(bool,X12))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f2195,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X70)) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X69))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X71))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X71) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X69)) ) | ($spl26 | $spl420 | $spl442)),
  inference(subsumption_resolution,[],[f2135,f443])).
fof(f2135,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X70)) | s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X69))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X71))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X71) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X69) | s(bool,'F_0') = s(bool,X70)) ) | ($spl420 | $spl442)),
  inference(superposition,[],[f1871,f1986])).
fof(f2264,plain,(
  $false | ($spl134 | $spl138 | $spl482)),
  inference(subsumption_resolution,[],[f2210,f525])).
fof(f2210,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ($spl138 | $spl482)),
  inference(backward_demodulation,[],[f2162,f541])).
fof(f2479,plain,(
  $false | ($spl2 | $spl488)),
  inference(resolution,[],[f2167,f425])).
fof(f2168,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | $spl490),
  inference(cnf_transformation,[],[f2168_D])).
fof(f2168_D,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f2166,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) | ~p(s(bool,X12)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | ($spl26 | $spl442)),
  inference(subsumption_resolution,[],[f2052,f443])).
fof(f2052,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) | ~p(s(bool,X12)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s(bool,'F_0') = s(bool,X12)) ) | ($spl26 | $spl442)),
  inference(superposition,[],[f443,f1986])).
fof(f2162,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) ) | $spl482),
  inference(cnf_transformation,[],[f2162_D])).
fof(f2162_D,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f2464,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl460 | $spl462 | $spl490)),
  inference(subsumption_resolution,[],[f2404,f2168])).
fof(f2404,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) ) | ($spl460 | $spl462)),
  inference(condensation,[],[f2378])).
fof(f2378,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) ) | ($spl460 | $spl462)),
  inference(backward_demodulation,[],[f2151,f2152])).
fof(f443,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl26),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f285,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.list.LIST_REL_def|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f577,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl144),
  inference(cnf_transformation,[],[f577_D])).
fof(f577_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f547,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl26)),
  inference(superposition,[],[f443,f437])).
fof(f818,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl186),
  inference(cnf_transformation,[],[f818_D])).
fof(f818_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f817,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl2 | $spl144 | $spl146)),
  inference(subsumption_resolution,[],[f805,f425])).
fof(f805,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl144 | $spl146)),
  inference(superposition,[],[f577,f579])).
fof(f1656,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) ) | $spl354),
  inference(cnf_transformation,[],[f1656_D])).
fof(f1656_D,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f1655,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) ) | ($spl62 | $spl146 | $spl186)),
  inference(forward_demodulation,[],[f1626,f579])).
fof(f1626,plain,(
  ( ! [X28,X29] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))) ) | ($spl62 | $spl186)),
  inference(superposition,[],[f818,f473])).
fof(f1702,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) ) | $spl362),
  inference(cnf_transformation,[],[f1702_D])).
fof(f1702_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f1690,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) ) | ($spl138 | $spl354)),
  inference(trivial_inequality_removal,[],[f1687])).
fof(f1687,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) ) | ($spl138 | $spl354)),
  inference(superposition,[],[f1656,f541])).
fof(f438,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK6(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl18),
  inference(cnf_transformation,[],[f438_D])).
fof(f438_D,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK6(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f277,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK6(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK6(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f81])).
fof(f81,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X48] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X48) | ? [X29] : s('type.num.num',X48) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))),
  file('thm.list.LIST_REL_def|split|2','thm.arithmetic.num_CASES')).
fof(f1747,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl366),
  inference(cnf_transformation,[],[f1747_D])).
fof(f1747_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f1710,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl18 | $spl362)),
  inference(superposition,[],[f1702,f438])).
fof(f474,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | $spl64),
  inference(cnf_transformation,[],[f474_D])).
fof(f474_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f363,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.list.LIST_REL_def|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f1871,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | $spl420),
  inference(cnf_transformation,[],[f1871_D])).
fof(f1871_D,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f1870,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))) | ~p(s(bool,X11)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl2 | $spl62 | $spl64 | $spl366)),
  inference(forward_demodulation,[],[f1869,f473])).
fof(f1869,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl2 | $spl64 | $spl366)),
  inference(subsumption_resolution,[],[f1839,f425])).
fof(f1839,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl64 | $spl366)),
  inference(superposition,[],[f474,f1747])).
fof(f442,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f286,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f1986,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,'F_0') = s(bool,X3)) ) | $spl442),
  inference(cnf_transformation,[],[f1986_D])).
fof(f1986_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,'F_0') = s(bool,X3)) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f1961,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s(bool,'F_0') = s(bool,X3)) ) | ($spl24 | $spl420)),
  inference(resolution,[],[f1871,f442])).
fof(f5728,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) | $spl1020),
  inference(cnf_transformation,[],[f5728_D])).
fof(f5728_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) <=> ~$spl1020),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1020])])).
fof(f2151,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl460),
  inference(cnf_transformation,[],[f2151_D])).
fof(f2151_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f5726,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3)))) ) | ($spl442 | $spl1010)),
  inference(duplicate_literal_removal,[],[f5677])).
fof(f5677,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl442 | $spl1010)),
  inference(superposition,[],[f5653,f1986])).
fof(f437,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f437_D])).
fof(f437_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f276,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.LIST_REL_def|split|2',aHL_BOOL_CASES)).
fof(f423,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | $spl1),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f263,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3)))))),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5])],[f171])).
fof(f171,plain,(
  ? [X0,X1,X2,X3,X4] : s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X2)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X2)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X9),s('type.list.list'(X7),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X9),s('type.list.list'(X7),X8)))))),
  file('thm.list.LIST_REL_def|split|2','thm.list.LIST_REL_def|split|2')).
fof(f579,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | $spl146),
  inference(cnf_transformation,[],[f579_D])).
fof(f579_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f570,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | ($spl1 | $spl16)),
  inference(trivial_inequality_removal,[],[f565])).
fof(f565,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(sK2,bool)),sK5),s('type.list.list'(sK1),'const.list.NIL_0'),s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | ($spl1 | $spl16)),
  inference(superposition,[],[f423,f437])).
fof(f493,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) ) | $spl102),
  inference(cnf_transformation,[],[f493_D])).
fof(f493_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f380,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK21(X4,X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK21(X4,X3,X2,X1,X0)),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',sK21(X4,X3,X2,X1,X0)),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21])],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(rectify,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))))) & ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : ! [X10,X27,X28] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.list.list'(X6),X27),s('type.list.list'(X7),X28)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X27))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),X28))) & ! [X29] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X27)))))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,'const.list.EL_2'(s('type.num.num',X29),s('type.list.list'(X6),X27))))),s(X7,'const.list.EL_2'(s('type.num.num',X29),s('type.list.list'(X7),X28)))))))))),
  file('thm.list.LIST_REL_def|split|2','thm.list.LIST_REL_EL_EQN')).
fof(f434,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(cnf_transformation,[],[f434_D])).
fof(f434_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f273,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.LIST_REL_def|split|2','thm.list.LENGTH|split|0')).
fof(f425,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f264,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.LIST_REL_def|split|2',aHL_TRUTH)).
fof(f5578,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | $spl992),
  inference(cnf_transformation,[],[f5578_D])).
fof(f5578_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) <=> ~$spl992),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl992])])).
fof(f5576,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | ($spl2 | $spl10 | $spl102 | $spl146)),
  inference(forward_demodulation,[],[f5575,f434])).
fof(f5575,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | ($spl2 | $spl102 | $spl146)),
  inference(subsumption_resolution,[],[f5550,f425])).
fof(f5550,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),'const.list.CONS_2'(s(sK2,sK4),s('type.list.list'(sK2),sK3))))) | ($spl102 | $spl146)),
  inference(superposition,[],[f493,f579])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) ) | $spl36),
  inference(cnf_transformation,[],[f460_D])).
fof(f460_D,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X45,X15] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X15))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X45),s('type.list.list'(X6),X15)))))),
  file('thm.list.LIST_REL_def|split|2','thm.list.LENGTH|split|1')).
fof(f5608,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))))) | $spl994),
  inference(cnf_transformation,[],[f5608_D])).
fof(f5608_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))))) <=> ~$spl994),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl994])])).
fof(f5582,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))))) | ($spl36 | $spl992)),
  inference(superposition,[],[f5578,f460])).
fof(f440,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f282,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.list.LIST_REL_def|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f435,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) ) | $spl12),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f274,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X29] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))))),
  file('thm.list.LIST_REL_def|split|2','thm.prim_rec.LESS_|48|')).
fof(f541,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl138),
  inference(cnf_transformation,[],[f541_D])).
fof(f541_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f539,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | ($spl12 | $spl20)),
  inference(resolution,[],[f435,f440])).
fof(f473,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) ) | $spl62),
  inference(cnf_transformation,[],[f473_D])).
fof(f473_D,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f359,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X48,X29] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X48),s('type.num.num',X29))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))),
  file('thm.list.LIST_REL_def|split|2','thm.prim_rec.LESS_MONO_EQ')).
fof(f5655,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))),s('type.num.num',X7)))) ) | $spl1012),
  inference(cnf_transformation,[],[f5655_D])).
fof(f5655_D,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))),s('type.num.num',X7)))) ) <=> ~$spl1012),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1012])])).
fof(f5654,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))),s('type.num.num',X7)))) ) | ($spl62 | $spl138 | $spl994)),
  inference(forward_demodulation,[],[f5624,f541])).
fof(f5624,plain,(
  ( ! [X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK2),sK3))),s('type.num.num',X7)))) ) | ($spl62 | $spl994)),
  inference(superposition,[],[f473,f5608])).
fof(f5761,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) ) | $spl1032),
  inference(cnf_transformation,[],[f5761_D])).
fof(f5761_D,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) ) <=> ~$spl1032),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1032])])).
fof(f5751,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) ) | ($spl1012 | $spl1020)),
  inference(backward_demodulation,[],[f5728,f5655])).
fof(f456,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl28),
  inference(cnf_transformation,[],[f456_D])).
fof(f456_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f336,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X29] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0'))))),
  file('thm.list.LIST_REL_def|split|2','thm.prim_rec.NOT_LESS_|48|')).
fof(f441,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) ) | $spl22),
  inference(cnf_transformation,[],[f441_D])).
fof(f441_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f281,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f525,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl134),
  inference(cnf_transformation,[],[f525_D])).
fof(f525_D,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f519,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | ($spl22 | $spl28)),
  inference(resolution,[],[f441,f456])).
fof(f5975,plain,(
  $false | ($spl134 | $spl1032)),
  inference(trivial_inequality_removal,[],[f5927])).
fof(f5927,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl134 | $spl1032)),
  inference(superposition,[],[f525,f5761])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 130
% Passive clauses: 646
% Generated clauses: 5906
% Final active clauses: 91
% Final passive clauses: 292
% Input formulas: 69
% Initial clauses: 151
% 
% Duplicate literals: 38
% Trivial inequalities: 129
% Fw subsumption resolutions: 357
% Fw demodulations: 1061
% Bw demodulations: 238
% Condensations: 67
% Global subsumptions: 17
% 
% Simple tautologies: 80
% Equational tautologies: 12
% Forward subsumptions: 2738
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 105
% Factoring: 1
% Forward superposition: 1404
% Backward superposition: 1549
% Self superposition: 34
% Equality factoring: 4
% Equality resolution: 21
% 
% Split clauses: 35
% Split components: 70
% SAT solver clauses: 887
% SAT solver unit clauses: 178
% SAT solver binary clauses: 43
% SAT solver learnt clauses: 35
% 
% Sat splits: 35
% Sat splitting refutations: 42
% 
% TWLsolver clauses: 816
% TWLsolver calls for satisfiability: 786
% 
% Memory used [KB]: 6012
% Time elapsed: 0.188 s
% ------------------------------
% ------------------------------
% Success in time 0.797 s
