Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (50917)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1330
% Passive clauses: 8068
% Generated clauses: 41274
% Final active clauses: 1293
% Final passive clauses: 6436
% Input formulas: 261
% Initial clauses: 515
% 
% Duplicate literals: 200
% Trivial inequalities: 116
% Fw subsumption resolutions: 2328
% Fw demodulations: 3072
% Bw demodulations: 36
% 
% Simple tautologies: 294
% Equational tautologies: 34
% Forward subsumptions: 26923
% 
% Binary resolution: 7316
% Factoring: 487
% Forward superposition: 12713
% Backward superposition: 14253
% Self superposition: 48
% Equality factoring: 2
% Equality resolution: 24
% 
% Split clauses: 360
% Split components: 722
% SAT solver clauses: 547
% SAT solver unit clauses: 31
% SAT solver binary clauses: 98
% 
% Sat splits: 360
% 
% TWLsolver clauses: 547
% TWLsolver calls for satisfiability: 438
% 
% Memory used [KB]: 22259
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f14463,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14438,f856,f1398,f1396,f2041,f2051,f1006,f1560,f855,f1397,f1395,f14405,f14380,f14382,f14408,f1455,f1457,f1204,f1699,f1025,f1579,f2043,f2053,f1087,f1603])).
fof(f1603,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl300),
  inference(cnf_transformation,[],[f1603_D])).
fof(f1603_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f1087,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(flattening,[],[f410])).
fof(f410,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X6))))),
  file('thm.extreal.lt_infty|split|6','thm.extreal.lt_refl')).
fof(f2053,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) ) | $spl816),
  inference(cnf_transformation,[],[f2053_D])).
fof(f2053_D,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) ) <=> ~$spl816),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl816])])).
fof(f2043,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) ) | ($spl254 | $spl300)),
  inference(resolution,[],[f1560,f1603])).
fof(f1579,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) | $spl276),
  inference(cnf_transformation,[],[f1579_D])).
fof(f1579_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f1025,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f691])).
fof(f691,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.lt_infty|split|6','thm.extreal.le_infty|split|2')).
fof(f1699,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) ) | $spl492),
  inference(cnf_transformation,[],[f1699_D])).
fof(f1699_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f1204,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f779])).
fof(f779,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f511])).
fof(f511,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f510])).
fof(f510,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.lt_infty|split|6','thm.extreal.extreal_lt_def')).
fof(f1457,plain,(
  ~p(s(bool,'F_0')) | $spl61),
  inference(cnf_transformation,[],[f1457_D])).
fof(f1457_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f1455,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f887,f857])).
fof(f857,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.lt_infty|split|6',aHL_TRUTH)).
fof(f887,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f656])).
fof(f656,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f34])).
fof(f34,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.extreal.lt_infty|split|6','thm.bool.NOT_CLAUSES|split|1')).
fof(f14408,plain,(
  $false | ($spl2 | $spl61 | $spl812)),
  inference(subsumption_resolution,[],[f14407,f1457])).
fof(f14407,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl812)),
  inference(forward_demodulation,[],[f1397,f2051])).
fof(f14382,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl3948),
  inference(cnf_transformation,[],[f14382_D])).
fof(f14382_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl3948),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3948])])).
fof(f14380,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl61 | $spl492 | $spl812)),
  inference(subsumption_resolution,[],[f14189,f1457])).
fof(f14189,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl492 | $spl812)),
  inference(superposition,[],[f1699,f2051])).
fof(f14405,plain,(
  $false | ($spl1 | $spl276 | $spl3948)),
  inference(subsumption_resolution,[],[f14384,f1395])).
fof(f14384,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK13) | ($spl276 | $spl3948)),
  inference(resolution,[],[f14382,f1579])).
fof(f1395,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK13) | $spl1),
  inference(cnf_transformation,[],[f1395_D])).
fof(f1395_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK13) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1397,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl2),
  inference(cnf_transformation,[],[f1397_D])).
fof(f1397_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f855,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK13)),
  inference(cnf_transformation,[],[f655])).
fof(f655,plain,(
  (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK13) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK13) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f654])).
fof(f654,plain,(
  ? [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))))),
  inference(nnf_transformation,[],[f572])).
fof(f572,plain,(
  ? [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) <~> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.lt_infty|split|6','thm.extreal.lt_infty|split|6')).
fof(f1560,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl254),
  inference(cnf_transformation,[],[f1560_D])).
fof(f1560_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f1006,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f681])).
fof(f681,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.extreal.lt_infty|split|6','thm.bool.EQ_CLAUSES|split|2')).
fof(f2051,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl812),
  inference(cnf_transformation,[],[f2051_D])).
fof(f2051_D,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) <=> ~$spl812),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl812])])).
fof(f2041,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | ($spl3 | $spl254)),
  inference(resolution,[],[f1560,f1398])).
fof(f1396,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK13) | $spl0),
  inference(cnf_transformation,[],[f1396_D])).
fof(f1396_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK13) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1398,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl3),
  inference(cnf_transformation,[],[f1398_D])).
fof(f1398_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f856,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK13)),
  inference(cnf_transformation,[],[f655])).
fof(f14438,plain,(
  $false | ($spl0 | $spl2 | $spl61 | $spl816)),
  inference(subsumption_resolution,[],[f14437,f1457])).
fof(f14437,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl2 | $spl816)),
  inference(forward_demodulation,[],[f14425,f2053])).
fof(f14425,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1396,f1397])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 540
% Passive clauses: 1958
% Generated clauses: 13566
% Final active clauses: 502
% Final passive clauses: 1380
% Input formulas: 261
% Initial clauses: 515
% 
% Duplicate literals: 93
% Trivial inequalities: 40
% Fw subsumption resolutions: 692
% Fw demodulations: 1136
% Bw demodulations: 67
% Condensations: 40
% Global subsumptions: 34
% 
% Simple tautologies: 256
% Equational tautologies: 6
% Forward subsumptions: 7105
% 
% Binary resolution: 1832
% Factoring: 58
% Forward superposition: 3467
% Backward superposition: 3489
% Self superposition: 39
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 2736
% SAT solver unit clauses: 702
% SAT solver binary clauses: 159
% SAT solver learnt clauses: 70
% SAT solver learnt literals: 11
% 
% Sat splits: 9
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 2597
% TWLsolver calls for satisfiability: 4870
% 
% Memory used [KB]: 13304
% Time elapsed: 0.266 s
% ------------------------------
% ------------------------------
% Success in time 0.9 s
