Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.measure
% (98043)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 248
% Passive clauses: 3040
% Generated clauses: 18467
% Final active clauses: 237
% Final passive clauses: 2552
% Input formulas: 28
% Initial clauses: 85
% 
% Duplicate literals: 952
% Trivial inequalities: 275
% Fw subsumption resolutions: 3093
% Fw demodulations: 2030
% Bw demodulations: 6
% 
% Simple tautologies: 358
% Equational tautologies: 18
% Forward subsumptions: 7154
% 
% Binary resolution: 2029
% Forward superposition: 6476
% Backward superposition: 3156
% Self superposition: 68
% Equality factoring: 32
% Equality resolution: 50
% 
% Split clauses: 846
% Split components: 1704
% SAT solver clauses: 1728
% SAT solver unit clauses: 16
% SAT solver binary clauses: 126
% 
% Sat splits: 846
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1725
% TWLsolver calls for satisfiability: 323
% 
% Memory used [KB]: 11257
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.measure
% (98060)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 702
% Generated clauses: 10921
% Final active clauses: 105
% Final passive clauses: 345
% Input formulas: 28
% Initial clauses: 85
% 
% Duplicate literals: 875
% Trivial inequalities: 142
% Fw subsumption resolutions: 695
% Fw demodulations: 958
% Bw demodulations: 229
% Condensations: 1141
% Global subsumptions: 8
% 
% Simple tautologies: 275
% Equational tautologies: 40
% Forward subsumptions: 3669
% Bw demodulations to eq. taut.: 13
% 
% Binary resolution: 548
% Forward superposition: 2374
% Backward superposition: 2938
% Self superposition: 136
% Equality factoring: 116
% Equality resolution: 20
% 
% Split clauses: 180
% Split components: 383
% SAT solver clauses: 1129
% SAT solver unit clauses: 87
% SAT solver binary clauses: 36
% SAT solver learnt clauses: 16
% 
% Sat splits: 180
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1097
% TWLsolver calls for satisfiability: 818
% 
% Memory used [KB]: 9722
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.measure
% (98091)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 475
% Passive clauses: 3701
% Generated clauses: 138661
% Final active clauses: 252
% Final passive clauses: 969
% Input formulas: 28
% Initial clauses: 85
% Discarded non-redundant clauses: 98395
% 
% Duplicate literals: 8010
% Trivial inequalities: 1741
% Fw subsumption resolutions: 13917
% Fw demodulations: 15789
% Bw demodulations: 882
% Condensations: 14653
% Global subsumptions: 86
% 
% Simple tautologies: 2035
% Equational tautologies: 1099
% Forward subsumptions: 74676
% Fw demodulations to eq. taut.: 52
% Bw demodulations to eq. taut.: 57
% 
% Unit resulting resolution: 3976
% Forward superposition: 39451
% Backward superposition: 37526
% Self superposition: 659
% Equality factoring: 772
% Equality resolution: 134
% 
% Split clauses: 222
% Split components: 450
% SAT solver clauses: 4297
% SAT solver unit clauses: 46
% SAT solver binary clauses: 235
% 
% Sat splits: 222
% Sat splitting refutations: 64
% 
% Memory used [KB]: 70745
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.measure
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.measure
% SZS output start Proof for thm.measure
fof(f4683,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f214,f215,f216,f217,f218,f243,f242,f219,f221,f232,f240,f279,f272,f336,f377,f379,f375,f374,f412,f413,f380,f409,f408,f378,f747,f748,f762,f764,f765,f766,f769,f779,f786,f799,f801,f802,f803,f806,f816,f823,f841,f843,f844,f845,f848,f858,f865,f902,f908,f909,f911,f906,f904,f912,f918,f919,f921,f916,f914,f876,f878,f879,f881,f883,f894,f922,f1132,f1318,f1320,f1321,f1322,f1325,f1335,f1342,f1353,f1355,f1356,f1358,f1360,f1371,f1394,f1396,f1397,f1398,f1401,f1411,f1419,f1434,f1436,f1437,f1438,f1441,f1451,f1459,f1649,f1651,f1652,f1653,f1656,f1666,f1673,f1859,f1861,f1862,f1864,f1866,f1877,f1379,f1892,f1893,f1906,f1894,f1895,f1910,f1896,f1904,f1897,f1912,f1908,f1902,f1900,f1921,f1923,f1924,f1925,f1928,f1938,f1945,f2475,f2477,f2478,f2480,f2482,f2493,f2501,f2515,f2517,f2518,f2520,f2522,f2533,f2541,f2557,f2559,f2560,f2561,f2564,f2574,f2582,f2591,f2593,f2594,f2595,f2598,f2608,f2616,f2629,f2631,f2632,f2634,f2636,f2647,f2657,f2686,f2688,f2689,f2690,f2693,f2703,f2711,f2734,f2736,f2737,f2739,f2741,f2752,f2762,f2791,f2793,f2794,f2795,f2798,f2808,f2816,f2827,f2829,f2830,f2831,f2834,f2844,f2862,f2886,f3056,f3058,f3059,f3060,f3063,f3073,f3081,f3261,f3263,f3264,f3265,f3268,f3278,f3286,f3294,f3296,f3297,f3299,f3301,f3312,f3321,f3445,f3447,f3453,f3454,f3455,f3456,f3459,f3460,f3464,f3465,f3477,f3478,f3481,f3483,f3446,f3487,f3489,f3490,f3492,f3494,f3505,f3448,f3704,f3706,f3707,f3708,f3711,f3721,f3729,f3858,f3860,f3866,f3867,f3868,f3869,f3872,f3873,f3877,f3878,f3890,f3891,f3894,f3896,f3859,f3484,f3899,f3901,f3902,f3904,f3906,f3917,f3861,f3939,f3941,f3942,f3944,f3946,f3957,f3984,f3986,f3987,f3989,f3991,f4002,f4024,f4026,f4027,f4029,f4031,f4042,f4050,f4072,f4074,f4075,f4076,f4079,f4089,f4100,f4263,f4265,f4266,f4268,f4270,f4281,f4011,f4312,f4314,f4315,f4316,f4319,f4329,f4341,f4448,f4450,f4451,f4453,f4455,f4466,f4475,f4602,f4608,f4610,f241,f4614,f245,f4619,f244,f4630,f246,f3965,f4677])).
fof(f4677,plain,(
  $false | ($spl28 | $spl31)),
  inference(subsumption_resolution,[],[f4648,f908])).
fof(f4648,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | $spl31),
  inference(resolution,[],[f911,f760])).
fof(f760,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X0))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X0),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) )),
  inference(resolution,[],[f212,f204])).
fof(f204,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X6))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X6))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))))) & (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK9(X4,X3,X2,X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),sK9(X4,X3,X2,X1,X0)))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X6))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))))) & (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))))),
  inference(rectify,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))))) & (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))))) & ((~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) <=> (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))))))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) <=> (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) <=> (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7,X6] : ! [X8,X9,X21] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X21),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9)))))) <=> (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X21),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8))),s(fun(X7,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9)))))))) & ! [X19] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,bool),X19),s(fun(fun(X7,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X21),s(fun(X7,bool),X19))),s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))))))),
  file('thm.measure.MEASURABLE_SUBSET','thm.measure.IN_MEASURABLE')).
fof(f212,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),
  inference(resolution,[],[f126,f185])).
fof(f185,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X19,X10] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X19),s(fun(X6,bool),X10)))) <=> ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X19)))) => p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.measure.MEASURABLE_SUBSET','thm.pred_set.SUBSET_DEF')).
fof(f126,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f63])).
fof(f63,plain,(
  ? [X0,X1,X2,X3] : ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))))))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2,X3] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))))))))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1] : ! [X2,X3] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X3))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9))),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),'const.measure.sigma_2'(s(fun(X7,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9))),s(fun(fun(X7,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9))),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),'const.measure.sigma_2'(s(fun(X7,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9))),s(fun(fun(X7,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X9))))))))))),
  file('thm.measure.MEASURABLE_SUBSET','thm.measure.MEASURABLE_SUBSET')).
fof(f3965,plain,(
  ( ! [X14,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X14))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X14))) | ~p(s(bool,X16))) ) | $spl104),
  inference(cnf_transformation,[],[f3965_D])).
fof(f3965_D,plain,(
  ( ! [X14,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X14))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X14))) | ~p(s(bool,X16))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f246,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl12),
  inference(cnf_transformation,[],[f246_D])).
fof(f246_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4630,plain,(
  ( ! [X216,X215,X217] : (p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X215,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X215))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X215)))) )),
  inference(subsumption_resolution,[],[f4569,f160])).
fof(f160,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) & p(s(bool,X10))) <=> p(s(bool,'F_0')))),
  file('thm.measure.MEASURABLE_SUBSET','thm.bool.AND_CLAUSES|split|2')).
fof(f4569,plain,(
  ( ! [X216,X215,X217] : (p(s(bool,'F_0')) | p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X217) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X215,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X215))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X215)))) )),
  inference(superposition,[],[f210,f3276])).
fof(f3276,plain,(
  ( ! [X71] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X71,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X71))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X71)))) )),
  inference(resolution,[],[f2643,f178])).
fof(f178,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.measure.MEASURABLE_SUBSET','thm.bool.EQ_CLAUSES|split|1')).
fof(f2643,plain,(
  ( ! [X52] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X52)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X52,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X52)))) )),
  inference(resolution,[],[f2474,f174])).
fof(f174,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.measure.MEASURABLE_SUBSET','thm.bool.EQ_CLAUSES|split|3')).
fof(f2474,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X1,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X1))))) )),
  inference(resolution,[],[f1393,f186])).
fof(f186,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f1393,plain,(
  ( ! [X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X3))))) )),
  inference(resolution,[],[f831,f184])).
fof(f184,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f831,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))))),
  inference(resolution,[],[f792,f188])).
fof(f188,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X5))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X4] : (~p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X4)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK6(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK6(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | (p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),sK7(X1,X0))))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),sK7(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),sK7(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X5))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X4] : (~p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X4)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2] : (p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(rectify,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : (~p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2] : (p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : (~p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & ((~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2] : (p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : (~p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f792,plain,(
  sP0(sK1,sK3)),
  inference(resolution,[],[f757,f197])).
fof(f197,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | sP0(X0,X1)) & (~sP0(X0,X1) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> sP0(X0,X1))),
  inference(definition_folding,[],[f75,f79])).
fof(f75,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : (~p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X18] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18)))))) & ! [X19] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X19),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18))),s(fun(X6,bool),X19))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18))))))) & ! [X20] : ((p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X6,bool),bool),X20)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,bool),bool),X20),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X20))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X18)))))))))),
  file('thm.measure.MEASURABLE_SUBSET','thm.measure.SIGMA_ALGEBRA')).
fof(f757,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),
  inference(resolution,[],[f212,f201])).
fof(f201,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f210,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f131])).
fof(f131,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11,X16,X12,X17] : ((s(bool,X11) = s(bool,X16) & (p(s(bool,X16)) => s(bool,X12) = s(bool,X17))) => ((p(s(bool,X11)) => p(s(bool,X12))) <=> (p(s(bool,X16)) => p(s(bool,X17)))))),
  file('thm.measure.MEASURABLE_SUBSET','thm.bool.IMP_CONG')).
fof(f244,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) | $spl8),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4619,plain,(
  ( ! [X171,X169,X170,X168] : (~p(s(bool,X169)) | ~p(s(bool,X170)) | s(bool,X170) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X168,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X168))) | s(bool,X169) != s(bool,X171) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X168)))) )),
  inference(subsumption_resolution,[],[f4538,f144])).
fof(f144,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X10)))),
  file('thm.measure.MEASURABLE_SUBSET','thm.bool.IMP_CLAUSES|split|4')).
fof(f4538,plain,(
  ( ! [X171,X169,X170,X168] : (p(s(bool,'F_0')) | ~p(s(bool,X169)) | ~p(s(bool,X170)) | s(bool,X170) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X168,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X168))) | s(bool,X169) != s(bool,X171) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X168)))) )),
  inference(superposition,[],[f134,f3276])).
fof(f134,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f245,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl10),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f4614,plain,(
  ( ! [X138,X136,X139,X137] : (p(s(bool,X137)) | ~p(s(bool,X138)) | ~p(s(bool,X139)) | s(bool,X138) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X136,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X136))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X136)))) )),
  inference(subsumption_resolution,[],[f4530,f160])).
fof(f4530,plain,(
  ( ! [X138,X136,X139,X137] : (p(s(bool,'F_0')) | p(s(bool,X137)) | ~p(s(bool,X138)) | ~p(s(bool,X139)) | s(bool,X138) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X136,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X136))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X136)))) )),
  inference(superposition,[],[f133,f3276])).
fof(f133,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f241,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f4610,plain,(
  ( ! [X118,X116,X119,X117] : (p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X117) != s(bool,X119) | s(bool,X118) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X116,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X116))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X116)))) )),
  inference(subsumption_resolution,[],[f4525,f160])).
fof(f4525,plain,(
  ( ! [X118,X116,X119,X117] : (p(s(bool,'F_0')) | p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X117) != s(bool,X119) | s(bool,X118) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X116,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X116))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X116)))) )),
  inference(superposition,[],[f132,f3276])).
fof(f132,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f4608,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X88,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X88))) | s(bool,X89) != s(bool,X91) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X88)))) )),
  inference(subsumption_resolution,[],[f4518,f144])).
fof(f4518,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,'F_0')) | ~p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X88,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X88))) | s(bool,X89) != s(bool,X91) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X88)))) )),
  inference(superposition,[],[f130,f3276])).
fof(f130,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f4602,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,X63) | s(bool,X62) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X60,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X60))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X60)))) )),
  inference(subsumption_resolution,[],[f4511,f160])).
fof(f4511,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'F_0')) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,X63) | s(bool,X62) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X60,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X60))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X60)))) )),
  inference(superposition,[],[f128,f3276])).
fof(f128,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f4475,plain,(
  ( ! [X4,X2] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X2)))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X2),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) ) | $spl114),
  inference(cnf_transformation,[],[f4475_D])).
fof(f4475_D,plain,(
  ( ! [X4,X2] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X2)))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X2),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4466,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X44),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f875,f210])).
fof(f875,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X0),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X0))))) )),
  inference(resolution,[],[f213,f184])).
fof(f213,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))),
  inference(resolution,[],[f126,f186])).
fof(f4455,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X30)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X30),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f875,f134])).
fof(f4453,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X22),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f875,f133])).
fof(f4451,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X14),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f875,f132])).
fof(f4450,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X10)))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X10),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f875,f130])).
fof(f4448,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X2)))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),X2),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f875,f128])).
fof(f4341,plain,(
  ( ! [X12,X10] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X10),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X10))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) ) | $spl112),
  inference(cnf_transformation,[],[f4341_D])).
fof(f4341_D,plain,(
  ( ! [X12,X10] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X10),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X10))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4329,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X51),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X51))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f760,f210])).
fof(f4319,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X38),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X38))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))) | s(bool,X40) != s(bool,X41)) )),
  inference(resolution,[],[f760,f134])).
fof(f4316,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X26),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X26))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f760,f133])).
fof(f4315,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X22),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X22))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f760,f132])).
fof(f4314,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X18),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X18))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))) | s(bool,X20) != s(bool,X21)) )),
  inference(resolution,[],[f760,f130])).
fof(f4312,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),X10),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),X10))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f760,f128])).
fof(f4011,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl106),
  inference(cnf_transformation,[],[f4011_D])).
fof(f4011_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4281,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl35),
  inference(resolution,[],[f916,f210])).
fof(f4270,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))) | s(bool,X22) != s(bool,X24)) ) | $spl35),
  inference(resolution,[],[f916,f134])).
fof(f4268,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl35),
  inference(resolution,[],[f916,f133])).
fof(f4266,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl35),
  inference(resolution,[],[f916,f132])).
fof(f4265,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))) | s(bool,X7) != s(bool,X9)) ) | $spl35),
  inference(resolution,[],[f916,f130])).
fof(f4263,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl35),
  inference(resolution,[],[f916,f128])).
fof(f4100,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl110),
  inference(cnf_transformation,[],[f4100_D])).
fof(f4100_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4089,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl34),
  inference(resolution,[],[f915,f210])).
fof(f915,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) | $spl34),
  inference(cnf_transformation,[],[f915_D])).
fof(f915_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4079,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))) | s(bool,X25) != s(bool,X26)) ) | $spl34),
  inference(resolution,[],[f915,f134])).
fof(f4076,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl34),
  inference(resolution,[],[f915,f133])).
fof(f4075,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl34),
  inference(resolution,[],[f915,f132])).
fof(f4074,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))) | s(bool,X10) != s(bool,X11)) ) | $spl34),
  inference(resolution,[],[f915,f130])).
fof(f4072,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl34),
  inference(resolution,[],[f915,f128])).
fof(f4050,plain,(
  ( ! [X14,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X14))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X14))) | ~p(s(bool,X16))) ) | $spl108),
  inference(cnf_transformation,[],[f4050_D])).
fof(f4050_D,plain,(
  ( ! [X14,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X14))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X14))) | ~p(s(bool,X16))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4042,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X56))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X56,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X56)))) )),
  inference(resolution,[],[f3857,f210])).
fof(f3857,plain,(
  ( ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X6,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X6)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X6)))) )),
  inference(subsumption_resolution,[],[f3774,f144])).
fof(f3774,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X6,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X6)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X6)))) )),
  inference(superposition,[],[f186,f2806])).
fof(f2806,plain,(
  ( ! [X71] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X71))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X71)))) )),
  inference(resolution,[],[f2529,f178])).
fof(f2529,plain,(
  ( ! [X40] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X40)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X40)))) )),
  inference(resolution,[],[f1433,f174])).
fof(f1433,plain,(
  ( ! [X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X3))))) )),
  inference(resolution,[],[f930,f184])).
fof(f930,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),
  inference(resolution,[],[f834,f188])).
fof(f834,plain,(
  sP0(sK2,sK4)),
  inference(resolution,[],[f758,f197])).
fof(f758,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))),
  inference(resolution,[],[f212,f202])).
fof(f202,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f4031,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X42))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X42,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X42))) | s(bool,X43) != s(bool,X45)) )),
  inference(resolution,[],[f3857,f134])).
fof(f4029,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X34))) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X34,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X34)))) )),
  inference(resolution,[],[f3857,f133])).
fof(f4027,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X26))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X26,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X26)))) )),
  inference(resolution,[],[f3857,f132])).
fof(f4026,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X22))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X22,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X22))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f3857,f130])).
fof(f4024,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X14)))) )),
  inference(resolution,[],[f3857,f128])).
fof(f4002,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl35),
  inference(resolution,[],[f916,f210])).
fof(f3991,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))) | s(bool,X22) != s(bool,X24)) ) | $spl35),
  inference(resolution,[],[f916,f134])).
fof(f3989,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl35),
  inference(resolution,[],[f916,f133])).
fof(f3987,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl35),
  inference(resolution,[],[f916,f132])).
fof(f3986,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))))) | s(bool,X7) != s(bool,X9)) ) | $spl35),
  inference(resolution,[],[f916,f130])).
fof(f3984,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) ) | $spl35),
  inference(resolution,[],[f916,f128])).
fof(f3957,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X56))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X56,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X56)))) )),
  inference(resolution,[],[f3444,f210])).
fof(f3444,plain,(
  ( ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X6,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X6)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X6)))) )),
  inference(subsumption_resolution,[],[f3361,f144])).
fof(f3361,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X6,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X6)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X6)))) )),
  inference(superposition,[],[f186,f2701])).
fof(f2701,plain,(
  ( ! [X71] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X71))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X71)))) )),
  inference(resolution,[],[f2489,f178])).
fof(f2489,plain,(
  ( ! [X40] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X40)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X40)))) )),
  inference(resolution,[],[f1393,f174])).
fof(f3946,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X42))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X42,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X42))) | s(bool,X43) != s(bool,X45)) )),
  inference(resolution,[],[f3444,f134])).
fof(f3944,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X34))) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X34,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X34)))) )),
  inference(resolution,[],[f3444,f133])).
fof(f3942,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X26))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X26,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X26)))) )),
  inference(resolution,[],[f3444,f132])).
fof(f3941,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X22))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X22,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X22))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f3444,f130])).
fof(f3939,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X14)))) )),
  inference(resolution,[],[f3444,f128])).
fof(f3861,plain,(
  ( ! [X23,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X23))) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X23)))) ) | $spl102),
  inference(cnf_transformation,[],[f3861_D])).
fof(f3861_D,plain,(
  ( ! [X23,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X23))) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X23)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f3917,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X44))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X44)))) )),
  inference(resolution,[],[f3850,f210])).
fof(f3850,plain,(
  ( ! [X158] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X158)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X158)))) )),
  inference(trivial_inequality_removal,[],[f3823])).
fof(f3823,plain,(
  ( ! [X158] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X158)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X158)))) )),
  inference(superposition,[],[f175,f2806])).
fof(f175,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.measure.MEASURABLE_SUBSET','thm.bool.EQ_CLAUSES|split|2')).
fof(f3906,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X30))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X30))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f3850,f134])).
fof(f3904,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X22))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X22)))) )),
  inference(resolution,[],[f3850,f133])).
fof(f3902,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X14)))) )),
  inference(resolution,[],[f3850,f132])).
fof(f3901,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X10))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X10))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f3850,f130])).
fof(f3899,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X2)))) )),
  inference(resolution,[],[f3850,f128])).
fof(f3484,plain,(
  ( ! [X198] : (s(bool,'T_0') = s(bool,X198)) ) | $spl96),
  inference(cnf_transformation,[],[f3484_D])).
fof(f3484_D,plain,(
  ( ! [X198] : (s(bool,'T_0') = s(bool,X198)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f3859,plain,(
  ( ! [X11] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X11)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X11)))) ) | $spl100),
  inference(cnf_transformation,[],[f3859_D])).
fof(f3859_D,plain,(
  ( ! [X11] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X11)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X11)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f3896,plain,(
  ( ! [X198,X197] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X197)))) | s(bool,'T_0') = s(bool,X198) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X197)))) )),
  inference(subsumption_resolution,[],[f3845,f209])).
fof(f209,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.measure.MEASURABLE_SUBSET',aHL_BOOL_CASES)).
fof(f3845,plain,(
  ( ! [X198,X197] : (s(bool,'F_0') != s(bool,X198) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X197)))) | s(bool,'T_0') = s(bool,X198) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X197)))) )),
  inference(superposition,[],[f451,f2806])).
fof(f451,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f427,f144])).
fof(f427,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,'F_0')) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f233,f209])).
fof(f233,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X26))) )),
  inference(resolution,[],[f126,f210])).
fof(f3894,plain,(
  ( ! [X187,X188] : (p(s(bool,X188)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X187)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X187)))) )),
  inference(subsumption_resolution,[],[f3840,f174])).
fof(f3840,plain,(
  ( ! [X187,X188] : (s(bool,'F_0') != s(bool,X188) | p(s(bool,X188)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X187)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X187)))) )),
  inference(superposition,[],[f233,f2806])).
fof(f3891,plain,(
  ( ! [X180,X178,X179] : (p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X178))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X178)))) )),
  inference(subsumption_resolution,[],[f3837,f160])).
fof(f3837,plain,(
  ( ! [X180,X178,X179] : (p(s(bool,'F_0')) | p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X178))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X178)))) )),
  inference(superposition,[],[f210,f2806])).
fof(f3890,plain,(
  ( ! [X173,X174,X172] : (p(s(bool,X174)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X172)))) | p(s(bool,X173)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X172)))) )),
  inference(subsumption_resolution,[],[f3835,f174])).
fof(f3835,plain,(
  ( ! [X173,X174,X172] : (s(bool,'F_0') != s(bool,X173) | p(s(bool,X174)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X172)))) | p(s(bool,X173)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X172)))) )),
  inference(superposition,[],[f210,f2806])).
fof(f3878,plain,(
  ( ! [X134,X132,X133,X131] : (~p(s(bool,X132)) | ~p(s(bool,X133)) | s(bool,X133) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X131))) | s(bool,X132) != s(bool,X134) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X131)))) )),
  inference(subsumption_resolution,[],[f3806,f144])).
fof(f3806,plain,(
  ( ! [X134,X132,X133,X131] : (p(s(bool,'F_0')) | ~p(s(bool,X132)) | ~p(s(bool,X133)) | s(bool,X133) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X131))) | s(bool,X132) != s(bool,X134) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X131)))) )),
  inference(superposition,[],[f134,f2806])).
fof(f3877,plain,(
  ( ! [X127,X130,X128,X129] : (p(s(bool,X128)) | ~p(s(bool,X129)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X127)))) | s(bool,X129) != s(bool,X130) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X127)))) )),
  inference(subsumption_resolution,[],[f3805,f174])).
fof(f3805,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'F_0') != s(bool,X128) | p(s(bool,X128)) | ~p(s(bool,X129)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X127)))) | s(bool,X129) != s(bool,X130) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X127)))) )),
  inference(superposition,[],[f134,f2806])).
fof(f3873,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,X100)) | ~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X99))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X99)))) )),
  inference(subsumption_resolution,[],[f3798,f160])).
fof(f3798,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,'F_0')) | p(s(bool,X100)) | ~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X99))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X99)))) )),
  inference(superposition,[],[f133,f2806])).
fof(f3872,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X93)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X91)))) | ~p(s(bool,X94)) | p(s(bool,X92)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X91)))) )),
  inference(subsumption_resolution,[],[f3796,f174])).
fof(f3796,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'F_0') != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X91)))) | ~p(s(bool,X94)) | p(s(bool,X92)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X91)))) )),
  inference(superposition,[],[f133,f2806])).
fof(f3869,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,X82) | s(bool,X81) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X79))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X79)))) )),
  inference(subsumption_resolution,[],[f3793,f160])).
fof(f3793,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,'F_0')) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,X82) | s(bool,X81) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X79))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X79)))) )),
  inference(superposition,[],[f132,f2806])).
fof(f3868,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X63)))) | p(s(bool,X64)) | s(bool,X65) != s(bool,X66) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X63)))) )),
  inference(subsumption_resolution,[],[f3789,f174])).
fof(f3789,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,'F_0') != s(bool,X64) | p(s(bool,X65)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X63)))) | p(s(bool,X64)) | s(bool,X65) != s(bool,X66) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X63)))) )),
  inference(superposition,[],[f132,f2806])).
fof(f3867,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X51))) | s(bool,X52) != s(bool,X54) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X51)))) )),
  inference(subsumption_resolution,[],[f3786,f144])).
fof(f3786,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,'F_0')) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X51))) | s(bool,X52) != s(bool,X54) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X51)))) )),
  inference(superposition,[],[f130,f2806])).
fof(f3866,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X43)))) | s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X43)))) )),
  inference(subsumption_resolution,[],[f3784,f174])).
fof(f3784,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'F_0') != s(bool,X44) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X43)))) | s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X43)))) )),
  inference(superposition,[],[f130,f2806])).
fof(f3860,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X23))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X23)))) )),
  inference(subsumption_resolution,[],[f3779,f160])).
fof(f3779,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,'F_0')) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X23))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X23)))) )),
  inference(superposition,[],[f128,f2806])).
fof(f3858,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X11)))) | p(s(bool,X12)) | s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X11)))) )),
  inference(subsumption_resolution,[],[f3776,f174])).
fof(f3776,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X11)))) | p(s(bool,X12)) | s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X11)))) )),
  inference(superposition,[],[f128,f2806])).
fof(f3729,plain,(
  ( ! [X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X32,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X32))) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X32)))) ) | $spl98),
  inference(cnf_transformation,[],[f3729_D])).
fof(f3729_D,plain,(
  ( ! [X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X32,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X32))) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X32)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f3721,plain,(
  ( ! [X74,X75,X73] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X73,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X73))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X73)))) )),
  inference(resolution,[],[f2748,f210])).
fof(f2748,plain,(
  ( ! [X52] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X52)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X52,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X52)))) )),
  inference(resolution,[],[f2514,f174])).
fof(f2514,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X1,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X1))))) )),
  inference(resolution,[],[f1433,f186])).
fof(f3711,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X60,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X60))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X60))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f2748,f134])).
fof(f3708,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X48,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X48))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X48)))) )),
  inference(resolution,[],[f2748,f133])).
fof(f3707,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X44,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X44))) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X44)))) )),
  inference(resolution,[],[f2748,f132])).
fof(f3706,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X40,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X40))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f2748,f130])).
fof(f3704,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X32,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X32))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X32)))) )),
  inference(resolution,[],[f2748,f128])).
fof(f3448,plain,(
  ( ! [X23,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X23))) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X23)))) ) | $spl94),
  inference(cnf_transformation,[],[f3448_D])).
fof(f3448_D,plain,(
  ( ! [X23,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X23))) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X23)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3505,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X44))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X44)))) )),
  inference(resolution,[],[f3437,f210])).
fof(f3437,plain,(
  ( ! [X158] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X158)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X158)))) )),
  inference(trivial_inequality_removal,[],[f3410])).
fof(f3410,plain,(
  ( ! [X158] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X158)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X158)))) )),
  inference(superposition,[],[f175,f2701])).
fof(f3494,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X30))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X30))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f3437,f134])).
fof(f3492,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X22))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X22)))) )),
  inference(resolution,[],[f3437,f133])).
fof(f3490,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X14)))) )),
  inference(resolution,[],[f3437,f132])).
fof(f3489,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X10))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X10))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f3437,f130])).
fof(f3487,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X2)))) )),
  inference(resolution,[],[f3437,f128])).
fof(f3446,plain,(
  ( ! [X11] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X11)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X11)))) ) | $spl92),
  inference(cnf_transformation,[],[f3446_D])).
fof(f3446_D,plain,(
  ( ! [X11] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X11)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X11)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3483,plain,(
  ( ! [X198,X197] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X197)))) | s(bool,'T_0') = s(bool,X198) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X197)))) )),
  inference(subsumption_resolution,[],[f3432,f209])).
fof(f3432,plain,(
  ( ! [X198,X197] : (s(bool,'F_0') != s(bool,X198) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X197)))) | s(bool,'T_0') = s(bool,X198) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X197)))) )),
  inference(superposition,[],[f451,f2701])).
fof(f3481,plain,(
  ( ! [X187,X188] : (p(s(bool,X188)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X187)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X187)))) )),
  inference(subsumption_resolution,[],[f3427,f174])).
fof(f3427,plain,(
  ( ! [X187,X188] : (s(bool,'F_0') != s(bool,X188) | p(s(bool,X188)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X187)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X187)))) )),
  inference(superposition,[],[f233,f2701])).
fof(f3478,plain,(
  ( ! [X180,X178,X179] : (p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X178))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X178)))) )),
  inference(subsumption_resolution,[],[f3424,f160])).
fof(f3424,plain,(
  ( ! [X180,X178,X179] : (p(s(bool,'F_0')) | p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X180) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X178))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X178)))) )),
  inference(superposition,[],[f210,f2701])).
fof(f3477,plain,(
  ( ! [X173,X174,X172] : (p(s(bool,X174)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X172)))) | p(s(bool,X173)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X172)))) )),
  inference(subsumption_resolution,[],[f3422,f174])).
fof(f3422,plain,(
  ( ! [X173,X174,X172] : (s(bool,'F_0') != s(bool,X173) | p(s(bool,X174)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X172)))) | p(s(bool,X173)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X172)))) )),
  inference(superposition,[],[f210,f2701])).
fof(f3465,plain,(
  ( ! [X134,X132,X133,X131] : (~p(s(bool,X132)) | ~p(s(bool,X133)) | s(bool,X133) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X131))) | s(bool,X132) != s(bool,X134) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X131)))) )),
  inference(subsumption_resolution,[],[f3393,f144])).
fof(f3393,plain,(
  ( ! [X134,X132,X133,X131] : (p(s(bool,'F_0')) | ~p(s(bool,X132)) | ~p(s(bool,X133)) | s(bool,X133) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X131))) | s(bool,X132) != s(bool,X134) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X131)))) )),
  inference(superposition,[],[f134,f2701])).
fof(f3464,plain,(
  ( ! [X127,X130,X128,X129] : (p(s(bool,X128)) | ~p(s(bool,X129)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X127)))) | s(bool,X129) != s(bool,X130) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X127)))) )),
  inference(subsumption_resolution,[],[f3392,f174])).
fof(f3392,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'F_0') != s(bool,X128) | p(s(bool,X128)) | ~p(s(bool,X129)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X127)))) | s(bool,X129) != s(bool,X130) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X127)))) )),
  inference(superposition,[],[f134,f2701])).
fof(f3460,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,X100)) | ~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X99))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X99)))) )),
  inference(subsumption_resolution,[],[f3385,f160])).
fof(f3385,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,'F_0')) | p(s(bool,X100)) | ~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X99))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X99)))) )),
  inference(superposition,[],[f133,f2701])).
fof(f3459,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X93)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X91)))) | ~p(s(bool,X94)) | p(s(bool,X92)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X91)))) )),
  inference(subsumption_resolution,[],[f3383,f174])).
fof(f3383,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'F_0') != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X91)))) | ~p(s(bool,X94)) | p(s(bool,X92)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X91)))) )),
  inference(superposition,[],[f133,f2701])).
fof(f3456,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,X82) | s(bool,X81) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X79))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X79)))) )),
  inference(subsumption_resolution,[],[f3380,f160])).
fof(f3380,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,'F_0')) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,X82) | s(bool,X81) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X79))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X79)))) )),
  inference(superposition,[],[f132,f2701])).
fof(f3455,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X63)))) | p(s(bool,X64)) | s(bool,X65) != s(bool,X66) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X63)))) )),
  inference(subsumption_resolution,[],[f3376,f174])).
fof(f3376,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,'F_0') != s(bool,X64) | p(s(bool,X65)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X63)))) | p(s(bool,X64)) | s(bool,X65) != s(bool,X66) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X63)))) )),
  inference(superposition,[],[f132,f2701])).
fof(f3454,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X51))) | s(bool,X52) != s(bool,X54) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X51)))) )),
  inference(subsumption_resolution,[],[f3373,f144])).
fof(f3373,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,'F_0')) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X51))) | s(bool,X52) != s(bool,X54) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X51)))) )),
  inference(superposition,[],[f130,f2701])).
fof(f3453,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X43)))) | s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X43)))) )),
  inference(subsumption_resolution,[],[f3371,f174])).
fof(f3371,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'F_0') != s(bool,X44) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X43)))) | s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X43)))) )),
  inference(superposition,[],[f130,f2701])).
fof(f3447,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X23))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X23)))) )),
  inference(subsumption_resolution,[],[f3366,f160])).
fof(f3366,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,'F_0')) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X23))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X23)))) )),
  inference(superposition,[],[f128,f2701])).
fof(f3445,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X11)))) | p(s(bool,X12)) | s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X11)))) )),
  inference(subsumption_resolution,[],[f3363,f174])).
fof(f3363,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X11)))) | p(s(bool,X12)) | s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X11)))) )),
  inference(superposition,[],[f128,f2701])).
fof(f3321,plain,(
  ( ! [X4,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X2)))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X2)))) ) | $spl90),
  inference(cnf_transformation,[],[f3321_D])).
fof(f3321_D,plain,(
  ( ! [X4,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X2)))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X2)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f3312,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X44)))) )),
  inference(resolution,[],[f761,f210])).
fof(f761,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X1))))) )),
  inference(resolution,[],[f212,f184])).
fof(f3301,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X30)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X30))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f761,f134])).
fof(f3299,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X22)))) )),
  inference(resolution,[],[f761,f133])).
fof(f3297,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X14)))) )),
  inference(resolution,[],[f761,f132])).
fof(f3296,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X10)))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X10))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f761,f130])).
fof(f3294,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),X2)))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),X2)))) )),
  inference(resolution,[],[f761,f128])).
fof(f3286,plain,(
  ( ! [X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X32,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X32))) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X32)))) ) | $spl88),
  inference(cnf_transformation,[],[f3286_D])).
fof(f3286_D,plain,(
  ( ! [X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X32,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X32))) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X32)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3278,plain,(
  ( ! [X74,X75,X73] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X73,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X73))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X73)))) )),
  inference(resolution,[],[f2643,f210])).
fof(f3268,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X60,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X60))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X60))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f2643,f134])).
fof(f3265,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X48,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X48))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X48)))) )),
  inference(resolution,[],[f2643,f133])).
fof(f3264,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X44,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X44))) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X44)))) )),
  inference(resolution,[],[f2643,f132])).
fof(f3263,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X40,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X40))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f2643,f130])).
fof(f3261,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X32,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X32))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X32)))) )),
  inference(resolution,[],[f2643,f128])).
fof(f3081,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl86),
  inference(cnf_transformation,[],[f3081_D])).
fof(f3081_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3073,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl28),
  inference(resolution,[],[f908,f210])).
fof(f3063,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X26) != s(bool,X27)) ) | $spl28),
  inference(resolution,[],[f908,f134])).
fof(f3060,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl28),
  inference(resolution,[],[f908,f133])).
fof(f3059,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl28),
  inference(resolution,[],[f908,f132])).
fof(f3058,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X11) != s(bool,X12)) ) | $spl28),
  inference(resolution,[],[f908,f130])).
fof(f3056,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl28),
  inference(resolution,[],[f908,f128])).
fof(f2886,plain,(
  $false | $spl27),
  inference(subsumption_resolution,[],[f906,f759])).
fof(f759,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))),
  inference(resolution,[],[f212,f203])).
fof(f203,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f2862,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl84),
  inference(cnf_transformation,[],[f2862_D])).
fof(f2862_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f2844,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) )),
  inference(resolution,[],[f759,f210])).
fof(f2834,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f759,f134])).
fof(f2831,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) )),
  inference(resolution,[],[f759,f133])).
fof(f2830,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) )),
  inference(resolution,[],[f759,f132])).
fof(f2829,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f759,f130])).
fof(f2827,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) )),
  inference(resolution,[],[f759,f128])).
fof(f2816,plain,(
  ( ! [X34,X32] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X32))) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X32)))) ) | $spl82),
  inference(cnf_transformation,[],[f2816_D])).
fof(f2816_D,plain,(
  ( ! [X34,X32] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X32))) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X32)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2808,plain,(
  ( ! [X74,X75,X73] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X73))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X73)))) )),
  inference(resolution,[],[f2529,f210])).
fof(f2798,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X60))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X60))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f2529,f134])).
fof(f2795,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X48))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X48)))) )),
  inference(resolution,[],[f2529,f133])).
fof(f2794,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X44))) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X44)))) )),
  inference(resolution,[],[f2529,f132])).
fof(f2793,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X40))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f2529,f130])).
fof(f2791,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X32))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X32)))) )),
  inference(resolution,[],[f2529,f128])).
fof(f2762,plain,(
  ( ! [X14,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X14)))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X14))) | ~p(s(bool,X16))) ) | $spl80),
  inference(cnf_transformation,[],[f2762_D])).
fof(f2762_D,plain,(
  ( ! [X14,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X14)))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X14))) | ~p(s(bool,X16))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2752,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X56,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X56)))) )),
  inference(resolution,[],[f2514,f210])).
fof(f2741,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X42)))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X42,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X42))) | s(bool,X43) != s(bool,X45)) )),
  inference(resolution,[],[f2514,f134])).
fof(f2739,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X34)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X34,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X34)))) )),
  inference(resolution,[],[f2514,f133])).
fof(f2737,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X26,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X26)))) )),
  inference(resolution,[],[f2514,f132])).
fof(f2736,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X22,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X22))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f2514,f130])).
fof(f2734,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK2,bool))),s(fun(fun(sK2,bool),bool),X14)))) )),
  inference(resolution,[],[f2514,f128])).
fof(f2711,plain,(
  ( ! [X34,X32] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X32))) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X32)))) ) | $spl78),
  inference(cnf_transformation,[],[f2711_D])).
fof(f2711_D,plain,(
  ( ! [X34,X32] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X32))) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X32)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2703,plain,(
  ( ! [X74,X75,X73] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X73))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X73)))) )),
  inference(resolution,[],[f2489,f210])).
fof(f2693,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X60))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X60))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f2489,f134])).
fof(f2690,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X48))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X48)))) )),
  inference(resolution,[],[f2489,f133])).
fof(f2689,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X44))) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X44)))) )),
  inference(resolution,[],[f2489,f132])).
fof(f2688,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X40))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f2489,f130])).
fof(f2686,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X32))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X32)))) )),
  inference(resolution,[],[f2489,f128])).
fof(f2657,plain,(
  ( ! [X14,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X14)))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X14))) | ~p(s(bool,X16))) ) | $spl76),
  inference(cnf_transformation,[],[f2657_D])).
fof(f2657_D,plain,(
  ( ! [X14,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X14)))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X14))) | ~p(s(bool,X16))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2647,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X56,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X56)))) )),
  inference(resolution,[],[f2474,f210])).
fof(f2636,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X42)))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X42,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X42))) | s(bool,X43) != s(bool,X45)) )),
  inference(resolution,[],[f2474,f134])).
fof(f2634,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X34)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X34,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X34)))) )),
  inference(resolution,[],[f2474,f133])).
fof(f2632,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X26,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X26)))) )),
  inference(resolution,[],[f2474,f132])).
fof(f2631,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X22,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X22))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f2474,f130])).
fof(f2629,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK5(X14,'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)),fun(sK1,bool))),s(fun(fun(sK1,bool),bool),X14)))) )),
  inference(resolution,[],[f2474,f128])).
fof(f2616,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl74),
  inference(cnf_transformation,[],[f2616_D])).
fof(f2616_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2608,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f1458,f210])).
fof(f1458,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),
  inference(subsumption_resolution,[],[f1432,f834])).
fof(f1432,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | ~sP0(sK2,sK4)),
  inference(resolution,[],[f930,f189])).
fof(f189,plain,(
  ( ! [X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X5))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f118])).
fof(f2598,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f1458,f134])).
fof(f2595,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f1458,f133])).
fof(f2594,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f1458,f132])).
fof(f2593,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f1458,f130])).
fof(f2591,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f1458,f128])).
fof(f2582,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) ) | $spl72),
  inference(cnf_transformation,[],[f2582_D])).
fof(f2582_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2574,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f1418,f210])).
fof(f1418,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))))),
  inference(subsumption_resolution,[],[f1392,f792])).
fof(f1392,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) | ~sP0(sK1,sK3)),
  inference(resolution,[],[f831,f189])).
fof(f2564,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f1418,f134])).
fof(f2561,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f1418,f133])).
fof(f2560,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f1418,f132])).
fof(f2559,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f1418,f130])).
fof(f2557,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f1418,f128])).
fof(f2541,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X2))) | ~p(s(bool,X4)) | p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X2))))) ) | $spl70),
  inference(cnf_transformation,[],[f2541_D])).
fof(f2541_D,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X2))) | ~p(s(bool,X4)) | p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X2))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2533,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X44)))) )),
  inference(resolution,[],[f1433,f210])).
fof(f2522,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X30)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X30))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f1433,f134])).
fof(f2520,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X22)))) )),
  inference(resolution,[],[f1433,f133])).
fof(f2518,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X14)))) )),
  inference(resolution,[],[f1433,f132])).
fof(f2517,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X10)))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X10))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f1433,f130])).
fof(f2515,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),X2)))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),X2)))) )),
  inference(resolution,[],[f1433,f128])).
fof(f2501,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X2))) | ~p(s(bool,X4)) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X2))))) ) | $spl68),
  inference(cnf_transformation,[],[f2501_D])).
fof(f2501_D,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X2))) | ~p(s(bool,X4)) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X2))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2493,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X44)))) )),
  inference(resolution,[],[f1393,f210])).
fof(f2482,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X30)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X30))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f1393,f134])).
fof(f2480,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X22)))) )),
  inference(resolution,[],[f1393,f133])).
fof(f2478,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X14)))) )),
  inference(resolution,[],[f1393,f132])).
fof(f2477,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X10)))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X10))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f1393,f130])).
fof(f2475,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),X2)))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),X2)))) )),
  inference(resolution,[],[f1393,f128])).
fof(f1945,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) ) | $spl66),
  inference(cnf_transformation,[],[f1945_D])).
fof(f1945_D,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1938,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f830,f210])).
fof(f830,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))))),
  inference(resolution,[],[f792,f187])).
fof(f187,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f118])).
fof(f1928,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))) | s(bool,X28) != s(bool,X29)) )),
  inference(resolution,[],[f830,f134])).
fof(f1925,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f830,f133])).
fof(f1924,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f830,f132])).
fof(f1923,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f830,f130])).
fof(f1921,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f830,f128])).
fof(f1900,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl53),
  inference(cnf_transformation,[],[f1900_D])).
fof(f1900_D,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f1902,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl55),
  inference(cnf_transformation,[],[f1902_D])).
fof(f1902_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f1908,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl61),
  inference(cnf_transformation,[],[f1908_D])).
fof(f1908_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f1912,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl65),
  inference(cnf_transformation,[],[f1912_D])).
fof(f1912_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f1897,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl33),
  inference(resolution,[],[f1352,f196])).
fof(f196,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),sK7(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK6(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f1352,plain,(
  ~sP0(sK2,'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | $spl33),
  inference(resolution,[],[f914,f198])).
fof(f198,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f119])).
fof(f1904,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl56),
  inference(cnf_transformation,[],[f1904_D])).
fof(f1904_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1896,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl33),
  inference(resolution,[],[f1352,f195])).
fof(f195,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),sK7(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK6(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f1910,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl62),
  inference(cnf_transformation,[],[f1910_D])).
fof(f1910_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1895,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl33),
  inference(resolution,[],[f1352,f194])).
fof(f194,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),sK7(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK6(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f1894,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl33),
  inference(resolution,[],[f1352,f193])).
fof(f193,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),sK7(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK6(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f1906,plain,(
  p(s(bool,'const.util_prob.countable_1'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2))))) | $spl58),
  inference(cnf_transformation,[],[f1906_D])).
fof(f1906_D,plain,(
  p(s(bool,'const.util_prob.countable_1'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1893,plain,(
  p(s(bool,'const.util_prob.countable_1'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.DIFF_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl33),
  inference(resolution,[],[f1352,f192])).
fof(f192,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),sK7(X1,X0))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK6(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f1892,plain,(
  p(s(bool,'const.util_prob.countable_1'(s(fun(fun(sK2,bool),bool),sK7('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK6('const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl33),
  inference(resolution,[],[f1352,f191])).
fof(f191,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.util_prob.countable_1'(s(fun(fun(X0,bool),bool),sK7(X1,X0))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK6(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f1379,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl44),
  inference(cnf_transformation,[],[f1379_D])).
fof(f1379_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1877,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl33),
  inference(resolution,[],[f914,f210])).
fof(f1866,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) | s(bool,X27) != s(bool,X29)) ) | $spl33),
  inference(resolution,[],[f914,f134])).
fof(f1864,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl33),
  inference(resolution,[],[f914,f133])).
fof(f1862,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl33),
  inference(resolution,[],[f914,f132])).
fof(f1861,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) | s(bool,X12) != s(bool,X14)) ) | $spl33),
  inference(resolution,[],[f914,f130])).
fof(f1859,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl33),
  inference(resolution,[],[f914,f128])).
fof(f1673,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl50),
  inference(cnf_transformation,[],[f1673_D])).
fof(f1673_D,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1666,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl32),
  inference(resolution,[],[f913,f210])).
fof(f913,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | $spl32),
  inference(cnf_transformation,[],[f913_D])).
fof(f913_D,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1656,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) | s(bool,X42) != s(bool,X43)) ) | $spl32),
  inference(resolution,[],[f913,f134])).
fof(f1653,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl32),
  inference(resolution,[],[f913,f133])).
fof(f1652,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl32),
  inference(resolution,[],[f913,f132])).
fof(f1651,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) | s(bool,X27) != s(bool,X28)) ) | $spl32),
  inference(resolution,[],[f913,f130])).
fof(f1649,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl32),
  inference(resolution,[],[f913,f128])).
fof(f1459,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl48),
  inference(cnf_transformation,[],[f1459_D])).
fof(f1459_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1451,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f930,f210])).
fof(f1441,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f930,f134])).
fof(f1438,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f930,f133])).
fof(f1437,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f930,f132])).
fof(f1436,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f930,f130])).
fof(f1434,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f930,f128])).
fof(f1419,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) ) | $spl46),
  inference(cnf_transformation,[],[f1419_D])).
fof(f1419_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1411,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f831,f210])).
fof(f1401,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f831,f134])).
fof(f1398,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f831,f133])).
fof(f1397,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f831,f132])).
fof(f1396,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f831,f130])).
fof(f1394,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f831,f128])).
fof(f1371,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl33),
  inference(resolution,[],[f914,f210])).
fof(f1360,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) | s(bool,X27) != s(bool,X29)) ) | $spl33),
  inference(resolution,[],[f914,f134])).
fof(f1358,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl33),
  inference(resolution,[],[f914,f133])).
fof(f1356,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl33),
  inference(resolution,[],[f914,f132])).
fof(f1355,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))) | s(bool,X12) != s(bool,X14)) ) | $spl33),
  inference(resolution,[],[f914,f130])).
fof(f1353,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) ) | $spl33),
  inference(resolution,[],[f914,f128])).
fof(f1342,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl42),
  inference(cnf_transformation,[],[f1342_D])).
fof(f1342_D,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1335,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl24),
  inference(resolution,[],[f903,f210])).
fof(f903,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | $spl24),
  inference(cnf_transformation,[],[f903_D])).
fof(f903_D,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1325,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X28) != s(bool,X29)) ) | $spl24),
  inference(resolution,[],[f903,f134])).
fof(f1322,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl24),
  inference(resolution,[],[f903,f133])).
fof(f1321,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl24),
  inference(resolution,[],[f903,f132])).
fof(f1320,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X13) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f903,f130])).
fof(f1318,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl24),
  inference(resolution,[],[f903,f128])).
fof(f1132,plain,(
  $false | $spl25),
  inference(subsumption_resolution,[],[f1105,f834])).
fof(f1105,plain,(
  ~sP0(sK2,sK4) | $spl25),
  inference(resolution,[],[f904,f187])).
fof(f922,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) ) | $spl40),
  inference(cnf_transformation,[],[f922_D])).
fof(f922_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f894,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f213,f210])).
fof(f883,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f213,f134])).
fof(f881,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f213,f133])).
fof(f879,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f213,f132])).
fof(f878,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f213,f130])).
fof(f876,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f213,f128])).
fof(f914,plain,(
  ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | $spl33),
  inference(cnf_transformation,[],[f914_D])).
fof(f914_D,plain,(
  ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f916,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) | $spl35),
  inference(cnf_transformation,[],[f916_D])).
fof(f916_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f921,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),sK9(sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK3,sK1,sK2)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) | $spl39),
  inference(cnf_transformation,[],[f921_D])).
fof(f921_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),sK9(sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK3,sK1,sK2)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f919,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),sK9(sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK3,sK1,sK2)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))),
  inference(subsumption_resolution,[],[f874,f757])).
fof(f874,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),sK9(sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK3,sK1,sK2)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),
  inference(resolution,[],[f213,f206])).
fof(f206,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),sK9(X4,X3,X2,X1,X0)))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f918,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK9(sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK3,sK1,sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | $spl36),
  inference(cnf_transformation,[],[f918_D])).
fof(f918_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK9(sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK3,sK1,sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f912,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK9(sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK3,sK1,sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))),
  inference(subsumption_resolution,[],[f873,f757])).
fof(f873,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK9(sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))),sK3,sK1,sK2)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),
  inference(resolution,[],[f213,f205])).
fof(f205,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK9(X4,X3,X2,X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f904,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | $spl25),
  inference(cnf_transformation,[],[f904_D])).
fof(f904_D,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f906,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | $spl27),
  inference(cnf_transformation,[],[f906_D])).
fof(f906_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f911,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) | $spl31),
  inference(cnf_transformation,[],[f911_D])).
fof(f911_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f909,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),
  inference(subsumption_resolution,[],[f872,f757])).
fof(f872,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),'const.pred_set.INTER_2'(s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),
  inference(resolution,[],[f213,f200])).
fof(f200,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),sK8(X5,X4,X3,X2,X1,X0)))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))),s(fun(X1,bool),X5)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))),s(fun(X1,bool),X5)))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),sK8(X5,X4,X3,X2,X1,X0)),s(fun(fun(X1,bool),bool),X4)))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),sK8(X5,X4,X3,X2,X1,X0)))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))),s(fun(X1,bool),X5)))))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),X6),s(fun(fun(X1,bool),bool),X4)))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X6))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))))))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))),s(fun(X1,bool),X5)))))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),X6),s(fun(fun(X1,bool),bool),X4)))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X6))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))))))))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.measure.subset_class_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))),s(fun(X1,bool),X5)))))) & ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),X6),s(fun(fun(X1,bool),bool),X4)))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X6))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))))))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : ((p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.measure.subset_class_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))),s(fun(X1,bool),X5)))))) & ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),X6),s(fun(fun(X1,bool),bool),X4)))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X6))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7] : ! [X21,X8,X9,X22] : ((p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))) & p(s(bool,'const.measure.subset_class_2'(s(fun(X7,bool),X22),s(fun(fun(X7,bool),bool),X9)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X21),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8))),s(fun(X7,bool),X22)))))) & ! [X19] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,bool),X19),s(fun(fun(X7,bool),bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X21),s(fun(X7,bool),X19))),s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X21),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),'const.measure.sigma_2'(s(fun(X7,bool),X22),s(fun(fun(X7,bool),bool),X9)))))))))),
  file('thm.measure.MEASURABLE_SUBSET','thm.measure.MEASURABLE_SIGMA')).
fof(f908,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | $spl28),
  inference(cnf_transformation,[],[f908_D])).
fof(f908_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f902,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))),
  inference(subsumption_resolution,[],[f871,f757])).
fof(f871,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),sK8('const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),sK3,sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2)),sK2,sK1)),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.util_prob.FUNSET_2'(s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))),s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))))),
  inference(resolution,[],[f213,f199])).
fof(f199,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),'const.measure.sigma_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),sK8(X5,X4,X3,X2,X1,X0)),s(fun(fun(X1,bool),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))),s(fun(X1,bool),X5)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X1,bool),X5),s(fun(fun(X1,bool),bool),X4)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f865,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))) ) | $spl22),
  inference(cnf_transformation,[],[f865_D])).
fof(f865_D,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f858,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))) )),
  inference(resolution,[],[f758,f210])).
fof(f848,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f758,f134])).
fof(f845,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))) )),
  inference(resolution,[],[f758,f133])).
fof(f844,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))) )),
  inference(resolution,[],[f758,f132])).
fof(f843,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f758,f130])).
fof(f841,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))) )),
  inference(resolution,[],[f758,f128])).
fof(f823,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))) ) | $spl20),
  inference(cnf_transformation,[],[f823_D])).
fof(f823_D,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f816,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))) )),
  inference(resolution,[],[f757,f210])).
fof(f806,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f757,f134])).
fof(f803,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))) )),
  inference(resolution,[],[f757,f133])).
fof(f802,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))) )),
  inference(resolution,[],[f757,f132])).
fof(f801,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f757,f130])).
fof(f799,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3)))) )),
  inference(resolution,[],[f757,f128])).
fof(f786,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) | $spl18),
  inference(cnf_transformation,[],[f786_D])).
fof(f786_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f779,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f212,f210])).
fof(f769,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f212,f134])).
fof(f766,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f212,f133])).
fof(f765,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f212,f132])).
fof(f764,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))) | s(bool,X9) != s(bool,X10)) )),
  inference(resolution,[],[f212,f130])).
fof(f762,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(fun(sK1,sK2),sK5('const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)),fun(sK1,sK2))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))) )),
  inference(resolution,[],[f212,f128])).
fof(f748,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f739])).
fof(f739,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(superposition,[],[f242,f228])).
fof(f228,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))),
  inference(resolution,[],[f126,f174])).
fof(f747,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f242])).
fof(f378,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) | $spl14),
  inference(cnf_transformation,[],[f378_D])).
fof(f378_D,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f408,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f389])).
fof(f389,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12)) | ~p(s(bool,X12))) ) | $spl8),
  inference(superposition,[],[f244,f178])).
fof(f409,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f388])).
fof(f388,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | $spl8),
  inference(superposition,[],[f244,f180])).
fof(f180,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.measure.MEASURABLE_SUBSET','thm.bool.EQ_CLAUSES|split|0')).
fof(f380,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) | $spl16),
  inference(cnf_transformation,[],[f380_D])).
fof(f380_D,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f413,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(subsumption_resolution,[],[f387,f244])).
fof(f387,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f244,f209])).
fof(f412,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | $spl8),
  inference(subsumption_resolution,[],[f386,f244])).
fof(f386,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f244,f209])).
fof(f374,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f357])).
fof(f357,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16)) | p(s(bool,X16))) ) | $spl2),
  inference(superposition,[],[f241,f174])).
fof(f375,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f356])).
fof(f356,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14)) | p(s(bool,X14))) ) | $spl2),
  inference(superposition,[],[f241,f176])).
fof(f176,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f109])).
fof(f379,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f353,f241])).
fof(f353,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl2),
  inference(superposition,[],[f241,f209])).
fof(f377,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl2),
  inference(subsumption_resolution,[],[f352,f241])).
fof(f352,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl2),
  inference(superposition,[],[f241,f209])).
fof(f336,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f276,f246])).
fof(f276,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl12),
  inference(resolution,[],[f246,f144])).
fof(f272,plain,(
  $false | $spl12),
  inference(resolution,[],[f246,f181])).
fof(f181,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.measure.MEASURABLE_SUBSET',aHL_FALSITY)).
fof(f279,plain,(
  $false | $spl12),
  inference(resolution,[],[f246,f126])).
fof(f240,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) ) | $spl0),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f232,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f126,f210])).
fof(f221,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f126,f134])).
fof(f219,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f126,f133])).
fof(f242,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) ) | $spl4),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f243,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) | $spl6),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f218,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f126,f132])).
fof(f217,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f126,f132])).
fof(f216,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f126,f130])).
fof(f215,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))))))))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f126,f128])).
fof(f214,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK1,sK2),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK3),s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),'const.measure.sigma_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK4)))))))))) )),
  inference(resolution,[],[f126,f128])).
% SZS output end Proof for thm.measure
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 140
% Passive clauses: 495
% Generated clauses: 4437
% Final active clauses: 110
% Final passive clauses: 257
% Input formulas: 28
% Initial clauses: 84
% 
% Selected by SInE selection: 27
% SInE iterations: 1
% 
% Duplicate literals: 62
% Trivial inequalities: 99
% Fw subsumption resolutions: 1001
% Fw demodulations: 327
% Bw demodulations: 21
% 
% Simple tautologies: 156
% Forward subsumptions: 2139
% 
% Binary resolution: 982
% Forward superposition: 626
% Backward superposition: 1260
% Equality resolution: 9
% 
% Split clauses: 258
% Split components: 583
% SAT solver clauses: 271
% SAT solver unit clauses: 7
% SAT solver binary clauses: 154
% 
% Sat splits: 258
% Sat splitting refutations: 8
% 
% Memory used [KB]: 3965
% Time elapsed: 0.118 s
% ------------------------------
% ------------------------------
% Success in time 11.831 s
