Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (66874)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 90
% Passive clauses: 3167
% Generated clauses: 16703
% Final active clauses: 77
% Final passive clauses: 2926
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 1031
% Trivial inequalities: 195
% Fw subsumption resolutions: 2003
% Fw demodulations: 1115
% Bw demodulations: 9
% 
% Simple tautologies: 322
% Equational tautologies: 97
% Forward subsumptions: 6713
% 
% Binary resolution: 75
% Forward superposition: 7220
% Backward superposition: 4249
% Self superposition: 348
% Equality factoring: 261
% Equality resolution: 14
% 
% Split clauses: 900
% Split components: 1812
% SAT solver clauses: 2133
% SAT solver unit clauses: 280
% SAT solver binary clauses: 654
% 
% Sat splits: 900
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2125
% TWLsolver calls for satisfiability: 105
% 
% Memory used [KB]: 7931
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (66886)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 48
% Passive clauses: 359
% Generated clauses: 7003
% Final active clauses: 47
% Final passive clauses: 305
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 1047
% Trivial inequalities: 73
% Fw subsumption resolutions: 379
% Fw demodulations: 84
% Bw demodulations: 5
% Condensations: 1567
% 
% Simple tautologies: 148
% Equational tautologies: 67
% Forward subsumptions: 2471
% 
% Binary resolution: 38
% Forward superposition: 2089
% Backward superposition: 1597
% Self superposition: 216
% Equality factoring: 198
% Equality resolution: 8
% 
% Split clauses: 81
% Split components: 163
% SAT solver clauses: 789
% SAT solver unit clauses: 13
% SAT solver binary clauses: 72
% 
% Sat splits: 81
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 789
% TWLsolver calls for satisfiability: 85
% 
% Memory used [KB]: 5884
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (66908)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 671
% Passive clauses: 5649
% Generated clauses: 163469
% Final active clauses: 120
% Final passive clauses: 644
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 11913
% Trivial inequalities: 757
% Fw subsumption resolutions: 8569
% Fw demodulations: 26843
% Bw demodulations: 3753
% Condensations: 16682
% Global subsumptions: 269
% 
% Simple tautologies: 886
% Equational tautologies: 3670
% Forward subsumptions: 87071
% Fw demodulations to eq. taut.: 48
% Bw demodulations to eq. taut.: 105
% 
% Unit resulting resolution: 17162
% Forward superposition: 44377
% Backward superposition: 25589
% Self superposition: 774
% Equality factoring: 384
% Equality resolution: 59
% 
% Split clauses: 421
% Split components: 843
% SAT solver clauses: 7206
% SAT solver unit clauses: 28
% SAT solver binary clauses: 318
% 
% Sat splits: 421
% Sat splitting refutations: 444
% 
% Memory used [KB]: 65500
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f23754,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f68,f71,f69,f146,f150,f148,f183,f197,f279,f283,f281,f316,f328,f642,f646,f644,f702,f725,f2941,f2946,f2951,f2959,f2945,f2943,f2960,f2965,f2969,f2967,f2970,f2950,f2948,f2971,f2955,f2953,f2972,f2976,f2974,f2977,f2981,f2979,f3088,f5555,f5560,f5565,f5573,f5559,f5557,f5574,f5578,f5576,f5579,f5583,f5581,f5584,f5564,f5562,f5585,f5569,f5567,f5586,f5590,f5588,f5591,f5595,f5593,f5522,f5541,f5543,f5548,f5550,f724,f5606,f5604,f3011,f643,f195,f5758,f7254,f7259,f7257,f7255,f7263,f7261,f7282,f7532,f7535,f7531,f7539,f7530,f7543,f8903,f8942,f8992,f9020,f9021,f9022,f9042,f9023,f9043,f9024,f9044,f9025,f9046,f9026,f9048,f9028,f2978,f10181,f8905,f10179,f10178,f8994,f10212,f10233,f10271,f10292,f10322,f10343,f10823,f10844,f10878,f10899,f10933,f10954,f11100,f11105,f11110,f11118,f11104,f11102,f11119,f11123,f11121,f11124,f11129,f11109,f11107,f11130,f11114,f11112,f11131,f11135,f11133,f11136,f11140,f11138,f11164,f7537,f11166,f11194,f11213,f10231,f11266,f7541,f11268,f11304,f7545,f11306,f11338,f11357,f10290,f11398,f11417,f10341,f11490,f8907,f11489,f10190,f11492,f11530,f8946,f11529,f10194,f11532,f11570,f8996,f11569,f10196,f11610,f9038,f11609,f9040,f11612,f11646,f11665,f9050,f11825,f11844,f10842,f11878,f11897,f10897,f11931,f11950,f10952,f13250,f13252,f13251,f15194,f15196,f15195,f17697,f17700,f17455,f17704,f17707,f17710,f17712,f2964,f2962,f17729,f10787,f17708,f17705,f17702,f17454,f8944,f10898,f17871,f17546,f17875,f17878,f17881,f11128,f11126,f17899,f11792,f17879,f17876,f17873,f17545,f11572,f11949,f20368,f20367,f20366,f20433,f20431,f20365,f20434,f20438,f20436,f22011,f22015,f22013,f22016,f22018,f22021,f22025,f22023,f22026,f22028,f22070,f22069,f22068,f22067,f22066,f22065,f22081,f22064,f22082,f22020,f22274,f22273,f22272,f22271,f22270,f22323,f22368,f22072,f22367,f22074,f22366,f22076,f22365,f22078,f22364,f22080,f22030,f22884,f22889,f22887,f22885,f23070,f23075,f23073,f23071,f23079,f23077,f23097,f23133,f23137,f23135,f23300,f23304,f23347,f23532,f23548,f22893,f22891,f23603,f23608,f23723,f23753])).
fof(f23753,plain,(
  $false | ($spl270 | $spl272 | $spl284)),
  inference(subsumption_resolution,[],[f23737,f23604])).
fof(f23604,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl270),
  inference(subsumption_resolution,[],[f23585,f44])).
fof(f44,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INTER_COMM',aHL_TRUTH)).
fof(f23585,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl270),
  inference(superposition,[],[f36,f22891])).
fof(f36,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7,X8,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INTER_COMM','thm.pred_set.IN_INTER')).
fof(f23737,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ($spl272 | $spl284)),
  inference(subsumption_resolution,[],[f23736,f44])).
fof(f23736,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ($spl272 | $spl284)),
  inference(forward_demodulation,[],[f23719,f23635])).
fof(f23635,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl284),
  inference(subsumption_resolution,[],[f23634,f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INTER_COMM',aHL_FALSITY)).
fof(f23634,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl284),
  inference(superposition,[],[f23136,f35])).
fof(f35,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INTER_COMM',aHL_BOOL_CASES)).
fof(f23136,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl284),
  inference(cnf_transformation,[],[f23136_D])).
fof(f23136_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f23719,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl272),
  inference(subsumption_resolution,[],[f23703,f43])).
fof(f23703,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl272),
  inference(superposition,[],[f38,f22892])).
fof(f22892,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl272),
  inference(cnf_transformation,[],[f22892_D])).
fof(f22892_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f38,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f23723,plain,(
  $false | ($spl272 | $spl282 | $spl284)),
  inference(subsumption_resolution,[],[f23722,f44])).
fof(f23722,plain,(
  ~p(s(bool,'T_0')) | ($spl272 | $spl282 | $spl284)),
  inference(forward_demodulation,[],[f23721,f23635])).
fof(f23721,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ($spl272 | $spl282)),
  inference(subsumption_resolution,[],[f23720,f44])).
fof(f23720,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ($spl272 | $spl282)),
  inference(forward_demodulation,[],[f23719,f23331])).
fof(f23331,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl282),
  inference(subsumption_resolution,[],[f23330,f43])).
fof(f23330,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl282),
  inference(superposition,[],[f23134,f35])).
fof(f23134,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl282),
  inference(cnf_transformation,[],[f23134_D])).
fof(f23134_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f23608,plain,(
  $false | ($spl270 | $spl285)),
  inference(subsumption_resolution,[],[f23607,f43])).
fof(f23607,plain,(
  p(s(bool,'F_0')) | ($spl270 | $spl285)),
  inference(forward_demodulation,[],[f23606,f23346])).
fof(f23346,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl285),
  inference(subsumption_resolution,[],[f23344,f44])).
fof(f23344,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl285),
  inference(superposition,[],[f23137,f35])).
fof(f23606,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl270),
  inference(subsumption_resolution,[],[f23586,f44])).
fof(f23586,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl270),
  inference(superposition,[],[f37,f22891])).
fof(f37,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f23603,plain,(
  $false | ($spl270 | $spl285)),
  inference(subsumption_resolution,[],[f23583,f44])).
fof(f23583,plain,(
  ~p(s(bool,'T_0')) | ($spl270 | $spl285)),
  inference(superposition,[],[f23333,f22891])).
fof(f23333,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))))))) ) | $spl285),
  inference(resolution,[],[f23137,f37])).
fof(f22891,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl270),
  inference(cnf_transformation,[],[f22891_D])).
fof(f22891_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f22893,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl273),
  inference(cnf_transformation,[],[f22893_D])).
fof(f22893_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f23548,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f491,f35])).
fof(f491,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(equality_resolution,[],[f55])).
fof(f55,plain,(
  ( ! [X2] : (s(fun(sK0,bool),X2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),X2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f31,f34])).
fof(f34,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))),
  file('thm.pred_set.INTER_COMM','thm.pred_set.EXTENSION')).
fof(f31,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f20])).
fof(f20,plain,(
  ? [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ~! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ~! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))),
  file('thm.pred_set.INTER_COMM','thm.pred_set.INTER_COMM')).
fof(f23532,plain,(
  $false | ($spl266 | $spl269)),
  inference(subsumption_resolution,[],[f23531,f23098])).
fof(f23098,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl266),
  inference(backward_demodulation,[],[f22887,f491])).
fof(f23531,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl269),
  inference(trivial_inequality_removal,[],[f23529])).
fof(f23529,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl269),
  inference(superposition,[],[f22889,f35])).
fof(f23347,plain,(
  $false | ($spl268 | $spl285)),
  inference(subsumption_resolution,[],[f23301,f23137])).
fof(f23301,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl268),
  inference(subsumption_resolution,[],[f23282,f44])).
fof(f23282,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl268),
  inference(superposition,[],[f36,f22888])).
fof(f22888,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl268),
  inference(cnf_transformation,[],[f22888_D])).
fof(f22888_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f23304,plain,(
  $false | ($spl268 | $spl283)),
  inference(subsumption_resolution,[],[f23303,f43])).
fof(f23303,plain,(
  p(s(bool,'F_0')) | ($spl268 | $spl283)),
  inference(forward_demodulation,[],[f23302,f23164])).
fof(f23164,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl283),
  inference(subsumption_resolution,[],[f23162,f44])).
fof(f23162,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl283),
  inference(superposition,[],[f23135,f35])).
fof(f23302,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl268),
  inference(subsumption_resolution,[],[f23283,f44])).
fof(f23283,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl268),
  inference(superposition,[],[f37,f22888])).
fof(f23300,plain,(
  $false | ($spl268 | $spl283)),
  inference(subsumption_resolution,[],[f23281,f44])).
fof(f23281,plain,(
  ~p(s(bool,'T_0')) | ($spl268 | $spl283)),
  inference(superposition,[],[f23151,f22888])).
fof(f23151,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))))))) ) | $spl283),
  inference(resolution,[],[f23135,f37])).
fof(f23135,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl283),
  inference(cnf_transformation,[],[f23135_D])).
fof(f23135_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f23137,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl285),
  inference(cnf_transformation,[],[f23137_D])).
fof(f23137_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f23133,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl266),
  inference(subsumption_resolution,[],[f23117,f43])).
fof(f23117,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl266),
  inference(superposition,[],[f38,f22887])).
fof(f23097,plain,(
  $false | ($spl269 | $spl273)),
  inference(subsumption_resolution,[],[f23096,f22889])).
fof(f23096,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl273),
  inference(trivial_inequality_removal,[],[f23095])).
fof(f23095,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl273),
  inference(superposition,[],[f22893,f35])).
fof(f23077,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl278),
  inference(cnf_transformation,[],[f23077_D])).
fof(f23077_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f23079,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl281),
  inference(cnf_transformation,[],[f23079_D])).
fof(f23079_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl281),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl281])])).
fof(f23071,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f553,f35])).
fof(f553,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(equality_resolution,[],[f56])).
fof(f56,plain,(
  ( ! [X3] : (s(fun(sK0,bool),X3) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f31,f34])).
fof(f23073,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl274),
  inference(cnf_transformation,[],[f23073_D])).
fof(f23073_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f23075,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl277),
  inference(cnf_transformation,[],[f23075_D])).
fof(f23075_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f23070,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f553,f35])).
fof(f22885,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f491,f35])).
fof(f22887,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) | $spl266),
  inference(cnf_transformation,[],[f22887_D])).
fof(f22887_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f22889,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl269),
  inference(cnf_transformation,[],[f22889_D])).
fof(f22889_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f22884,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f491,f35])).
fof(f22030,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl250),
  inference(cnf_transformation,[],[f22030_D])).
fof(f22030_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f22080,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl260),
  inference(cnf_transformation,[],[f22080_D])).
fof(f22080_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f22364,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl90 | $spl244)),
  inference(trivial_inequality_removal,[],[f22352])).
fof(f22352,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl90 | $spl244)),
  inference(superposition,[],[f7972,f22023])).
fof(f7972,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f7324,f35])).
fof(f7324,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1)))) ) | $spl90),
  inference(superposition,[],[f7262,f32])).
fof(f32,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f12])).
fof(f12,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.INTER_COMM',aHL_EXT)).
fof(f7262,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl90),
  inference(cnf_transformation,[],[f7262_D])).
fof(f7262_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f22078,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl259),
  inference(cnf_transformation,[],[f22078_D])).
fof(f22078_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f22365,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl87 | $spl244)),
  inference(trivial_inequality_removal,[],[f22350])).
fof(f22350,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl87 | $spl244)),
  inference(superposition,[],[f7960,f22023])).
fof(f7960,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl87),
  inference(superposition,[],[f7265,f35])).
fof(f7265,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1)))) ) | $spl87),
  inference(superposition,[],[f7259,f32])).
fof(f22076,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl257),
  inference(cnf_transformation,[],[f22076_D])).
fof(f22076_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f22366,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl244)),
  inference(trivial_inequality_removal,[],[f22348])).
fof(f22348,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl244)),
  inference(superposition,[],[f6722,f22023])).
fof(f6722,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f5934,f35])).
fof(f5934,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1)))) ) | $spl18),
  inference(superposition,[],[f718,f32])).
fof(f718,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl18),
  inference(trivial_inequality_removal,[],[f701])).
fof(f701,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | $spl18),
  inference(superposition,[],[f52,f674])).
fof(f674,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl18),
  inference(trivial_inequality_removal,[],[f656])).
fof(f656,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl18),
  inference(superposition,[],[f261,f644])).
fof(f261,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X5))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK1,sK0)),s(fun(sK0,bool),X5)))) )),
  inference(superposition,[],[f52,f35])).
fof(f52,plain,(
  ( ! [X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK1,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X3)))) )),
  inference(superposition,[],[f31,f34])).
fof(f22074,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl254),
  inference(cnf_transformation,[],[f22074_D])).
fof(f22074_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f22367,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl244)),
  inference(trivial_inequality_removal,[],[f22346])).
fof(f22346,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl244)),
  inference(superposition,[],[f6704,f22023])).
fof(f6704,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f5850,f35])).
fof(f5850,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1)))) ) | $spl18),
  inference(superposition,[],[f644,f32])).
fof(f22072,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl253),
  inference(cnf_transformation,[],[f22072_D])).
fof(f22072_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f22368,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl244),
  inference(trivial_inequality_removal,[],[f22344])).
fof(f22344,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl244),
  inference(superposition,[],[f64,f22023])).
fof(f64,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0))))) )),
  inference(superposition,[],[f50,f35])).
fof(f50,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1)))) )),
  inference(superposition,[],[f31,f32])).
fof(f22323,plain,(
  $false | ($spl3 | $spl244 | $spl248)),
  inference(subsumption_resolution,[],[f22285,f71])).
fof(f22285,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl244 | $spl248)),
  inference(backward_demodulation,[],[f22028,f22023])).
fof(f22270,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl90 | $spl244)),
  inference(trivial_inequality_removal,[],[f22258])).
fof(f22258,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl90 | $spl244)),
  inference(superposition,[],[f7972,f22023])).
fof(f22271,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl87 | $spl244)),
  inference(trivial_inequality_removal,[],[f22256])).
fof(f22256,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl87 | $spl244)),
  inference(superposition,[],[f7960,f22023])).
fof(f22272,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl244)),
  inference(trivial_inequality_removal,[],[f22254])).
fof(f22254,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl244)),
  inference(superposition,[],[f6722,f22023])).
fof(f22273,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl244)),
  inference(trivial_inequality_removal,[],[f22252])).
fof(f22252,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl244)),
  inference(superposition,[],[f6704,f22023])).
fof(f22274,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl244),
  inference(trivial_inequality_removal,[],[f22250])).
fof(f22250,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl244),
  inference(superposition,[],[f64,f22023])).
fof(f22020,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl242),
  inference(cnf_transformation,[],[f22020_D])).
fof(f22020_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f22082,plain,(
  ( ! [X3] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) ) | $spl264),
  inference(cnf_transformation,[],[f22082_D])).
fof(f22082_D,plain,(
  ( ! [X3] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f22064,plain,(
  ( ! [X3] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0))))) ) | ($spl18 | $spl236)),
  inference(trivial_inequality_removal,[],[f22054])).
fof(f22054,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0))))) ) | ($spl18 | $spl236)),
  inference(superposition,[],[f8525,f22013])).
fof(f8525,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,X3,bool,sK0)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),X3))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X3,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6453,f35])).
fof(f6453,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),X0)))) ) | $spl18),
  inference(superposition,[],[f6447,f32])).
fof(f6447,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6440,f44])).
fof(f6440,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1)))) ) | $spl18),
  inference(superposition,[],[f6007,f674])).
fof(f6007,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1)))) ) | $spl18),
  inference(superposition,[],[f5872,f33])).
fof(f33,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) != s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f26])).
fof(f5872,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5857,f43])).
fof(f5857,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))))))) ) | $spl18),
  inference(superposition,[],[f37,f644])).
fof(f22081,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),X1)))) ) | $spl262),
  inference(cnf_transformation,[],[f22081_D])).
fof(f22081_D,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),X1)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f22065,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0))))) ) | ($spl18 | $spl236)),
  inference(trivial_inequality_removal,[],[f22052])).
fof(f22052,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0))))) ) | ($spl18 | $spl236)),
  inference(superposition,[],[f8519,f22013])).
fof(f8519,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,X3,bool,sK0)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),X4))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X3,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6401,f35])).
fof(f6401,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),X1)))) ) | $spl18),
  inference(superposition,[],[f6399,f32])).
fof(f6399,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6392,f44])).
fof(f6392,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))) ) | $spl18),
  inference(superposition,[],[f5982,f674])).
fof(f5982,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0)))) ) | $spl18),
  inference(superposition,[],[f5871,f33])).
fof(f5871,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5856,f43])).
fof(f5856,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))))) ) | $spl18),
  inference(superposition,[],[f36,f644])).
fof(f22066,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl90 | $spl236)),
  inference(trivial_inequality_removal,[],[f22050])).
fof(f22050,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl90 | $spl236)),
  inference(superposition,[],[f7969,f22013])).
fof(f7969,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f7323,f35])).
fof(f7323,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X0)))) ) | $spl90),
  inference(superposition,[],[f7262,f32])).
fof(f22067,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl87 | $spl236)),
  inference(trivial_inequality_removal,[],[f22048])).
fof(f22048,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl87 | $spl236)),
  inference(superposition,[],[f7957,f22013])).
fof(f7957,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl87),
  inference(superposition,[],[f7264,f35])).
fof(f7264,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X0)))) ) | $spl87),
  inference(superposition,[],[f7259,f32])).
fof(f22068,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl18 | $spl236)),
  inference(trivial_inequality_removal,[],[f22046])).
fof(f22046,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl18 | $spl236)),
  inference(superposition,[],[f6719,f22013])).
fof(f6719,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f5933,f35])).
fof(f5933,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0)))) ) | $spl18),
  inference(superposition,[],[f718,f32])).
fof(f22069,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl18 | $spl236)),
  inference(trivial_inequality_removal,[],[f22044])).
fof(f22044,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl18 | $spl236)),
  inference(superposition,[],[f6701,f22013])).
fof(f6701,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f5849,f35])).
fof(f5849,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0)))) ) | $spl18),
  inference(superposition,[],[f644,f32])).
fof(f22070,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl236),
  inference(trivial_inequality_removal,[],[f22042])).
fof(f22042,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl236),
  inference(superposition,[],[f62,f22013])).
fof(f62,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0))))) )),
  inference(superposition,[],[f49,f35])).
fof(f49,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0)))) )),
  inference(superposition,[],[f31,f32])).
fof(f22028,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl248),
  inference(cnf_transformation,[],[f22028_D])).
fof(f22028_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f22026,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f21996,f43])).
fof(f21996,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f21829,f10397])).
fof(f10397,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f10396])).
fof(f10396,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6704,f35])).
fof(f21829,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl18),
  inference(subsumption_resolution,[],[f21810,f44])).
fof(f21810,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl18),
  inference(superposition,[],[f6865,f6881])).
fof(f6881,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(subsumption_resolution,[],[f6878,f43])).
fof(f6878,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(superposition,[],[f6858,f35])).
fof(f6858,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(subsumption_resolution,[],[f6847,f44])).
fof(f6847,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(superposition,[],[f5917,f674])).
fof(f5917,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5903,f44])).
fof(f5903,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X2))))) ) | $spl18),
  inference(superposition,[],[f38,f674])).
fof(f6865,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X2))))))) ) | $spl18),
  inference(resolution,[],[f6858,f38])).
fof(f22023,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl244),
  inference(cnf_transformation,[],[f22023_D])).
fof(f22023_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f22025,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl246),
  inference(cnf_transformation,[],[f22025_D])).
fof(f22025_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f22021,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f21995,f43])).
fof(f21995,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f21829,f10393])).
fof(f10393,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f10391])).
fof(f10391,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6703,f35])).
fof(f6703,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f5850,f35])).
fof(f22018,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl240),
  inference(cnf_transformation,[],[f22018_D])).
fof(f22018_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f22016,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f21994,f43])).
fof(f21994,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl18),
  inference(superposition,[],[f21829,f10177])).
fof(f10177,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f10176])).
fof(f10176,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6701,f35])).
fof(f22013,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl236),
  inference(cnf_transformation,[],[f22013_D])).
fof(f22013_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f22015,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl238),
  inference(cnf_transformation,[],[f22015_D])).
fof(f22015_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f22011,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f21993,f43])).
fof(f21993,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl18),
  inference(superposition,[],[f21829,f10165])).
fof(f10165,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f10163])).
fof(f10163,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6700,f35])).
fof(f6700,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f5849,f35])).
fof(f20436,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0)))) | $spl232),
  inference(cnf_transformation,[],[f20436_D])).
fof(f20436_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0)))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f20438,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | $spl234),
  inference(cnf_transformation,[],[f20438_D])).
fof(f20438_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f20434,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X5))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),X5)))) ) | $spl230),
  inference(cnf_transformation,[],[f20434_D])).
fof(f20434_D,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X5))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),X5)))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f20365,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X8))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X8,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X8,sK0)),s(fun(sK0,bool),X8))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0))))) ) | $spl90),
  inference(trivial_inequality_removal,[],[f20359])).
fof(f20359,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X8))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X8,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X8,sK0)),s(fun(sK0,bool),X8))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f109,f13717])).
fof(f13717,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl90),
  inference(trivial_inequality_removal,[],[f13716])).
fof(f13716,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f7972,f35])).
fof(f109,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK2,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK2,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X5))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X4,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X4,sK0)),s(fun(sK0,bool),X4)))) )),
  inference(superposition,[],[f48,f34])).
fof(f48,plain,(
  ( ! [X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK2,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X3)))) )),
  inference(superposition,[],[f31,f34])).
fof(f20431,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0)))) | $spl226),
  inference(cnf_transformation,[],[f20431_D])).
fof(f20431_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0)))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f20433,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | $spl228),
  inference(cnf_transformation,[],[f20433_D])).
fof(f20433_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f20366,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X7))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X7,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X7,sK0)),s(fun(sK0,bool),X7))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0))))) ) | $spl90),
  inference(trivial_inequality_removal,[],[f20358])).
fof(f20358,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X7))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X7,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X7,sK0)),s(fun(sK0,bool),X7))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f109,f13706])).
fof(f13706,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl90),
  inference(trivial_inequality_removal,[],[f13704])).
fof(f13704,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f7971,f35])).
fof(f7971,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f7324,f35])).
fof(f20367,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X6))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X6,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X6,sK0)),s(fun(sK0,bool),X6))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0))))) ) | $spl90),
  inference(trivial_inequality_removal,[],[f20357])).
fof(f20357,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X6))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X6,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X6,sK0)),s(fun(sK0,bool),X6))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f109,f13688])).
fof(f13688,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl90),
  inference(trivial_inequality_removal,[],[f13687])).
fof(f13687,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f7969,f35])).
fof(f20368,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X5))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0))))) ) | $spl90),
  inference(trivial_inequality_removal,[],[f20356])).
fof(f20356,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),X5))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK1,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f109,f13676])).
fof(f13676,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl90),
  inference(trivial_inequality_removal,[],[f13674])).
fof(f13674,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f7968,f35])).
fof(f7968,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl90),
  inference(superposition,[],[f7323,f35])).
fof(f11949,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl216),
  inference(cnf_transformation,[],[f11949_D])).
fof(f11949_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f11572,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl206),
  inference(cnf_transformation,[],[f11572_D])).
fof(f11572_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f17545,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(trivial_inequality_removal,[],[f17525])).
fof(f17525,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f12209,f9623])).
fof(f9623,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(subsumption_resolution,[],[f9620,f43])).
fof(f9620,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(superposition,[],[f7144,f35])).
fof(f7144,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(resolution,[],[f7124,f5918])).
fof(f5918,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5904,f44])).
fof(f5904,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X3))))) ) | $spl18),
  inference(superposition,[],[f38,f674])).
fof(f7124,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(subsumption_resolution,[],[f7116,f44])).
fof(f7116,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(superposition,[],[f5918,f6881])).
fof(f12209,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f12208])).
fof(f12208,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(X1,sK1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6722,f35])).
fof(f17873,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f17517,f71])).
fof(f17517,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f10397,f9623])).
fof(f17876,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f17498,f71])).
fof(f17498,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f9623,f10397])).
fof(f17879,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f15000,f43])).
fof(f15000,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl88 | $spl90)),
  inference(superposition,[],[f8292,f13717])).
fof(f8292,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl88),
  inference(resolution,[],[f8289,f7312])).
fof(f7312,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2))))))) ) | $spl88),
  inference(subsumption_resolution,[],[f7297,f44])).
fof(f7297,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X3))))) ) | $spl88),
  inference(superposition,[],[f38,f7261])).
fof(f8289,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl88),
  inference(subsumption_resolution,[],[f8281,f44])).
fof(f8281,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl88),
  inference(superposition,[],[f7312,f8083])).
fof(f8083,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl88),
  inference(subsumption_resolution,[],[f8080,f43])).
fof(f8080,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl88),
  inference(superposition,[],[f8049,f35])).
fof(f8049,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl88),
  inference(subsumption_resolution,[],[f8037,f44])).
fof(f8037,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl88),
  inference(superposition,[],[f7311,f7261])).
fof(f7311,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2))))))) ) | $spl88),
  inference(subsumption_resolution,[],[f7296,f44])).
fof(f7296,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X2))))) ) | $spl88),
  inference(superposition,[],[f38,f7261])).
fof(f11792,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f11735,f43])).
fof(f11735,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f7144,f10397])).
fof(f17899,plain,(
  $false | ($spl179 | $spl217)),
  inference(subsumption_resolution,[],[f17898,f11950])).
fof(f17898,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl179),
  inference(trivial_inequality_removal,[],[f17897])).
fof(f17897,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl179),
  inference(superposition,[],[f11125,f35])).
fof(f11125,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl179),
  inference(cnf_transformation,[],[f11125_D])).
fof(f11125_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f11126,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl178),
  inference(cnf_transformation,[],[f11126_D])).
fof(f11126_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f11128,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl180),
  inference(cnf_transformation,[],[f11128_D])).
fof(f11128_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f17881,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f14888,f43])).
fof(f14888,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl88 | $spl90)),
  inference(superposition,[],[f8292,f13706])).
fof(f17878,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f17497,f71])).
fof(f17497,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f9623,f10393])).
fof(f17875,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f17516,f71])).
fof(f17516,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f10393,f9623])).
fof(f17546,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(trivial_inequality_removal,[],[f17524])).
fof(f17524,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f12198,f9623])).
fof(f12198,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f12196])).
fof(f12196,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6721,f35])).
fof(f6721,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(X0,sK1,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f5934,f35])).
fof(f17871,plain,(
  $false | ($spl3 | $spl18 | $spl136)),
  inference(subsumption_resolution,[],[f17870,f71])).
fof(f17870,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl18 | $spl136)),
  inference(forward_demodulation,[],[f10196,f9623])).
fof(f10898,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl156),
  inference(cnf_transformation,[],[f10898_D])).
fof(f10898_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f8944,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl108),
  inference(cnf_transformation,[],[f8944_D])).
fof(f8944_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f17454,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(trivial_inequality_removal,[],[f17430])).
fof(f17430,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f12180,f9592])).
fof(f9592,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(subsumption_resolution,[],[f9589,f43])).
fof(f9589,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(superposition,[],[f7125,f35])).
fof(f7125,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(subsumption_resolution,[],[f7117,f44])).
fof(f7117,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(superposition,[],[f5918,f7000])).
fof(f7000,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl18),
  inference(subsumption_resolution,[],[f6997,f43])).
fof(f6997,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl18),
  inference(superposition,[],[f6860,f35])).
fof(f6860,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl18),
  inference(resolution,[],[f6858,f5917])).
fof(f12180,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f12179])).
fof(f12179,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6719,f35])).
fof(f17702,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f17422,f71])).
fof(f17422,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f10177,f9592])).
fof(f17705,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f17403,f71])).
fof(f17403,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f9592,f10177])).
fof(f17708,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f14775,f43])).
fof(f14775,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl88 | $spl90)),
  inference(superposition,[],[f8290,f13688])).
fof(f8290,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)))))) | $spl88),
  inference(subsumption_resolution,[],[f8282,f44])).
fof(f8282,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)))))) | $spl88),
  inference(superposition,[],[f7312,f8160])).
fof(f8160,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl88),
  inference(subsumption_resolution,[],[f8157,f43])).
fof(f8157,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl88),
  inference(superposition,[],[f8062,f35])).
fof(f8062,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl88),
  inference(resolution,[],[f8049,f7311])).
fof(f10787,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f10732,f43])).
fof(f10732,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f7125,f10177])).
fof(f17729,plain,(
  $false | ($spl37 | $spl157)),
  inference(subsumption_resolution,[],[f17728,f10899])).
fof(f17728,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl37),
  inference(trivial_inequality_removal,[],[f17727])).
fof(f17727,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl37),
  inference(superposition,[],[f2961,f35])).
fof(f2961,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl37),
  inference(cnf_transformation,[],[f2961_D])).
fof(f2961_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f2962,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f2962_D])).
fof(f2962_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2964,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl38),
  inference(cnf_transformation,[],[f2964_D])).
fof(f2964_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f17712,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f10649,f43])).
fof(f10649,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f7125,f10165])).
fof(f17710,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f14663,f43])).
fof(f14663,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl88 | $spl90)),
  inference(superposition,[],[f8290,f13676])).
fof(f17707,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f17402,f71])).
fof(f17402,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f9592,f10165])).
fof(f17704,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f17421,f71])).
fof(f17421,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f10165,f9592])).
fof(f17455,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(trivial_inequality_removal,[],[f17429])).
fof(f17429,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f12168,f9592])).
fof(f12168,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f12166])).
fof(f12166,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f6718,f35])).
fof(f6718,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f5933,f35])).
fof(f17700,plain,(
  $false | ($spl3 | $spl18 | $spl134)),
  inference(subsumption_resolution,[],[f17699,f71])).
fof(f17699,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl18 | $spl134)),
  inference(forward_demodulation,[],[f10194,f9592])).
fof(f17697,plain,(
  $false | ($spl3 | $spl18 | $spl132)),
  inference(subsumption_resolution,[],[f17696,f71])).
fof(f17696,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl18 | $spl132)),
  inference(forward_demodulation,[],[f10190,f9555])).
fof(f9555,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | $spl18),
  inference(subsumption_resolution,[],[f9552,f43])).
fof(f9552,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | $spl18),
  inference(superposition,[],[f6975,f35])).
fof(f6975,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl18),
  inference(resolution,[],[f6860,f5917])).
fof(f15195,plain,(
  ( ! [X84] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X84)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X84)))) ) | $spl222),
  inference(cnf_transformation,[],[f15195_D])).
fof(f15195_D,plain,(
  ( ! [X84] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X84)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X84)))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f15196,plain,(
  ( ! [X85] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X85))))) ) | $spl224),
  inference(cnf_transformation,[],[f15196_D])).
fof(f15196_D,plain,(
  ( ! [X85] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X85))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f15194,plain,(
  ( ! [X85,X84] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X85)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X84)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X84)))) ) | $spl88),
  inference(subsumption_resolution,[],[f15168,f43])).
fof(f15168,plain,(
  ( ! [X85,X84] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X85)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X84)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X84)))) ) | $spl88),
  inference(superposition,[],[f38,f8826])).
fof(f8826,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X29),s(fun(sK0,bool),X30))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X29)))) ) | $spl88),
  inference(subsumption_resolution,[],[f8821,f43])).
fof(f8821,plain,(
  ( ! [X30,X29] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X29),s(fun(sK0,bool),X30))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X29)))) ) | $spl88),
  inference(superposition,[],[f8742,f35])).
fof(f8742,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X10)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X10),s(fun(sK0,bool),X11)))))) ) | $spl88),
  inference(resolution,[],[f8710,f36])).
fof(f8710,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X7)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X7)))) ) | $spl88),
  inference(resolution,[],[f8050,f37])).
fof(f8050,plain,(
  ( ! [X14] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X14)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X14)))) ) | $spl88),
  inference(subsumption_resolution,[],[f8047,f44])).
fof(f8047,plain,(
  ( ! [X14] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X14)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),X14)))) ) | $spl88),
  inference(superposition,[],[f7311,f35])).
fof(f13251,plain,(
  ( ! [X92] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X92)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X92)))) ) | $spl218),
  inference(cnf_transformation,[],[f13251_D])).
fof(f13251_D,plain,(
  ( ! [X92] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X92)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X92)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f13252,plain,(
  ( ! [X93] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X93))))) ) | $spl220),
  inference(cnf_transformation,[],[f13252_D])).
fof(f13252_D,plain,(
  ( ! [X93] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X93))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f13250,plain,(
  ( ! [X92,X93] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X93)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X92)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X92)))) ) | $spl18),
  inference(subsumption_resolution,[],[f13220,f43])).
fof(f13220,plain,(
  ( ! [X92,X93] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X93)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X92)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X92)))) ) | $spl18),
  inference(superposition,[],[f38,f8661])).
fof(f8661,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X29),s(fun(sK0,bool),X30))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f8656,f43])).
fof(f8656,plain,(
  ( ! [X30,X29] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X29),s(fun(sK0,bool),X30))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X29)))) ) | $spl18),
  inference(superposition,[],[f8577,f35])).
fof(f8577,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X10)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X10),s(fun(sK0,bool),X11)))))) ) | $spl18),
  inference(resolution,[],[f8534,f36])).
fof(f8534,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X7)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X7)))) ) | $spl18),
  inference(resolution,[],[f6859,f37])).
fof(f6859,plain,(
  ( ! [X14] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X14)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X14)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6856,f44])).
fof(f6856,plain,(
  ( ! [X14] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X14)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X14)))) ) | $spl18),
  inference(superposition,[],[f5917,f35])).
fof(f10952,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | $spl158),
  inference(cnf_transformation,[],[f10952_D])).
fof(f10952_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f11950,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl217),
  inference(cnf_transformation,[],[f11950_D])).
fof(f11950_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f11931,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | ($spl3 | $spl18 | $spl178)),
  inference(superposition,[],[f32,f11794])).
fof(f11794,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18 | $spl178)),
  inference(subsumption_resolution,[],[f11793,f71])).
fof(f11793,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl178)),
  inference(forward_demodulation,[],[f11792,f11126])).
fof(f10897,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) | $spl154),
  inference(cnf_transformation,[],[f10897_D])).
fof(f10897_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f11897,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl215),
  inference(cnf_transformation,[],[f11897_D])).
fof(f11897_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f11878,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) | ($spl3 | $spl18 | $spl174)),
  inference(superposition,[],[f32,f11791])).
fof(f11791,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18 | $spl174)),
  inference(subsumption_resolution,[],[f11790,f71])).
fof(f11790,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl174)),
  inference(forward_demodulation,[],[f11789,f11121])).
fof(f11789,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f11734,f43])).
fof(f11734,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f7125,f10397])).
fof(f10842,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl150),
  inference(cnf_transformation,[],[f10842_D])).
fof(f10842_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f11844,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl213),
  inference(cnf_transformation,[],[f11844_D])).
fof(f11844_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f11825,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | ($spl3 | $spl18 | $spl182)),
  inference(superposition,[],[f32,f11797])).
fof(f11797,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | ($spl3 | $spl18 | $spl182)),
  inference(subsumption_resolution,[],[f11796,f71])).
fof(f11796,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | ($spl18 | $spl182)),
  inference(forward_demodulation,[],[f11795,f11133])).
fof(f11795,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f11742,f43])).
fof(f11742,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f6975,f10397])).
fof(f9050,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl130),
  inference(cnf_transformation,[],[f9050_D])).
fof(f9050_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f11665,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl211),
  inference(cnf_transformation,[],[f11665_D])).
fof(f11665_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f11646,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl208),
  inference(superposition,[],[f32,f11612])).
fof(f11612,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl208),
  inference(cnf_transformation,[],[f11612_D])).
fof(f11612_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f9040,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | $spl118),
  inference(cnf_transformation,[],[f9040_D])).
fof(f9040_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f11609,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl186)),
  inference(trivial_inequality_removal,[],[f11599])).
fof(f11599,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl186)),
  inference(superposition,[],[f6704,f11138])).
fof(f9038,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | $spl117),
  inference(cnf_transformation,[],[f9038_D])).
fof(f9038_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f11610,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl186),
  inference(trivial_inequality_removal,[],[f11597])).
fof(f11597,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl186),
  inference(superposition,[],[f64,f11138])).
fof(f10196,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl136),
  inference(cnf_transformation,[],[f10196_D])).
fof(f10196_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f11569,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl178)),
  inference(trivial_inequality_removal,[],[f11559])).
fof(f11559,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl178)),
  inference(superposition,[],[f6704,f11126])).
fof(f8996,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl115),
  inference(cnf_transformation,[],[f8996_D])).
fof(f8996_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f11570,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl178),
  inference(trivial_inequality_removal,[],[f11557])).
fof(f11557,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl178),
  inference(superposition,[],[f64,f11126])).
fof(f11532,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl204),
  inference(cnf_transformation,[],[f11532_D])).
fof(f11532_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f10194,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | $spl134),
  inference(cnf_transformation,[],[f10194_D])).
fof(f10194_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f11529,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl174)),
  inference(trivial_inequality_removal,[],[f11519])).
fof(f11519,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl174)),
  inference(superposition,[],[f6704,f11121])).
fof(f8946,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | $spl111),
  inference(cnf_transformation,[],[f8946_D])).
fof(f8946_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f11530,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl174),
  inference(trivial_inequality_removal,[],[f11517])).
fof(f11517,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl174),
  inference(superposition,[],[f64,f11121])).
fof(f11492,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl202),
  inference(cnf_transformation,[],[f11492_D])).
fof(f11492_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f10190,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | $spl132),
  inference(cnf_transformation,[],[f10190_D])).
fof(f10190_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f11489,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | ($spl18 | $spl182)),
  inference(trivial_inequality_removal,[],[f11479])).
fof(f11479,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | ($spl18 | $spl182)),
  inference(superposition,[],[f6704,f11133])).
fof(f8907,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | $spl107),
  inference(cnf_transformation,[],[f8907_D])).
fof(f8907_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f11490,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl182),
  inference(trivial_inequality_removal,[],[f11477])).
fof(f11477,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl182),
  inference(superposition,[],[f64,f11133])).
fof(f10341,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | $spl146),
  inference(cnf_transformation,[],[f10341_D])).
fof(f10341_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11417,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl201),
  inference(cnf_transformation,[],[f11417_D])).
fof(f11417_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f11398,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | ($spl3 | $spl18 | $spl162)),
  inference(superposition,[],[f32,f11308])).
fof(f11308,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18 | $spl162)),
  inference(subsumption_resolution,[],[f11307,f71])).
fof(f11307,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl162)),
  inference(forward_demodulation,[],[f11303,f7170])).
fof(f7170,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(subsumption_resolution,[],[f7167,f43])).
fof(f7167,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(superposition,[],[f7124,f35])).
fof(f11303,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl162)),
  inference(trivial_inequality_removal,[],[f11293])).
fof(f11293,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl162)),
  inference(superposition,[],[f6704,f11102])).
fof(f10290,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl142),
  inference(cnf_transformation,[],[f10290_D])).
fof(f10290_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f11357,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl199),
  inference(cnf_transformation,[],[f11357_D])).
fof(f11357_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f11338,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl18 | $spl170)),
  inference(superposition,[],[f32,f11270])).
fof(f11270,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl3 | $spl18 | $spl170)),
  inference(subsumption_resolution,[],[f11269,f71])).
fof(f11269,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl170)),
  inference(forward_demodulation,[],[f11265,f7000])).
fof(f11265,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl170)),
  inference(trivial_inequality_removal,[],[f11255])).
fof(f11255,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl18 | $spl170)),
  inference(superposition,[],[f6704,f11112])).
fof(f11306,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl196),
  inference(cnf_transformation,[],[f11306_D])).
fof(f11306_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f7545,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl103),
  inference(cnf_transformation,[],[f7545_D])).
fof(f7545_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f11304,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl162),
  inference(trivial_inequality_removal,[],[f11291])).
fof(f11291,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl162),
  inference(superposition,[],[f64,f11102])).
fof(f11268,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl194),
  inference(cnf_transformation,[],[f11268_D])).
fof(f11268_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7541,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl99),
  inference(cnf_transformation,[],[f7541_D])).
fof(f7541_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f11266,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl170),
  inference(trivial_inequality_removal,[],[f11253])).
fof(f11253,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl170),
  inference(superposition,[],[f64,f11112])).
fof(f10231,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | $spl138),
  inference(cnf_transformation,[],[f10231_D])).
fof(f10231_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f11213,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl193),
  inference(cnf_transformation,[],[f11213_D])).
fof(f11213_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl193),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl193])])).
fof(f11194,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | ($spl3 | $spl18 | $spl166)),
  inference(superposition,[],[f32,f11168])).
fof(f11168,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18 | $spl166)),
  inference(subsumption_resolution,[],[f11167,f71])).
fof(f11167,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl166)),
  inference(forward_demodulation,[],[f11163,f6881])).
fof(f11163,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl166)),
  inference(trivial_inequality_removal,[],[f11153])).
fof(f11153,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl18 | $spl166)),
  inference(superposition,[],[f6704,f11107])).
fof(f11166,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl190),
  inference(cnf_transformation,[],[f11166_D])).
fof(f11166_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7537,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl95),
  inference(cnf_transformation,[],[f7537_D])).
fof(f7537_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f11164,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl166),
  inference(trivial_inequality_removal,[],[f11151])).
fof(f11151,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl166),
  inference(superposition,[],[f64,f11107])).
fof(f11138,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl186),
  inference(cnf_transformation,[],[f11138_D])).
fof(f11138_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f11140,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl188),
  inference(cnf_transformation,[],[f11140_D])).
fof(f11140_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f11136,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f11057,f43])).
fof(f11057,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f7145,f10393])).
fof(f7145,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))))) | $spl18),
  inference(resolution,[],[f7124,f5917])).
fof(f11133,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl182),
  inference(cnf_transformation,[],[f11133_D])).
fof(f11133_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f11135,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl184),
  inference(cnf_transformation,[],[f11135_D])).
fof(f11135_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f11131,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f11056,f43])).
fof(f11056,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f6975,f10393])).
fof(f11112,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl170),
  inference(cnf_transformation,[],[f11112_D])).
fof(f11112_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f11114,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl172),
  inference(cnf_transformation,[],[f11114_D])).
fof(f11114_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f11130,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f11055,f71])).
fof(f11055,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f7000,f10393])).
fof(f11107,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl166),
  inference(cnf_transformation,[],[f11107_D])).
fof(f11107_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f11109,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl168),
  inference(cnf_transformation,[],[f11109_D])).
fof(f11109_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f11129,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f11054,f71])).
fof(f11054,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f6881,f10393])).
fof(f11124,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f11049,f43])).
fof(f11049,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f7144,f10393])).
fof(f11121,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl174),
  inference(cnf_transformation,[],[f11121_D])).
fof(f11121_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f11123,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl176),
  inference(cnf_transformation,[],[f11123_D])).
fof(f11123_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f11119,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f11048,f43])).
fof(f11048,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f7125,f10393])).
fof(f11102,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl162),
  inference(cnf_transformation,[],[f11102_D])).
fof(f11102_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f11104,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl164),
  inference(cnf_transformation,[],[f11104_D])).
fof(f11104_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f11118,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f11047,f71])).
fof(f11047,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f7170,f10393])).
fof(f11110,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f11030,f71])).
fof(f11030,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f10393,f7000])).
fof(f11105,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f11029,f71])).
fof(f11029,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f10393,f6881])).
fof(f11100,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f11028,f71])).
fof(f11028,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK1,bool,sK0)))) | $spl18),
  inference(superposition,[],[f10393,f7170])).
fof(f10954,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl161),
  inference(cnf_transformation,[],[f10954_D])).
fof(f10954_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f10933,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | ($spl3 | $spl18 | $spl40)),
  inference(superposition,[],[f32,f10792])).
fof(f10792,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18 | $spl40)),
  inference(subsumption_resolution,[],[f10791,f71])).
fof(f10791,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl40)),
  inference(forward_demodulation,[],[f10790,f2967])).
fof(f10790,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f10733,f43])).
fof(f10733,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f7144,f10177])).
fof(f10899,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl157),
  inference(cnf_transformation,[],[f10899_D])).
fof(f10899_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f10878,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) | ($spl3 | $spl18 | $spl36)),
  inference(superposition,[],[f32,f10789])).
fof(f10789,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18 | $spl36)),
  inference(subsumption_resolution,[],[f10788,f71])).
fof(f10788,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl36)),
  inference(forward_demodulation,[],[f10787,f2962])).
fof(f10844,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl153),
  inference(cnf_transformation,[],[f10844_D])).
fof(f10844_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f10823,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | ($spl3 | $spl18 | $spl44)),
  inference(superposition,[],[f32,f10795])).
fof(f10795,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | ($spl3 | $spl18 | $spl44)),
  inference(subsumption_resolution,[],[f10794,f71])).
fof(f10794,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | ($spl18 | $spl44)),
  inference(forward_demodulation,[],[f10793,f2974])).
fof(f10793,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f10740,f43])).
fof(f10740,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl18),
  inference(superposition,[],[f6975,f10177])).
fof(f10343,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl149),
  inference(cnf_transformation,[],[f10343_D])).
fof(f10343_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f10322,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | ($spl3 | $spl18 | $spl24)),
  inference(superposition,[],[f32,f10192])).
fof(f10192,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18 | $spl24)),
  inference(subsumption_resolution,[],[f10191,f71])).
fof(f10191,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f10180,f7170])).
fof(f10180,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl24)),
  inference(trivial_inequality_removal,[],[f10172])).
fof(f10172,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl24)),
  inference(superposition,[],[f6701,f2943])).
fof(f10292,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl145),
  inference(cnf_transformation,[],[f10292_D])).
fof(f10292_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f10271,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl18 | $spl32)),
  inference(superposition,[],[f32,f10188])).
fof(f10188,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl3 | $spl18 | $spl32)),
  inference(subsumption_resolution,[],[f10187,f71])).
fof(f10187,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl18 | $spl32)),
  inference(forward_demodulation,[],[f10182,f7000])).
fof(f10182,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl18 | $spl32)),
  inference(trivial_inequality_removal,[],[f10169])).
fof(f10169,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl18 | $spl32)),
  inference(superposition,[],[f6701,f2953])).
fof(f10233,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl141),
  inference(cnf_transformation,[],[f10233_D])).
fof(f10233_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f10212,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | ($spl3 | $spl18 | $spl28)),
  inference(superposition,[],[f32,f10186])).
fof(f10186,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18 | $spl28)),
  inference(subsumption_resolution,[],[f10185,f71])).
fof(f10185,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl28)),
  inference(forward_demodulation,[],[f10183,f6881])).
fof(f10183,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl28)),
  inference(trivial_inequality_removal,[],[f10168])).
fof(f10168,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl28)),
  inference(superposition,[],[f6701,f2948])).
fof(f8994,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl112),
  inference(cnf_transformation,[],[f8994_D])).
fof(f8994_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f10178,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl40)),
  inference(trivial_inequality_removal,[],[f10174])).
fof(f10174,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl40)),
  inference(superposition,[],[f6701,f2967])).
fof(f10179,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl36)),
  inference(trivial_inequality_removal,[],[f10173])).
fof(f10173,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl18 | $spl36)),
  inference(superposition,[],[f6701,f2962])).
fof(f8905,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl104),
  inference(cnf_transformation,[],[f8905_D])).
fof(f8905_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f10181,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | ($spl18 | $spl44)),
  inference(trivial_inequality_removal,[],[f10170])).
fof(f10170,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | ($spl18 | $spl44)),
  inference(superposition,[],[f6701,f2974])).
fof(f2978,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl49),
  inference(cnf_transformation,[],[f2978_D])).
fof(f2978_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f9028,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl50),
  inference(superposition,[],[f32,f2981])).
fof(f9048,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | $spl128),
  inference(cnf_transformation,[],[f9048_D])).
fof(f9048_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f9026,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | ($spl50 | $spl90)),
  inference(superposition,[],[f7323,f2981])).
fof(f9046,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | $spl127),
  inference(cnf_transformation,[],[f9046_D])).
fof(f9046_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f9025,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | ($spl50 | $spl87)),
  inference(superposition,[],[f7264,f2981])).
fof(f9044,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))))) ) | $spl124),
  inference(cnf_transformation,[],[f9044_D])).
fof(f9044_D,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f9024,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))))) ) | ($spl18 | $spl50)),
  inference(superposition,[],[f6453,f2981])).
fof(f9043,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),X0)))) ) | $spl122),
  inference(cnf_transformation,[],[f9043_D])).
fof(f9043_D,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),X0)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f9023,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),X0)))) ) | ($spl18 | $spl50)),
  inference(superposition,[],[f6401,f2981])).
fof(f9042,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | $spl121),
  inference(cnf_transformation,[],[f9042_D])).
fof(f9042_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f9022,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | ($spl18 | $spl50)),
  inference(superposition,[],[f5933,f2981])).
fof(f9021,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | ($spl18 | $spl50)),
  inference(superposition,[],[f5849,f2981])).
fof(f9020,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))))) | $spl50),
  inference(superposition,[],[f49,f2981])).
fof(f8992,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl40),
  inference(trivial_inequality_removal,[],[f8982])).
fof(f8982,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl40),
  inference(superposition,[],[f62,f2967])).
fof(f8942,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl36),
  inference(trivial_inequality_removal,[],[f8932])).
fof(f8932,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl36),
  inference(superposition,[],[f62,f2962])).
fof(f8903,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl44),
  inference(trivial_inequality_removal,[],[f8893])).
fof(f8893,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl44),
  inference(superposition,[],[f62,f2974])).
fof(f7543,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl100),
  inference(cnf_transformation,[],[f7543_D])).
fof(f7543_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f7530,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl24),
  inference(trivial_inequality_removal,[],[f7526])).
fof(f7526,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl24),
  inference(superposition,[],[f62,f2943])).
fof(f7539,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl96),
  inference(cnf_transformation,[],[f7539_D])).
fof(f7539_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f7531,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl32),
  inference(trivial_inequality_removal,[],[f7525])).
fof(f7525,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl32),
  inference(superposition,[],[f62,f2953])).
fof(f7535,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl92),
  inference(cnf_transformation,[],[f7535_D])).
fof(f7535_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f7532,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl28),
  inference(trivial_inequality_removal,[],[f7524])).
fof(f7524,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl28),
  inference(superposition,[],[f62,f2948])).
fof(f7282,plain,(
  $false | ($spl87 | $spl91)),
  inference(subsumption_resolution,[],[f7281,f7259])).
fof(f7281,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl91),
  inference(trivial_inequality_removal,[],[f7280])).
fof(f7280,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl91),
  inference(superposition,[],[f7263,f35])).
fof(f7261,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl88),
  inference(cnf_transformation,[],[f7261_D])).
fof(f7261_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f7263,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl91),
  inference(cnf_transformation,[],[f7263_D])).
fof(f7263_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f7255,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl9),
  inference(superposition,[],[f7247,f35])).
fof(f7247,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl9),
  inference(equality_resolution,[],[f5810])).
fof(f5810,plain,(
  ( ! [X3] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X3) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK2,sK0)),s(fun(sK0,bool),X3)))) ) | $spl9),
  inference(superposition,[],[f194,f34])).
fof(f194,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2) | $spl9),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f7257,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl84),
  inference(cnf_transformation,[],[f7257_D])).
fof(f7257_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7259,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl87),
  inference(cnf_transformation,[],[f7259_D])).
fof(f7259_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f7254,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl9),
  inference(superposition,[],[f7247,f35])).
fof(f5758,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f5607])).
fof(f5607,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | $spl8),
  inference(backward_demodulation,[],[f195,f31])).
fof(f195,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | $spl8),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f643,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl19),
  inference(cnf_transformation,[],[f643_D])).
fof(f643_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f3011,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | $spl20),
  inference(superposition,[],[f34,f646])).
fof(f5604,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl80),
  inference(cnf_transformation,[],[f5604_D])).
fof(f5604_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5606,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl82),
  inference(cnf_transformation,[],[f5606_D])).
fof(f5606_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f724,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl22),
  inference(cnf_transformation,[],[f724_D])).
fof(f724_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5550,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(trivial_inequality_removal,[],[f5512])).
fof(f5512,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(superposition,[],[f90,f5069])).
fof(f5069,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,X0,bool,sK0))))) ) | $spl20),
  inference(trivial_inequality_removal,[],[f5067])).
fof(f5067,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,X0,bool,sK0))))) ) | $spl20),
  inference(superposition,[],[f3708,f35])).
fof(f3708,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK2,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X0,bool,sK0))))) ) | $spl20),
  inference(superposition,[],[f2998,f35])).
fof(f2998,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK2,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK2,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0)))) ) | $spl20),
  inference(superposition,[],[f646,f32])).
fof(f90,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),sK2))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X5))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X5,sK0)),s(fun(sK0,bool),X5)))) )),
  inference(superposition,[],[f47,f35])).
fof(f47,plain,(
  ( ! [X2] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X2,sK0)),s(fun(sK0,bool),X2))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2)))) )),
  inference(superposition,[],[f31,f34])).
fof(f5548,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(trivial_inequality_removal,[],[f5514])).
fof(f5514,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(superposition,[],[f47,f5069])).
fof(f5543,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(trivial_inequality_removal,[],[f5519])).
fof(f5519,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(superposition,[],[f253,f5069])).
fof(f253,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK1,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X5))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK1,sK0)),s(fun(sK0,bool),sK1)))) )),
  inference(superposition,[],[f52,f35])).
fof(f5541,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(trivial_inequality_removal,[],[f5521])).
fof(f5521,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(superposition,[],[f52,f5069])).
fof(f5522,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK2,bool,sK0)))) | $spl20),
  inference(superposition,[],[f34,f5069])).
fof(f5593,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl76),
  inference(cnf_transformation,[],[f5593_D])).
fof(f5593_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5595,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl78),
  inference(cnf_transformation,[],[f5595_D])).
fof(f5595_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5591,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5509,f43])).
fof(f5509,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f4038,f5069])).
fof(f4038,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))))) | $spl19),
  inference(resolution,[],[f4031,f3076])).
fof(f3076,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2))))))) ) | $spl19),
  inference(subsumption_resolution,[],[f3059,f44])).
fof(f3059,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X2))))) ) | $spl19),
  inference(superposition,[],[f38,f2994])).
fof(f2994,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl19),
  inference(trivial_inequality_removal,[],[f2993])).
fof(f2993,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl19),
  inference(superposition,[],[f643,f35])).
fof(f4031,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl19),
  inference(subsumption_resolution,[],[f4022,f44])).
fof(f4022,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl19),
  inference(superposition,[],[f3077,f3804])).
fof(f3804,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl19),
  inference(subsumption_resolution,[],[f3801,f43])).
fof(f3801,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl19),
  inference(superposition,[],[f3781,f35])).
fof(f3781,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl19),
  inference(subsumption_resolution,[],[f3770,f44])).
fof(f3770,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl19),
  inference(superposition,[],[f3076,f2994])).
fof(f3077,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK1))))))) ) | $spl19),
  inference(subsumption_resolution,[],[f3060,f44])).
fof(f3060,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X3))))) ) | $spl19),
  inference(superposition,[],[f38,f2994])).
fof(f5588,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),bool,sK0)))) | $spl72),
  inference(cnf_transformation,[],[f5588_D])).
fof(f5588_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),bool,sK0)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f5590,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),bool,sK0)))) | $spl74),
  inference(cnf_transformation,[],[f5590_D])).
fof(f5590_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),bool,sK0)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5586,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),bool,sK0)))) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5508,f43])).
fof(f5508,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f3896,f5069])).
fof(f3896,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))))) | $spl19),
  inference(resolution,[],[f3783,f3076])).
fof(f3783,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl19),
  inference(resolution,[],[f3781,f3076])).
fof(f5567,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl60),
  inference(cnf_transformation,[],[f5567_D])).
fof(f5567_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f5569,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl62),
  inference(cnf_transformation,[],[f5569_D])).
fof(f5569_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5585,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl3 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5507,f71])).
fof(f5507,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f3921,f5069])).
fof(f3921,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))) | $spl19),
  inference(subsumption_resolution,[],[f3918,f43])).
fof(f3918,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))) | $spl19),
  inference(superposition,[],[f3783,f35])).
fof(f5562,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl56),
  inference(cnf_transformation,[],[f5562_D])).
fof(f5562_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5564,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl58),
  inference(cnf_transformation,[],[f5564_D])).
fof(f5564_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5584,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl3 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5506,f71])).
fof(f5506,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f3804,f5069])).
fof(f5581,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl68),
  inference(cnf_transformation,[],[f5581_D])).
fof(f5581_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5583,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl70),
  inference(cnf_transformation,[],[f5583_D])).
fof(f5583_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5579,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5501,f43])).
fof(f5501,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f4037,f5069])).
fof(f4037,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl19),
  inference(resolution,[],[f4031,f3077])).
fof(f5576,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl64),
  inference(cnf_transformation,[],[f5576_D])).
fof(f5576_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5578,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl66),
  inference(cnf_transformation,[],[f5578_D])).
fof(f5578_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5574,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5500,f43])).
fof(f5500,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f4032,f5069])).
fof(f4032,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)))))) | $spl19),
  inference(subsumption_resolution,[],[f4023,f44])).
fof(f4023,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK1)))))) | $spl19),
  inference(superposition,[],[f3077,f3921])).
fof(f5557,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl52),
  inference(cnf_transformation,[],[f5557_D])).
fof(f5557_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5559,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl54),
  inference(cnf_transformation,[],[f5559_D])).
fof(f5559_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5573,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl3 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5499,f71])).
fof(f5499,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f4063,f5069])).
fof(f4063,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl19),
  inference(subsumption_resolution,[],[f4060,f43])).
fof(f4060,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl19),
  inference(superposition,[],[f4031,f35])).
fof(f5565,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl3 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5478,f71])).
fof(f5478,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f5069,f3921])).
fof(f5560,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl3 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5477,f71])).
fof(f5477,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f5069,f3804])).
fof(f5555,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl3 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f5476,f71])).
fof(f5476,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl19 | $spl20)),
  inference(superposition,[],[f5069,f4063])).
fof(f3088,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f725,f2994])).
fof(f2979,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl48),
  inference(cnf_transformation,[],[f2979_D])).
fof(f2979_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2981,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl50),
  inference(cnf_transformation,[],[f2981_D])).
fof(f2981_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2977,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f2898,f43])).
fof(f2898,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl18),
  inference(superposition,[],[f1743,f2640])).
fof(f2640,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f2638])).
fof(f2638,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f1393,f35])).
fof(f1393,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0))))) ) | $spl18),
  inference(superposition,[],[f647,f35])).
fof(f647,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X0)))) ) | $spl18),
  inference(superposition,[],[f644,f32])).
fof(f1743,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))))) | $spl18),
  inference(resolution,[],[f1739,f726])).
fof(f726,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f705,f44])).
fof(f705,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X2))))) ) | $spl18),
  inference(superposition,[],[f38,f674])).
fof(f1739,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(subsumption_resolution,[],[f1731,f44])).
fof(f1731,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(superposition,[],[f727,f1504])).
fof(f1504,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(subsumption_resolution,[],[f1501,f43])).
fof(f1501,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(superposition,[],[f1481,f35])).
fof(f1481,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(subsumption_resolution,[],[f1470,f44])).
fof(f1470,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(superposition,[],[f726,f674])).
fof(f727,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f706,f44])).
fof(f706,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),X3))))) ) | $spl18),
  inference(superposition,[],[f38,f674])).
fof(f2974,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl44),
  inference(cnf_transformation,[],[f2974_D])).
fof(f2974_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2976,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl46),
  inference(cnf_transformation,[],[f2976_D])).
fof(f2976_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2972,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f2897,f43])).
fof(f2897,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl18),
  inference(superposition,[],[f1598,f2640])).
fof(f1598,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))))) | $spl18),
  inference(resolution,[],[f1483,f726])).
fof(f1483,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))))) | $spl18),
  inference(resolution,[],[f1481,f726])).
fof(f2953,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f2953_D])).
fof(f2953_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2955,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl34),
  inference(cnf_transformation,[],[f2955_D])).
fof(f2955_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2971,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f2896,f71])).
fof(f2896,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl18),
  inference(superposition,[],[f1623,f2640])).
fof(f1623,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl18),
  inference(subsumption_resolution,[],[f1620,f43])).
fof(f1620,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))))) | $spl18),
  inference(superposition,[],[f1483,f35])).
fof(f2948,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f2948_D])).
fof(f2948_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2950,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl30),
  inference(cnf_transformation,[],[f2950_D])).
fof(f2950_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2970,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f2895,f71])).
fof(f2895,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f1504,f2640])).
fof(f2967,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl40),
  inference(cnf_transformation,[],[f2967_D])).
fof(f2967_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2969,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl42),
  inference(cnf_transformation,[],[f2969_D])).
fof(f2969_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2965,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f2890,f43])).
fof(f2890,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f1742,f2640])).
fof(f1742,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(resolution,[],[f1739,f727])).
fof(f2960,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(subsumption_resolution,[],[f2889,f43])).
fof(f2889,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f1740,f2640])).
fof(f1740,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(subsumption_resolution,[],[f1732,f44])).
fof(f1732,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(superposition,[],[f727,f1623])).
fof(f2943,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl24),
  inference(cnf_transformation,[],[f2943_D])).
fof(f2943_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2945,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl26),
  inference(cnf_transformation,[],[f2945_D])).
fof(f2945_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2959,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f2888,f71])).
fof(f2888,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f1768,f2640])).
fof(f1768,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(subsumption_resolution,[],[f1765,f43])).
fof(f1765,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(superposition,[],[f1739,f35])).
fof(f2951,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f2871,f71])).
fof(f2871,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))),bool,sK0)))) | $spl18),
  inference(superposition,[],[f2640,f1623])).
fof(f2946,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f2870,f71])).
fof(f2870,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f2640,f1504])).
fof(f2941,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl18)),
  inference(subsumption_resolution,[],[f2869,f71])).
fof(f2869,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f2640,f1768])).
fof(f725,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl23),
  inference(cnf_transformation,[],[f725_D])).
fof(f725_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f702,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | $spl18),
  inference(superposition,[],[f34,f674])).
fof(f644,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl18),
  inference(cnf_transformation,[],[f644_D])).
fof(f644_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f646,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl20),
  inference(cnf_transformation,[],[f646_D])).
fof(f646_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f642,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))),
  inference(equality_resolution,[],[f456])).
fof(f456,plain,(
  ( ! [X4] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X4,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X4,sK0)),s(fun(sK0,bool),X4)))) )),
  inference(trivial_inequality_removal,[],[f454])).
fof(f454,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X4,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X4,sK0)),s(fun(sK0,bool),X4)))) )),
  inference(superposition,[],[f81,f35])).
fof(f81,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X4,sK0)),s(fun(sK0,bool),X4))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,X4,sK0)),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f47,f35])).
fof(f328,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,sK2,bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f328_D])).
fof(f328_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,sK2,bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f316,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,sK2,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | $spl12),
  inference(superposition,[],[f32,f299])).
fof(f299,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,sK2,bool,sK0)))) | $spl12),
  inference(trivial_inequality_removal,[],[f287])).
fof(f287,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,sK2,bool,sK0)))) | $spl12),
  inference(superposition,[],[f60,f281])).
fof(f60,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK3(X1,sK2,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(X1,sK2,bool,sK0))))) )),
  inference(superposition,[],[f46,f35])).
fof(f46,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK3(X1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK3(X1,sK2,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))) )),
  inference(superposition,[],[f31,f32])).
fof(f281,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK2,bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f281_D])).
fof(f281_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK2,bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f283,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK2,bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f283_D])).
fof(f283_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK2,bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f279,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,sK2,bool,sK0))))),
  inference(equality_resolution,[],[f97])).
fof(f97,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(X0,sK2,bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f95])).
fof(f95,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(X0,sK2,bool,sK0))))) )),
  inference(superposition,[],[f59,f35])).
fof(f59,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(X0,sK2,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(X0,sK2,bool,sK0))))) )),
  inference(superposition,[],[f46,f35])).
fof(f197,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK2,sK1,bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK2,sK1,bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f183,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK2,sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | $spl4),
  inference(superposition,[],[f32,f166])).
fof(f166,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK2,sK1,bool,sK0)))) | $spl4),
  inference(trivial_inequality_removal,[],[f154])).
fof(f154,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK2,sK1,bool,sK0)))) | $spl4),
  inference(superposition,[],[f64,f148])).
fof(f148,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK1,bool,sK0)))) | $spl4),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK1,bool,sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f150,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK1,bool,sK0)))) | $spl6),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK1,bool,sK0)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f146,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,sK1,bool,sK0))))),
  inference(equality_resolution,[],[f67])).
fof(f67,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,X0,bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f65])).
fof(f65,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK2,X0,bool,sK0))))) )),
  inference(superposition,[],[f57,f35])).
fof(f57,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK2,X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X0,bool,sK0))))) )),
  inference(superposition,[],[f45,f35])).
fof(f45,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK2,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK2,X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0)))) )),
  inference(superposition,[],[f31,f32])).
fof(f69,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X1,bool,sK0))))) ) | $spl0),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X1,bool,sK0))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f71,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f68,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X1,bool,sK0))))) )),
  inference(subsumption_resolution,[],[f66,f57])).
fof(f66,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK2,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK2,X1,bool,sK0))))) )),
  inference(superposition,[],[f57,f35])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 879
% Passive clauses: 10749
% Generated clauses: 23687
% Final active clauses: 529
% Final passive clauses: 6990
% Input formulas: 8
% Initial clauses: 14
% 
% Selected by SInE selection: 8
% SInE iterations: 1
% 
% Duplicate literals: 63
% Trivial inequalities: 1396
% Fw subsumption resolutions: 1450
% Fw demodulations: 1848
% Bw demodulations: 187
% 
% Simple tautologies: 268
% Forward subsumptions: 8227
% 
% Binary resolution: 878
% Forward superposition: 12817
% Backward superposition: 4716
% Equality resolution: 213
% 
% Split clauses: 138
% Split components: 286
% SAT solver clauses: 157
% SAT solver unit clauses: 1
% SAT solver binary clauses: 20
% 
% Sat splits: 138
% Sat splitting refutations: 18
% 
% Memory used [KB]: 11001
% Time elapsed: 0.840 s
% ------------------------------
% ------------------------------
% Success in time 12.548 s
