Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (144953)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 447
% Passive clauses: 6200
% Generated clauses: 39884
% Final active clauses: 354
% Final passive clauses: 5503
% Input formulas: 133
% Initial clauses: 271
% 
% Duplicate literals: 106
% Trivial inequalities: 236
% Fw subsumption resolutions: 4281
% Fw demodulations: 7038
% Bw demodulations: 89
% 
% Simple tautologies: 357
% Equational tautologies: 25
% Forward subsumptions: 21467
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 1919
% Factoring: 2
% Forward superposition: 11061
% Backward superposition: 14661
% Self superposition: 65
% Equality factoring: 14
% Equality resolution: 12
% 
% Split clauses: 159
% Split components: 319
% SAT solver clauses: 306
% SAT solver unit clauses: 9
% SAT solver binary clauses: 97
% 
% Sat splits: 159
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 303
% TWLsolver calls for satisfiability: 101
% 
% Memory used [KB]: 19061
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (144971)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 400
% Passive clauses: 3752
% Generated clauses: 30696
% Final active clauses: 179
% Final passive clauses: 1489
% Input formulas: 133
% Initial clauses: 271
% 
% Duplicate literals: 113
% Trivial inequalities: 160
% Fw subsumption resolutions: 2183
% Fw demodulations: 7552
% Bw demodulations: 1953
% Condensations: 369
% Global subsumptions: 20
% 
% Simple tautologies: 472
% Equational tautologies: 13
% Forward subsumptions: 11493
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 1277
% Factoring: 1
% Forward superposition: 5907
% Backward superposition: 4907
% Self superposition: 101
% Equality factoring: 4
% Equality resolution: 11
% 
% Split clauses: 689
% Split components: 1448
% SAT solver clauses: 4180
% SAT solver unit clauses: 331
% SAT solver binary clauses: 126
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 31
% 
% Sat splits: 689
% Sat splitting refutations: 100
% 
% TWLsolver clauses: 4092
% TWLsolver calls for satisfiability: 1899
% 
% Memory used [KB]: 18293
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (144990)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 494
% Passive clauses: 2366
% Generated clauses: 602459
% Final active clauses: 433
% Final passive clauses: 1062
% Input formulas: 133
% Initial clauses: 271
% Discarded non-redundant clauses: 1244
% 
% Duplicate literals: 8
% Trivial inequalities: 13
% Fw subsumption resolutions: 142
% Fw demodulations: 415581
% Bw demodulations: 367
% Condensations: 20
% Global subsumptions: 3
% 
% Simple tautologies: 86
% Equational tautologies: 30
% Forward subsumptions: 183957
% Fw demodulations to eq. taut.: 79
% 
% Unit resulting resolution: 177908
% Factoring: 1
% Forward superposition: 4070
% Backward superposition: 3744
% Self superposition: 74
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 2403
% SAT solver unit clauses: 1037
% SAT solver binary clauses: 957
% 
% Sat splits: 9
% Sat splitting refutations: 14
% 
% Memory used [KB]: 209932
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f30378,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f614,f615,f616,f617,f712,f715,f716,f719,f720,f623,f727,f625,f729,f734,f737,f741,f746,f752,f756,f761,f765,f769,f771,f774,f779,f785,f789,f794,f669,f670,f671,f672,f673,f674,f675,f676,f677,f684,f685,f688,f689,f696,f697,f700,f701,f702,f704,f900,f901,f904,f905,f1103,f1104,f1108,f1197,f1198,f1201,f1202,f1228,f1373,f1374,f1378,f1383,f1384,f1388,f1391,f1400,f1401,f1405,f1410,f1411,f1415,f1419,f1428,f1429,f1433,f1438,f1439,f1443,f1447,f1456,f1457,f1461,f1466,f1467,f1471,f1480,f1481,f1484,f1485,f2170,f2171,f2175,f2206,f2207,f2211,f2243,f2244,f2248,f2276,f2277,f2280,f2281,f1733,f2296,f2294,f2298,f2290,f2301,f2289,f2300,f2302,f2530,f2834,f2859,f2860,f2861,f2862,f2953,f2954,f2955,f2956,f2957,f2868,f2958,f2870,f2960,f2962,f2964,f2966,f2969,f2972,f2975,f758,f2978,f763,f2980,f2981,f2982,f2983,f2986,f2989,f2992,f2914,f2915,f2916,f2917,f2918,f2919,f2920,f2922,f2929,f2930,f2933,f2934,f2941,f2942,f2945,f2946,f2947,f2949,f3195,f3387,f3388,f3392,f1134,f3520,f1732,f3521,f1735,f3524,f1739,f3525,f1734,f1741,f3990,f3991,f3995,f4015,f4016,f4020,f4041,f4042,f4046,f4067,f4068,f4072,f4077,f4078,f1475,f4093,f4094,f4098,f4103,f4104,f4115,f4120,f4121,f4124,f4125,f4704,f4705,f4709,f4671,f4670,f2921,f811,f4735,f791,f5565,f5567,f5675,f5676,f790,f6009,f5918,f6010,f5920,f6012,f6014,f6016,f6018,f6021,f6024,f749,f6026,f6027,f6028,f6029,f6032,f6035,f6038,f782,f5964,f800,f5965,f5966,f5967,f5968,f5969,f5970,f5972,f5979,f5980,f5983,f5984,f5991,f5992,f5995,f5996,f6051,f6052,f6056,f6086,f6087,f6091,f6122,f6123,f6127,f6158,f6159,f6163,f6194,f6195,f6199,f6230,f6231,f6234,f6235,f6700,f6701,f6738,f6739,f6776,f6777,f6811,f6812,f6847,f6848,f6999,f7003,f7001,f7018,f7019,f7054,f7055,f7093,f7094,f7130,f7131,f7238,f7241,f7240,f7189,f7190,f7242,f7191,f7192,f7243,f7193,f7244,f7246,f6996,f6994,f6992,f6990,f5999,f832,f5997,f831,f7258,f7259,f714,f5912,f5911,f711,f708,f5910,f5909,f710,f709,f7227,f7261,f7226,f7265,f7223,f7266,f7222,f7267,f7215,f7263,f7269,f7212,f7211,f7210,f7209,f7208,f7564,f7569,f7572,f7571,f7534,f7535,f7536,f7537,f7538,f7671,f7519,f7678,f7518,f7517,f7679,f7516,f7515,f7680,f798,f7889,f7890,f7914,f7915,f8240,f7565,f8244,f8245,f7570,f7983,f7984,f7985,f7986,f7987,f7998,f7999,f8282,f8000,f8001,f8261,f8002,f8003,f8283,f8004,f8005,f8262,f8006,f8263,f8007,f8284,f8013,f8287,f8288,f8289,f8290,f8292,f8018,f8293,f8557,f8558,f8583,f8584,f8609,f8610,f8638,f8639,f8666,f8667,f8694,f8695,f8696,f8789,f8697,f8791,f8792,f8793,f8794,f739,f8795,f8797,f747,f8801,f8799,f8803,f8806,f8807,f773,f8808,f8809,f8812,f8747,f8816,f8749,f8756,f8757,f8760,f8761,f8805,f9340,f9431,f9444,f753,f9487,f9554,f9595,f9697,f9739,f10084,f10403,f10405,f10945,f10948,f10949,f10947,f10956,f11153,f11064,f11154,f11066,f11156,f11158,f11160,f11162,f11163,f11164,f11116,f809,f11125,f11126,f11130,f11137,f11138,f11141,f11142,f11324,f11325,f11326,f11327,f11328,f11547,f11548,f11549,f11550,f11551,f11590,f11696,f11698,f12611,f12614,f12615,f12616,f12620,f12622,f12625,f12619,f12643,f12787,f12645,f806,f13010,f13011,f13012,f12612,f13013,f13014,f13015,f13016,f12618,f13003,f13018,f13004,f13019,f13008,f807,f13021,f13120,f13232,f13231,f13230,f13229,f13227,f13209,f13208,f13207,f13205,f13196,f13195,f13194,f13283,f13284,f13285,f13286,f13287,f13288,f13289,f13290,f13310,f13324,f13325,f13329,f13358,f13359,f13380,f13122,f13121,f13502,f13501,f13494,f13479,f13478,f13477,f13476,f13475,f13474,f13472,f13471,f13470,f13448,f13556,f13557,f13558,f13567,f13568,f13829,f13833,f13837,f13840,f13832,f14337,f15001,f14922,f15010,f15005,f14924,f15009,f15012,f766,f15015,f15016,f776,f15019,f15024,f14966,f15032,f14973,f15036,f14974,f15037,f14977,f15039,f14978,f15014,f15041,f14985,f15042,f14986,f15045,f14989,f15047,f14990,f15044,f15049,f15245,f15246,f15270,f15271,f15296,f15297,f15318,f15319,f15322,f15323,f15796,f15797,f15822,f15823,f15851,f15852,f15878,f15879,f16019,f16020,f16054,f16055,f16090,f16091,f16126,f16127,f16162,f16163,f16283,f16284,f16320,f16321,f16863,f16864,f16889,f16890,f17001,f17002,f17029,f17030,f17058,f730,f17059,f17094,f17185,f17202,f735,f17245,f17313,f17728,f17730,f18591,f18594,f18595,f18596,f18593,f18851,f18966,f18983,f19083,f19085,f19137,f19139,f19434,f19168,f19436,f19256,f19563,f19260,f19564,f8811,f19383,f19645,f19384,f19647,f19385,f19649,f8810,f20132,f20133,f20134,f20205,f20135,f19432,f20164,f20211,f20278,f20280,f20362,f20373,f20693,f20694,f20695,f15003,f20696,f20209,f20767,f20769,f20361,f15007,f21429,f21438,f15021,f21439,f15020,f21441,f12613,f21442,f21443,f18334,f12627,f18333,f12624,f18330,f12621,f18329,f21445,f12617,f10893,f10894,f21446,f10895,f10896,f21447,f10897,f21448,f14572,f11111,f11112,f11113,f11114,f11115,f21874,f21785,f21875,f718,f21787,f21877,f732,f743,f21879,f21832,f21833,f803,f21834,f21835,f804,f21836,f805,f802,f21839,f813,f21846,f21847,f21850,f820,f21851,f21858,f823,f21859,f826,f21862,f828,f21863,f825,f830,f22093,f22094,f22098,f1390,f22118,f22119,f22123,f1418,f22144,f22145,f22149,f1446,f22166,f13830,f22167,f13834,f22170,f13839,f22171,f15575,f13842,f22646,f22647,f22651,f1474,f22672,f22673,f22677,f4114,f22701,f22702,f22706,f8657,f22728,f22729,f22733,f8685,f796,f22856,f4372,f22857,f4375,f22860,f4379,f22861,f4374,f4381,f23336,f23337,f2191,f23371,f23372,f2228,f23407,f23408,f2265,f23443,f23444,f4726,f23479,f23480,f6219,f23518,f23519,f6723,f23555,f23556,f6761,f23666,f23667,f16870,f23812,f23813,f16896,f23837,f23838,f17010,f23864,f23865,f17039,f23892,f23893,f17066,f24685,f24686,f7274,f24687,f24688,f7273,f24689,f7272,f24800,f7566,f780,f24774,f24775,f7579,f24776,f24777,f7580,f24778,f7581,f7271,f24837,f24838,f24880,f24839,f24840,f24881,f24841,f24882,f24918,f24919,f24963,f24920,f24921,f24964,f24922,f24965,f24879,f25008,f9685,f25113,f9445,f25309,f25325,f786,f25363,f9806,f9429,f26130,f26132,f26713,f26715,f26797,f26808,f26825,f26827,f26796,f27740,f27768,f28093,f28095,f726,f724,f28185,f770,f28154,f817,f28155,f818,f28159,f768,f822,f28294,f28296,f11969,f30240,f30242,f20207,f722,f30374])).
fof(f30374,plain,(
  $false | ($spl7 | $spl56)),
  inference(subsumption_resolution,[],[f30371,f797])).
fof(f797,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl56),
  inference(cnf_transformation,[],[f797_D])).
fof(f797_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f30371,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl7),
  inference(trivial_inequality_removal,[],[f30320])).
fof(f30320,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl7),
  inference(superposition,[],[f714,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))))),
  inference(ennf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) => ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))) => ! [X11] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X9),s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X10))),s('type.num.num',X9))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_SUB')).
fof(f722,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl10),
  inference(cnf_transformation,[],[f722_D])).
fof(f722_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f20207,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl369),
  inference(cnf_transformation,[],[f20207_D])).
fof(f20207_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f30242,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl245),
  inference(superposition,[],[f8816,f476])).
fof(f476,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.SUB_LEFT_SUB',aHL_BOOL_CASES)).
fof(f30240,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl245),
  inference(superposition,[],[f8816,f475])).
fof(f475,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) | s(bool,'F_0') = s(bool,X17))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.BOOL_CASES_AX')).
fof(f11969,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl260),
  inference(cnf_transformation,[],[f11969_D])).
fof(f11969_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f28296,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl240),
  inference(superposition,[],[f8804,f476])).
fof(f8804,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl240),
  inference(cnf_transformation,[],[f8804_D])).
fof(f8804_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f28294,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl240),
  inference(superposition,[],[f8804,f475])).
fof(f822,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) | $spl82),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f768,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl39),
  inference(cnf_transformation,[],[f768_D])).
fof(f768_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f28159,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))),
  inference(superposition,[],[f391,f474])).
fof(f474,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f189])).
fof(f189,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | ? [X7] : s('type.num.num',X6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.num_CASES')).
fof(f391,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6])],[f264])).
fof(f264,plain,(
  ? [X0,X1,X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X2))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X8))),s('type.num.num',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X8))),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_LEFT_SUB')).
fof(f818,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) ) | $spl78),
  inference(cnf_transformation,[],[f818_D])).
fof(f818_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f28155,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f391,f429])).
fof(f429,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X6,X7] : (s('type.num.num',X6) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_EQ_EQ_|48|')).
fof(f817,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) ) | $spl76),
  inference(cnf_transformation,[],[f817_D])).
fof(f817_D,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f28154,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f770,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) ) | $spl40),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f28185,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(forward_demodulation,[],[f28124,f480])).
fof(f480,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB|split|0')).
fof(f28124,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(superposition,[],[f391,f429])).
fof(f724,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl12),
  inference(cnf_transformation,[],[f724_D])).
fof(f724_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f726,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | $spl15),
  inference(cnf_transformation,[],[f726_D])).
fof(f726_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f28095,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f28093,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f27768,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f27767,f834])).
fof(f834,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl1),
  inference(resolution,[],[f708,f512])).
fof(f512,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.NOT_LESS')).
fof(f27767,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl12),
  inference(subsumption_resolution,[],[f27483,f609])).
fof(f609,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f461])).
fof(f461,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X12] : ! [X18] : (s(X12,X18) = s(X12,X18) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.REFL_CLAUSE')).
fof(f27483,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl12),
  inference(superposition,[],[f513,f724])).
fof(f513,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f27740,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f27434,f600])).
fof(f600,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.SUB_LEFT_SUB',aHL_TRUTH)).
fof(f27434,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl12)),
  inference(backward_demodulation,[],[f724,f708])).
fof(f26796,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | $spl400),
  inference(cnf_transformation,[],[f26796_D])).
fof(f26796_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f26827,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl18),
  inference(superposition,[],[f731,f476])).
fof(f731,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | $spl18),
  inference(cnf_transformation,[],[f731_D])).
fof(f731_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f26825,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl18),
  inference(superposition,[],[f731,f475])).
fof(f26808,plain,(
  $false | ($spl19 | $spl56)),
  inference(subsumption_resolution,[],[f26807,f797])).
fof(f26807,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl19),
  inference(subsumption_resolution,[],[f26757,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X12] : ! [X9,X17] : s(X12,X17) = s(X12,'const.bool.COND_3'(s(bool,X9),s(X12,X17),s(X12,X17)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.COND_ID')).
fof(f26757,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl19),
  inference(superposition,[],[f732,f433])).
fof(f433,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_EQ_|48|')).
fof(f26797,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | $spl401),
  inference(cnf_transformation,[],[f26797_D])).
fof(f26797_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) <=> ~$spl401),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl401])])).
fof(f26715,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl19),
  inference(superposition,[],[f732,f476])).
fof(f26713,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl19),
  inference(superposition,[],[f732,f475])).
fof(f26132,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl240),
  inference(superposition,[],[f8804,f476])).
fof(f26130,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl240),
  inference(superposition,[],[f8804,f475])).
fof(f9429,plain,(
  ( ! [X62] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl246),
  inference(cnf_transformation,[],[f9429_D])).
fof(f9429_D,plain,(
  ( ! [X62] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f9806,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl257),
  inference(cnf_transformation,[],[f9806_D])).
fof(f9806_D,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f25363,plain,(
  ( ! [X62] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl50),
  inference(superposition,[],[f786,f602])).
fof(f602,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X0)))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X11)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUC_SUB')).
fof(f786,plain,(
  ( ! [X13] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | $spl50),
  inference(cnf_transformation,[],[f786_D])).
fof(f786_D,plain,(
  ( ! [X13] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f25325,plain,(
  ( ! [X24,X23] : (sP3(X23,X24,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) ) | $spl398),
  inference(cnf_transformation,[],[f25325_D])).
fof(f25325_D,plain,(
  ( ! [X24,X23] : (sP3(X23,X24,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f25309,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f1050,f429])).
fof(f1050,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) | $spl1),
  inference(resolution,[],[f833,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f310])).
fof(f310,plain,(
  ! [X3,X1,X2] : ((~sP3(X3,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP3(X3,X1,X2)))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X3,X1,X2] : ((~sP3(X3,X1,X2) | (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3)))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP3(X3,X1,X2)))),
  inference(nnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X3,X1,X2] : (sP3(X3,X1,X2) <=> (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f833,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl1),
  inference(resolution,[],[f708,f516])).
fof(f516,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.LESS_IMP_LESS_OR_EQ')).
fof(f9445,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) ) | $spl250),
  inference(cnf_transformation,[],[f9445_D])).
fof(f9445_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f25113,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X19)))) ) | $spl46),
  inference(subsumption_resolution,[],[f25056,f479])).
fof(f479,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_|48||split|1')).
fof(f25056,plain,(
  ( ! [X19,X18] : (s('type.num.num',X18) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X19)))) ) | $spl46),
  inference(superposition,[],[f780,f429])).
fof(f9685,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK5))) | $spl255),
  inference(cnf_transformation,[],[f9685_D])).
fof(f9685_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK5))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f25008,plain,(
  ( ! [X62] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl46),
  inference(superposition,[],[f780,f602])).
fof(f24879,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | $spl385),
  inference(cnf_transformation,[],[f24879_D])).
fof(f24879_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) <=> ~$spl385),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl385])])).
fof(f24965,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) | $spl396),
  inference(cnf_transformation,[],[f24965_D])).
fof(f24965_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f24922,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) ) | $spl323),
  inference(superposition,[],[f15044,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ((s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))) & (s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))) | s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | (s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) => (s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) => (s('type.num.num',X8) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.ADD_EQ_SUB')).
fof(f24964,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2)))) | s('type.num.num',sK4) != s('type.num.num',X2)) ) | $spl394),
  inference(cnf_transformation,[],[f24964_D])).
fof(f24964_D,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2)))) | s('type.num.num',sK4) != s('type.num.num',X2)) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f24921,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) ) | $spl323),
  inference(superposition,[],[f15044,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) | s('type.num.num',X1) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ((s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) | s('type.num.num',X1) = s('type.num.num',X2)) & (s('type.num.num',X1) != s('type.num.num',X2) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2))))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X8,X7,X6] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))) <=> s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.CANCEL_SUB')).
fof(f24920,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2))))) ) | $spl323),
  inference(superposition,[],[f15044,f419])).
fof(f24963,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',sK6) != s('type.num.num',X0)) ) | $spl392),
  inference(cnf_transformation,[],[f24963_D])).
fof(f24963_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',sK6) != s('type.num.num',X0)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f24919,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | s('type.num.num',sK6) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) ) | $spl323),
  inference(superposition,[],[f15044,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) | s('type.num.num',X1) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ((s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) | s('type.num.num',X1) = s('type.num.num',X2)) & (s('type.num.num',X1) != s('type.num.num',X2) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))))))),
  inference(nnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X8,X7,X6] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X8))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_CANCEL')).
fof(f24918,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))))) | s('type.num.num',sK6) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl323),
  inference(superposition,[],[f15044,f417])).
fof(f24882,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) | $spl390),
  inference(cnf_transformation,[],[f24882_D])).
fof(f24882_D,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f24841,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) ) | $spl301),
  inference(superposition,[],[f15005,f420])).
fof(f24881,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2)))) | s('type.num.num',sK4) != s('type.num.num',X2)) ) | $spl388),
  inference(cnf_transformation,[],[f24881_D])).
fof(f24881_D,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2)))) | s('type.num.num',sK4) != s('type.num.num',X2)) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f24840,plain,(
  ( ! [X3] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) ) | $spl301),
  inference(superposition,[],[f15005,f419])).
fof(f24839,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2))))) ) | $spl301),
  inference(superposition,[],[f15005,f419])).
fof(f24880,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',sK6) != s('type.num.num',X0)) ) | $spl386),
  inference(cnf_transformation,[],[f24880_D])).
fof(f24880_D,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',sK6) != s('type.num.num',X0)) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f24838,plain,(
  ( ! [X1] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))))) | s('type.num.num',sK6) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) ) | $spl301),
  inference(superposition,[],[f15005,f417])).
fof(f24837,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))))) | s('type.num.num',sK6) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl301),
  inference(superposition,[],[f15005,f417])).
fof(f7271,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | $spl187),
  inference(cnf_transformation,[],[f7271_D])).
fof(f7271_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f7581,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl202),
  inference(cnf_transformation,[],[f7581_D])).
fof(f7581_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f24778,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f420])).
fof(f7580,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2)))) | s('type.num.num',sK4) != s('type.num.num',X2)) ) | $spl200),
  inference(cnf_transformation,[],[f7580_D])).
fof(f7580_D,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2)))) | s('type.num.num',sK4) != s('type.num.num',X2)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f24777,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f419])).
fof(f24776,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2))))) ) | $spl87),
  inference(superposition,[],[f825,f419])).
fof(f7579,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) | $spl198),
  inference(cnf_transformation,[],[f7579_D])).
fof(f7579_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f24775,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f417])).
fof(f24774,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f417])).
fof(f780,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) | $spl46),
  inference(cnf_transformation,[],[f780_D])).
fof(f780_D,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7566,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl195),
  inference(cnf_transformation,[],[f7566_D])).
fof(f7566_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f24800,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl87),
  inference(forward_demodulation,[],[f24739,f480])).
fof(f24739,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl87),
  inference(superposition,[],[f825,f429])).
fof(f7272,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl188),
  inference(cnf_transformation,[],[f7272_D])).
fof(f7272_D,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f24689,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f420])).
fof(f7273,plain,(
  ( ! [X3] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | s('type.num.num',sK4) != s('type.num.num',X3)) ) | $spl190),
  inference(cnf_transformation,[],[f7273_D])).
fof(f7273_D,plain,(
  ( ! [X3] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | s('type.num.num',sK4) != s('type.num.num',X3)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f24688,plain,(
  ( ! [X3] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f419])).
fof(f24687,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2))))) ) | $spl9),
  inference(superposition,[],[f718,f419])).
fof(f7274,plain,(
  ( ! [X1] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) | $spl192),
  inference(cnf_transformation,[],[f7274_D])).
fof(f7274_D,plain,(
  ( ! [X1] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f24686,plain,(
  ( ! [X1] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f417])).
fof(f24685,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f417])).
fof(f17066,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl340),
  inference(cnf_transformation,[],[f17066_D])).
fof(f17066_D,plain,(
  ( ! [X21,X20] : (sP0(X20,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X21)) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f23893,plain,(
  ( ! [X24,X23,X25] : (sP0(X23,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl56),
  inference(superposition,[],[f7709,f429])).
fof(f7709,plain,(
  ( ! [X37,X36] : (sP0(X36,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)),X37)) ) | $spl56),
  inference(resolution,[],[f797,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f319])).
fof(f319,plain,(
  ! [X0,X2,X3] : ((~sP0(X0,X2,X3) | p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP0(X0,X2,X3)))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0,X2,X3] : ((~sP0(X0,X2,X3) | (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP0(X0,X2,X3)))),
  inference(nnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X2,X3] : (sP0(X0,X2,X3) <=> (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f23892,plain,(
  ( ! [X21,X22,X20] : (sP0(X20,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl56),
  inference(superposition,[],[f7709,f429])).
fof(f17039,plain,(
  ( ! [X21,X20] : (sP0(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) ) | $spl338),
  inference(cnf_transformation,[],[f17039_D])).
fof(f17039_D,plain,(
  ( ! [X21,X20] : (sP0(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f23865,plain,(
  ( ! [X24,X23,X25] : (sP0(X23,X24,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl56),
  inference(superposition,[],[f7708,f429])).
fof(f7708,plain,(
  ( ! [X35,X34] : (sP0(X34,X35,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) ) | $spl56),
  inference(resolution,[],[f797,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f320])).
fof(f23864,plain,(
  ( ! [X21,X22,X20] : (sP0(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl56),
  inference(superposition,[],[f7708,f429])).
fof(f17010,plain,(
  ( ! [X21,X20] : (sP1(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) ) | $spl336),
  inference(cnf_transformation,[],[f17010_D])).
fof(f17010_D,plain,(
  ( ! [X21,X20] : (sP1(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f23838,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,X24,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl56),
  inference(superposition,[],[f7706,f429])).
fof(f7706,plain,(
  ( ! [X30,X31] : (sP1(X30,X31,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) ) | $spl56),
  inference(resolution,[],[f797,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f316])).
fof(f316,plain,(
  ! [X2,X1,X3] : ((~sP1(X2,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP1(X2,X1,X3)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X2,X1,X3] : ((~sP1(X2,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP1(X2,X1,X3)))),
  inference(nnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X2,X1,X3] : (sP1(X2,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f23837,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl56),
  inference(superposition,[],[f7706,f429])).
fof(f16896,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) ) | $spl334),
  inference(cnf_transformation,[],[f16896_D])).
fof(f16896_D,plain,(
  ( ! [X21,X20] : (sP2(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f23813,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,X24,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl56),
  inference(superposition,[],[f7703,f429])).
fof(f7703,plain,(
  ( ! [X24,X25] : (sP2(X24,X25,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) ) | $spl56),
  inference(resolution,[],[f797,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X3] : ((~sP2(X0,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X3)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X3] : ((~sP2(X0,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X3)))),
  inference(nnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X3] : (sP2(X0,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f23812,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl56),
  inference(superposition,[],[f7703,f429])).
fof(f16870,plain,(
  ( ! [X21,X20] : (sP3(X20,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl332),
  inference(cnf_transformation,[],[f16870_D])).
fof(f16870_D,plain,(
  ( ! [X21,X20] : (sP3(X20,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X21)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f23667,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl56),
  inference(superposition,[],[f7700,f429])).
fof(f7700,plain,(
  ( ! [X19,X18] : (sP3(X18,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)),X19)) ) | $spl56),
  inference(resolution,[],[f797,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f311])).
fof(f23666,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl56),
  inference(superposition,[],[f7700,f429])).
fof(f6761,plain,(
  ( ! [X21,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl162),
  inference(cnf_transformation,[],[f6761_D])).
fof(f6761_D,plain,(
  ( ! [X21,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f23556,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f943,f429])).
fof(f943,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | $spl53),
  inference(resolution,[],[f791,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f320])).
fof(f23555,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f943,f429])).
fof(f6723,plain,(
  ( ! [X21,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl160),
  inference(cnf_transformation,[],[f6723_D])).
fof(f6723_D,plain,(
  ( ! [X21,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f23519,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f940,f429])).
fof(f940,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X16,X17)) ) | $spl53),
  inference(resolution,[],[f791,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f317])).
fof(f23518,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f940,f429])).
fof(f6219,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl158),
  inference(cnf_transformation,[],[f6219_D])).
fof(f6219_D,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f23480,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f939,f429])).
fof(f939,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X15)) ) | $spl53),
  inference(resolution,[],[f791,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f317])).
fof(f23479,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f939,f429])).
fof(f4726,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl156),
  inference(cnf_transformation,[],[f4726_D])).
fof(f4726_D,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f23444,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f937,f429])).
fof(f937,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl53),
  inference(resolution,[],[f791,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f23443,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f937,f429])).
fof(f2265,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl132),
  inference(cnf_transformation,[],[f2265_D])).
fof(f2265_D,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f23408,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f936,f429])).
fof(f936,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X9)) ) | $spl53),
  inference(resolution,[],[f791,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f314])).
fof(f23407,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f936,f429])).
fof(f2228,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl130),
  inference(cnf_transformation,[],[f2228_D])).
fof(f2228_D,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f23372,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f934,f429])).
fof(f934,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X4,X5)) ) | $spl53),
  inference(resolution,[],[f791,f395])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f311])).
fof(f23371,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f934,f429])).
fof(f2191,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl128),
  inference(cnf_transformation,[],[f2191_D])).
fof(f2191_D,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f23337,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f933,f429])).
fof(f933,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) ) | $spl53),
  inference(resolution,[],[f791,f393])).
fof(f23336,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f933,f429])).
fof(f4381,plain,(
  s('type.num.num',sK8(sK5)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK5))))) | $spl154),
  inference(cnf_transformation,[],[f4381_D])).
fof(f4381_D,plain,(
  s('type.num.num',sK8(sK5)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK5))))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f4374,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl148),
  inference(cnf_transformation,[],[f4374_D])).
fof(f4374_D,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f22861,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK5)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK5))))) | $spl54),
  inference(superposition,[],[f796,f474])).
fof(f4379,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK5)),s('type.num.num','const.num.|48|_0')))) | $spl153),
  inference(cnf_transformation,[],[f4379_D])).
fof(f4379_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK5)),s('type.num.num','const.num.|48|_0')))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f22860,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK5)),s('type.num.num','const.num.|48|_0')))) | $spl54),
  inference(superposition,[],[f796,f464])).
fof(f464,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.LESS_EQ_|48|')).
fof(f4375,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK5)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK5)),s('type.num.num',X13)))) ) | $spl150),
  inference(cnf_transformation,[],[f4375_D])).
fof(f4375_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK5)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK5)),s('type.num.num',X13)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f22857,plain,(
  ( ! [X13] : (s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK5)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK5)),s('type.num.num',X13)))) ) | $spl54),
  inference(superposition,[],[f796,f429])).
fof(f4372,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK5))))) ) | $spl146),
  inference(cnf_transformation,[],[f4372_D])).
fof(f4372_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK5))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f22856,plain,(
  ( ! [X12] : (s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK5))))) ) | $spl54),
  inference(superposition,[],[f796,f429])).
fof(f796,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5)))) | $spl54),
  inference(cnf_transformation,[],[f796_D])).
fof(f796_D,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8685,plain,(
  ( ! [X21,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) | $spl232),
  inference(cnf_transformation,[],[f8685_D])).
fof(f8685_D,plain,(
  ( ! [X21,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f22733,plain,(
  ( ! [X35,X34] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f847,f474])).
fof(f847,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X22,X23)) ) | $spl1),
  inference(resolution,[],[f708,f407])).
fof(f22729,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f847,f429])).
fof(f22728,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f847,f429])).
fof(f8657,plain,(
  ( ! [X21,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) | $spl230),
  inference(cnf_transformation,[],[f8657_D])).
fof(f8657_D,plain,(
  ( ! [X21,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f22706,plain,(
  ( ! [X35,X34] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f844,f474])).
fof(f844,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X16,X17)) ) | $spl1),
  inference(resolution,[],[f708,f403])).
fof(f22702,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f844,f429])).
fof(f22701,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f844,f429])).
fof(f4114,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21)) ) | $spl142),
  inference(cnf_transformation,[],[f4114_D])).
fof(f4114_D,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f22677,plain,(
  ( ! [X35,X34] : (sP1(X34,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f843,f474])).
fof(f843,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X15)) ) | $spl1),
  inference(resolution,[],[f708,f402])).
fof(f22673,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f843,f429])).
fof(f22672,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f843,f429])).
fof(f1474,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) | $spl114),
  inference(cnf_transformation,[],[f1474_D])).
fof(f1474_D,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f22651,plain,(
  ( ! [X35,X34] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f841,f474])).
fof(f841,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X10,X11)) ) | $spl1),
  inference(resolution,[],[f708,f399])).
fof(f22647,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f22646,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f13842,plain,(
  s('type.num.num',sK8(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK4))))) | $spl296),
  inference(cnf_transformation,[],[f13842_D])).
fof(f13842_D,plain,(
  s('type.num.num',sK8(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK4))))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f15575,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl330),
  inference(cnf_transformation,[],[f15575_D])).
fof(f15575_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f22171,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK4))))) | $spl24),
  inference(superposition,[],[f743,f474])).
fof(f13839,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK4)),s('type.num.num','const.num.|48|_0')))) | $spl295),
  inference(cnf_transformation,[],[f13839_D])).
fof(f13839_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK4)),s('type.num.num','const.num.|48|_0')))) <=> ~$spl295),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl295])])).
fof(f22170,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK4)),s('type.num.num','const.num.|48|_0')))) | $spl24),
  inference(superposition,[],[f743,f464])).
fof(f13834,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK4)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK4)),s('type.num.num',X13)))) ) | $spl292),
  inference(cnf_transformation,[],[f13834_D])).
fof(f13834_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK4)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK4)),s('type.num.num',X13)))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f22167,plain,(
  ( ! [X13] : (s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK4)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK4)),s('type.num.num',X13)))) ) | $spl24),
  inference(superposition,[],[f743,f429])).
fof(f13830,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK4))))) ) | $spl288),
  inference(cnf_transformation,[],[f13830_D])).
fof(f13830_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK4))))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f22166,plain,(
  ( ! [X12] : (s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK4))))) ) | $spl24),
  inference(superposition,[],[f743,f429])).
fof(f1446,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21)) ) | $spl110),
  inference(cnf_transformation,[],[f1446_D])).
fof(f1446_D,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f22149,plain,(
  ( ! [X35,X34] : (sP2(X34,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f840,f474])).
fof(f840,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X9)) ) | $spl1),
  inference(resolution,[],[f708,f398])).
fof(f22145,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f22144,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f1418,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) | $spl106),
  inference(cnf_transformation,[],[f1418_D])).
fof(f1418_D,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f22123,plain,(
  ( ! [X35,X34] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f838,f474])).
fof(f838,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X4,X5)) ) | $spl1),
  inference(resolution,[],[f708,f395])).
fof(f22119,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f22118,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f1390,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) ) | $spl102),
  inference(cnf_transformation,[],[f1390_D])).
fof(f1390_D,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f22098,plain,(
  ( ! [X35,X34] : (sP3(X34,X35,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f837,f474])).
fof(f837,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) | $spl1),
  inference(resolution,[],[f708,f393])).
fof(f22094,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f22093,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f830,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) | $spl92),
  inference(cnf_transformation,[],[f830_D])).
fof(f830_D,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f825,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl87),
  inference(cnf_transformation,[],[f825_D])).
fof(f825_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f21863,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))))),
  inference(superposition,[],[f391,f474])).
fof(f828,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) | $spl91),
  inference(cnf_transformation,[],[f828_D])).
fof(f828_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f21862,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f826,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num',X13)))) ) | $spl88),
  inference(cnf_transformation,[],[f826_D])).
fof(f826_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num',X13)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f21859,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f823,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) ) | $spl84),
  inference(cnf_transformation,[],[f823_D])).
fof(f823_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f21858,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) )),
  inference(superposition,[],[f391,f429])).
fof(f21851,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))),
  inference(superposition,[],[f391,f474])).
fof(f820,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl81),
  inference(cnf_transformation,[],[f820_D])).
fof(f820_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f21850,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f21847,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f391,f429])).
fof(f21846,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f813,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) | $spl75),
  inference(cnf_transformation,[],[f813_D])).
fof(f813_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f21839,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),
  inference(superposition,[],[f391,f433])).
fof(f802,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl61),
  inference(cnf_transformation,[],[f802_D])).
fof(f802_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f805,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5)))) ) | $spl66),
  inference(cnf_transformation,[],[f805_D])).
fof(f805_D,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f21836,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f420])).
fof(f804,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2)))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) ) | $spl64),
  inference(cnf_transformation,[],[f804_D])).
fof(f804_D,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2)))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f21835,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f419])).
fof(f21834,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2))))) )),
  inference(superposition,[],[f391,f419])).
fof(f803,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | s('type.num.num',sK5) != s('type.num.num',X0)) ) | $spl62),
  inference(cnf_transformation,[],[f803_D])).
fof(f803_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | s('type.num.num',sK5) != s('type.num.num',X0)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f21833,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f21832,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))))) | s('type.num.num',sK5) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f21879,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(forward_demodulation,[],[f21816,f480])).
fof(f21816,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(superposition,[],[f391,f429])).
fof(f743,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl24),
  inference(cnf_transformation,[],[f743_D])).
fof(f743_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f732,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | $spl19),
  inference(cnf_transformation,[],[f732_D])).
fof(f732_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f21877,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f21876,f480])).
fof(f21876,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f21797,f478])).
fof(f478,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f21797,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(superposition,[],[f391,f474])).
fof(f21787,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f718,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl9),
  inference(cnf_transformation,[],[f718_D])).
fof(f718_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f21875,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f21786,f460])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X12] : ! [X20,X21] : s(X12,X20) = s(X12,'const.bool.COND_3'(s(bool,'T_0'),s(X12,X20),s(X12,X21)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.COND_CLAUSES|split|0')).
fof(f21786,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f21785,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f21874,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f21784,f460])).
fof(f21784,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f11115,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f420])).
fof(f11114,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f419])).
fof(f11113,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2))))) )),
  inference(superposition,[],[f391,f419])).
fof(f11112,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f11111,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))))) | s('type.num.num',sK5) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f14572,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl81),
  inference(forward_demodulation,[],[f14571,f478])).
fof(f14571,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl81),
  inference(forward_demodulation,[],[f10927,f449])).
fof(f449,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.ADD_SYM')).
fof(f10927,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl81),
  inference(superposition,[],[f820,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ! [X2] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))))),
  inference(ennf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) => ! [X2] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X6,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X8)))) => ! [X7] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num',X7))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_LESS_EQ_ADD')).
fof(f21448,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5)))) ) | $spl382),
  inference(cnf_transformation,[],[f21448_D])).
fof(f21448_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5)))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f10897,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | $spl81),
  inference(superposition,[],[f820,f420])).
fof(f21447,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2)))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) ) | $spl380),
  inference(cnf_transformation,[],[f21447_D])).
fof(f21447_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2)))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f10896,plain,(
  ( ! [X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | $spl81),
  inference(superposition,[],[f820,f419])).
fof(f10895,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2))))) ) | $spl81),
  inference(superposition,[],[f820,f419])).
fof(f21446,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | s('type.num.num',sK5) != s('type.num.num',X0)) ) | $spl378),
  inference(cnf_transformation,[],[f21446_D])).
fof(f21446_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | s('type.num.num',sK5) != s('type.num.num',X0)) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f10894,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | $spl81),
  inference(superposition,[],[f820,f417])).
fof(f10893,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK5) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | $spl81),
  inference(superposition,[],[f820,f417])).
fof(f12617,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | $spl264),
  inference(cnf_transformation,[],[f12617_D])).
fof(f12617_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f21445,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl376),
  inference(cnf_transformation,[],[f21445_D])).
fof(f21445_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f18329,plain,(
  ( ! [X12] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | $spl44),
  inference(superposition,[],[f776,f429])).
fof(f12621,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num',X13)))) ) | $spl268),
  inference(cnf_transformation,[],[f12621_D])).
fof(f12621_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num',X13)))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f18330,plain,(
  ( ! [X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num',X13)))) ) | $spl44),
  inference(superposition,[],[f776,f429])).
fof(f12624,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num','const.num.|48|_0')))) | $spl271),
  inference(cnf_transformation,[],[f12624_D])).
fof(f12624_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl271])])).
fof(f18333,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num','const.num.|48|_0')))) | $spl44),
  inference(superposition,[],[f776,f464])).
fof(f12627,plain,(
  s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) | $spl272),
  inference(cnf_transformation,[],[f12627_D])).
fof(f12627_D,plain,(
  s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f18334,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) | $spl44),
  inference(superposition,[],[f776,f474])).
fof(f21443,plain,(
  ( ! [X25] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))) ) | $spl16),
  inference(subsumption_resolution,[],[f17145,f479])).
fof(f17145,plain,(
  ( ! [X25] : (s('type.num.num',X25) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))) ) | $spl16),
  inference(superposition,[],[f730,f474])).
fof(f21442,plain,(
  ( ! [X25] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))) ) | $spl20),
  inference(subsumption_resolution,[],[f17217,f480])).
fof(f17217,plain,(
  ( ! [X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))) ) | $spl20),
  inference(superposition,[],[f735,f474])).
fof(f12613,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | $spl262),
  inference(cnf_transformation,[],[f12613_D])).
fof(f12613_D,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f21441,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl44),
  inference(forward_demodulation,[],[f18312,f478])).
fof(f18312,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl44),
  inference(superposition,[],[f776,f474])).
fof(f15020,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl308),
  inference(cnf_transformation,[],[f15020_D])).
fof(f15020_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f21439,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl48),
  inference(forward_demodulation,[],[f19149,f478])).
fof(f19149,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl48),
  inference(superposition,[],[f781,f474])).
fof(f781,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl48),
  inference(cnf_transformation,[],[f781_D])).
fof(f781_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f15021,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl309),
  inference(cnf_transformation,[],[f15021_D])).
fof(f15021_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) <=> ~$spl309),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl309])])).
fof(f21438,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl243),
  inference(forward_demodulation,[],[f20145,f478])).
fof(f20145,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl243),
  inference(superposition,[],[f8811,f474])).
fof(f21429,plain,(
  $false | ($spl7 | $spl56 | $spl290)),
  inference(subsumption_resolution,[],[f21426,f14062])).
fof(f14062,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ($spl56 | $spl290)),
  inference(backward_demodulation,[],[f13832,f797])).
fof(f21426,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ($spl7 | $spl290)),
  inference(trivial_inequality_removal,[],[f21385])).
fof(f21385,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ($spl7 | $spl290)),
  inference(superposition,[],[f21371,f441])).
fof(f21371,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl7 | $spl290)),
  inference(forward_demodulation,[],[f714,f13832])).
fof(f15007,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl302),
  inference(cnf_transformation,[],[f15007_D])).
fof(f15007_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f20361,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) | $spl374),
  inference(cnf_transformation,[],[f20361_D])).
fof(f20361_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f20769,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl18 | $spl290)),
  inference(superposition,[],[f20381,f476])).
fof(f20381,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) | ($spl18 | $spl290)),
  inference(forward_demodulation,[],[f731,f13832])).
fof(f20767,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl18 | $spl290)),
  inference(superposition,[],[f20381,f475])).
fof(f20209,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl371),
  inference(cnf_transformation,[],[f20209_D])).
fof(f20209_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl371),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl371])])).
fof(f20696,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl245 | $spl290)),
  inference(superposition,[],[f20268,f476])).
fof(f20268,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl245 | $spl290)),
  inference(forward_demodulation,[],[f8816,f13832])).
fof(f15003,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl298),
  inference(cnf_transformation,[],[f15003_D])).
fof(f15003_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f20695,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl245 | $spl290)),
  inference(superposition,[],[f20268,f476])).
fof(f20694,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl245 | $spl290)),
  inference(superposition,[],[f20268,f475])).
fof(f20693,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl245 | $spl290)),
  inference(superposition,[],[f20268,f475])).
fof(f20373,plain,(
  $false | ($spl19 | $spl56 | $spl290)),
  inference(subsumption_resolution,[],[f20372,f14062])).
fof(f20372,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ($spl19 | $spl290)),
  inference(subsumption_resolution,[],[f20322,f458])).
fof(f20322,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | ($spl19 | $spl290)),
  inference(superposition,[],[f18993,f433])).
fof(f18993,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) | ($spl19 | $spl290)),
  inference(forward_demodulation,[],[f732,f13832])).
fof(f20362,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) | $spl375),
  inference(cnf_transformation,[],[f20362_D])).
fof(f20362_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) <=> ~$spl375),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl375])])).
fof(f20280,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl19 | $spl290)),
  inference(superposition,[],[f18993,f476])).
fof(f20278,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl19 | $spl290)),
  inference(superposition,[],[f18993,f475])).
fof(f20211,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl373),
  inference(cnf_transformation,[],[f20211_D])).
fof(f20211_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f20164,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) ) | $spl243),
  inference(superposition,[],[f8811,f429])).
fof(f19432,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl350),
  inference(cnf_transformation,[],[f19432_D])).
fof(f19432_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f20135,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl243),
  inference(superposition,[],[f8811,f476])).
fof(f20205,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl366),
  inference(cnf_transformation,[],[f20205_D])).
fof(f20205_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f20134,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl243),
  inference(superposition,[],[f8811,f476])).
fof(f20133,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl243),
  inference(superposition,[],[f8811,f475])).
fof(f20132,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl243),
  inference(superposition,[],[f8811,f475])).
fof(f8810,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl242),
  inference(cnf_transformation,[],[f8810_D])).
fof(f8810_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f19649,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl365),
  inference(cnf_transformation,[],[f19649_D])).
fof(f19649_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl365),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl365])])).
fof(f19385,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl36 | $spl48)),
  inference(superposition,[],[f766,f781])).
fof(f19647,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))) | $spl363),
  inference(cnf_transformation,[],[f19647_D])).
fof(f19647_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))) <=> ~$spl363),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl363])])).
fof(f19384,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl30 | $spl48)),
  inference(superposition,[],[f753,f781])).
fof(f19645,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl361),
  inference(cnf_transformation,[],[f19645_D])).
fof(f19645_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl361),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl361])])).
fof(f19383,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl26 | $spl48)),
  inference(superposition,[],[f747,f781])).
fof(f8811,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl243),
  inference(cnf_transformation,[],[f8811_D])).
fof(f8811_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f19564,plain,(
  ( ! [X130] : (s('type.num.num',X130) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X130),s('type.num.num',sK4)))) ) | $spl358),
  inference(cnf_transformation,[],[f19564_D])).
fof(f19564_D,plain,(
  ( ! [X130] : (s('type.num.num',X130) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X130),s('type.num.num',sK4)))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f19260,plain,(
  ( ! [X130] : (s('type.num.num',X130) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X130),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl48),
  inference(superposition,[],[f431,f781])).
fof(f431,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f327])).
fof(f19563,plain,(
  ( ! [X126] : (s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X126)))) ) | $spl356),
  inference(cnf_transformation,[],[f19563_D])).
fof(f19563_D,plain,(
  ( ! [X126] : (s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X126)))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f19256,plain,(
  ( ! [X126] : (s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X126))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl48),
  inference(superposition,[],[f430,f781])).
fof(f430,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f327])).
fof(f19436,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl354),
  inference(cnf_transformation,[],[f19436_D])).
fof(f19436_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f19168,plain,(
  ( ! [X14] : (s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) ) | $spl48),
  inference(superposition,[],[f781,f429])).
fof(f19434,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl352),
  inference(cnf_transformation,[],[f19434_D])).
fof(f19434_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f19139,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl48),
  inference(superposition,[],[f781,f476])).
fof(f19137,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl48),
  inference(superposition,[],[f781,f475])).
fof(f19085,plain,(
  $false | $spl49),
  inference(subsumption_resolution,[],[f19084,f519])).
fof(f519,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X7] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.ZERO_LESS_EQ')).
fof(f19084,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl49),
  inference(subsumption_resolution,[],[f19016,f460])).
fof(f19016,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl49),
  inference(superposition,[],[f782,f473])).
fof(f473,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.EQ_CLAUSES|split|1')).
fof(f19083,plain,(
  $false | $spl49),
  inference(subsumption_resolution,[],[f19082,f519])).
fof(f19082,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl49),
  inference(subsumption_resolution,[],[f19015,f460])).
fof(f19015,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl49),
  inference(superposition,[],[f782,f471])).
fof(f471,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.EQ_CLAUSES|split|0')).
fof(f18983,plain,(
  ( ! [X23,X22] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X23) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X23)))) ) | $spl36),
  inference(subsumption_resolution,[],[f18923,f479])).
fof(f18923,plain,(
  ( ! [X23,X22] : (s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X23) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X23)))) ) | $spl36),
  inference(superposition,[],[f766,f429])).
fof(f18966,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl349),
  inference(cnf_transformation,[],[f18966_D])).
fof(f18966_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f18851,plain,(
  ( ! [X62] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl36),
  inference(superposition,[],[f766,f602])).
fof(f18593,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl346),
  inference(cnf_transformation,[],[f18593_D])).
fof(f18593_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f18596,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) | ($spl44 | $spl148 | $spl290)),
  inference(forward_demodulation,[],[f18334,f14102])).
fof(f14102,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl148 | $spl290)),
  inference(backward_demodulation,[],[f13832,f4374])).
fof(f18595,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num','const.num.|48|_0')))) | ($spl44 | $spl148 | $spl290)),
  inference(forward_demodulation,[],[f18333,f14102])).
fof(f18594,plain,(
  ( ! [X13] : (s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num',X13)))) ) | ($spl44 | $spl148 | $spl290)),
  inference(forward_demodulation,[],[f18330,f14102])).
fof(f18591,plain,(
  ( ! [X12] : (s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | ($spl44 | $spl148 | $spl290)),
  inference(forward_demodulation,[],[f18329,f14102])).
fof(f17730,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl240 | $spl290)),
  inference(superposition,[],[f14108,f476])).
fof(f14108,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ($spl240 | $spl290)),
  inference(backward_demodulation,[],[f13832,f8804])).
fof(f17728,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ($spl240 | $spl290)),
  inference(superposition,[],[f14108,f475])).
fof(f17313,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl345),
  inference(cnf_transformation,[],[f17313_D])).
fof(f17313_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl345),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl345])])).
fof(f17245,plain,(
  ( ! [X62] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl20),
  inference(superposition,[],[f735,f602])).
fof(f735,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | $spl20),
  inference(cnf_transformation,[],[f735_D])).
fof(f735_D,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f17202,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X19)))) ) | $spl16),
  inference(subsumption_resolution,[],[f17141,f479])).
fof(f17141,plain,(
  ( ! [X19,X18] : (s('type.num.num',X18) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X19)))) ) | $spl16),
  inference(superposition,[],[f730,f429])).
fof(f17185,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK4))) | $spl343),
  inference(cnf_transformation,[],[f17185_D])).
fof(f17185_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK4))) <=> ~$spl343),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl343])])).
fof(f17094,plain,(
  ( ! [X62] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl16),
  inference(superposition,[],[f730,f602])).
fof(f17059,plain,(
  ( ! [X24,X23,X25] : (sP0(X23,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14107,f429])).
fof(f14107,plain,(
  ( ! [X37,X36] : (sP0(X36,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)),X37)) ) | ($spl56 | $spl290)),
  inference(backward_demodulation,[],[f13832,f7709])).
fof(f730,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) | $spl16),
  inference(cnf_transformation,[],[f730_D])).
fof(f730_D,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f17058,plain,(
  ( ! [X21,X22,X20] : (sP0(X20,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14107,f429])).
fof(f17030,plain,(
  ( ! [X24,X23,X25] : (sP0(X23,X24,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14106,f429])).
fof(f14106,plain,(
  ( ! [X35,X34] : (sP0(X34,X35,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) ) | ($spl56 | $spl290)),
  inference(backward_demodulation,[],[f13832,f7708])).
fof(f17029,plain,(
  ( ! [X21,X22,X20] : (sP0(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14106,f429])).
fof(f17002,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,X24,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14105,f429])).
fof(f14105,plain,(
  ( ! [X30,X31] : (sP1(X30,X31,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) ) | ($spl56 | $spl290)),
  inference(backward_demodulation,[],[f13832,f7706])).
fof(f17001,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14105,f429])).
fof(f16890,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,X24,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14104,f429])).
fof(f14104,plain,(
  ( ! [X24,X25] : (sP2(X24,X25,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) ) | ($spl56 | $spl290)),
  inference(backward_demodulation,[],[f13832,f7703])).
fof(f16889,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14104,f429])).
fof(f16864,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14103,f429])).
fof(f14103,plain,(
  ( ! [X19,X18] : (sP3(X18,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)),X19)) ) | ($spl56 | $spl290)),
  inference(backward_demodulation,[],[f13832,f7700])).
fof(f16863,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl56 | $spl290)),
  inference(superposition,[],[f14103,f429])).
fof(f16321,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14083,f429])).
fof(f14083,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | ($spl53 | $spl290)),
  inference(backward_demodulation,[],[f13832,f943])).
fof(f16320,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14083,f429])).
fof(f16284,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14082,f429])).
fof(f14082,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X16,X17)) ) | ($spl53 | $spl290)),
  inference(backward_demodulation,[],[f13832,f940])).
fof(f16283,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14082,f429])).
fof(f16163,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14081,f429])).
fof(f14081,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X15)) ) | ($spl53 | $spl290)),
  inference(backward_demodulation,[],[f13832,f939])).
fof(f16162,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14081,f429])).
fof(f16127,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14080,f429])).
fof(f14080,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | ($spl53 | $spl290)),
  inference(backward_demodulation,[],[f13832,f937])).
fof(f16126,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14080,f429])).
fof(f16091,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14079,f429])).
fof(f14079,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X9)) ) | ($spl53 | $spl290)),
  inference(backward_demodulation,[],[f13832,f936])).
fof(f16090,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14079,f429])).
fof(f16055,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14078,f429])).
fof(f14078,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X4,X5)) ) | ($spl53 | $spl290)),
  inference(backward_demodulation,[],[f13832,f934])).
fof(f16054,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14078,f429])).
fof(f16020,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14077,f429])).
fof(f14077,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) ) | ($spl53 | $spl290)),
  inference(backward_demodulation,[],[f13832,f933])).
fof(f16019,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl53 | $spl290)),
  inference(superposition,[],[f14077,f429])).
fof(f15879,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14075,f429])).
fof(f14075,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)),X22,X23)) ) | ($spl1 | $spl290)),
  inference(backward_demodulation,[],[f13832,f847])).
fof(f15878,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14075,f429])).
fof(f15852,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14074,f429])).
fof(f14074,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)),X16,X17)) ) | ($spl1 | $spl290)),
  inference(backward_demodulation,[],[f13832,f844])).
fof(f15851,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14074,f429])).
fof(f15823,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14073,f429])).
fof(f14073,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)),X15)) ) | ($spl1 | $spl290)),
  inference(backward_demodulation,[],[f13832,f843])).
fof(f15822,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14073,f429])).
fof(f15797,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14072,f429])).
fof(f14072,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)),X10,X11)) ) | ($spl1 | $spl290)),
  inference(backward_demodulation,[],[f13832,f841])).
fof(f15796,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14072,f429])).
fof(f15323,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK5)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK5))))) | ($spl54 | $spl290)),
  inference(superposition,[],[f14061,f474])).
fof(f14061,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5)))) | ($spl54 | $spl290)),
  inference(backward_demodulation,[],[f13832,f796])).
fof(f15322,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK5)),s('type.num.num','const.num.|48|_0')))) | ($spl54 | $spl290)),
  inference(superposition,[],[f14061,f464])).
fof(f15319,plain,(
  ( ! [X13] : (s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK5)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK5)),s('type.num.num',X13)))) ) | ($spl54 | $spl290)),
  inference(superposition,[],[f14061,f429])).
fof(f15318,plain,(
  ( ! [X12] : (s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK5))))) ) | ($spl54 | $spl290)),
  inference(superposition,[],[f14061,f429])).
fof(f15297,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14071,f429])).
fof(f14071,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)),X9)) ) | ($spl1 | $spl290)),
  inference(backward_demodulation,[],[f13832,f840])).
fof(f15296,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14071,f429])).
fof(f15271,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14070,f429])).
fof(f14070,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)),X4,X5)) ) | ($spl1 | $spl290)),
  inference(backward_demodulation,[],[f13832,f838])).
fof(f15270,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14070,f429])).
fof(f15246,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14069,f429])).
fof(f14069,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) ) | ($spl1 | $spl290)),
  inference(backward_demodulation,[],[f13832,f837])).
fof(f15245,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | ($spl1 | $spl290)),
  inference(superposition,[],[f14069,f429])).
fof(f15049,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))))) | $spl328),
  inference(cnf_transformation,[],[f15049_D])).
fof(f15049_D,plain,(
  s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))))) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f15044,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl323),
  inference(cnf_transformation,[],[f15044_D])).
fof(f15044_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl323),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl323])])).
fof(f14990,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))))) | $spl290),
  inference(superposition,[],[f14055,f474])).
fof(f14055,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) | $spl290),
  inference(backward_demodulation,[],[f13832,f391])).
fof(f15047,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))),s('type.num.num','const.num.|48|_0')))) | $spl327),
  inference(cnf_transformation,[],[f15047_D])).
fof(f15047_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f14989,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))),s('type.num.num','const.num.|48|_0')))) | $spl290),
  inference(superposition,[],[f14055,f464])).
fof(f15045,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))),s('type.num.num',X13)))) ) | $spl324),
  inference(cnf_transformation,[],[f15045_D])).
fof(f15045_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))),s('type.num.num',X13)))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f14986,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))),s('type.num.num',X13)))) ) | $spl290),
  inference(superposition,[],[f14055,f429])).
fof(f15042,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))))) ) | $spl320),
  inference(cnf_transformation,[],[f15042_D])).
fof(f15042_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f14985,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))))) ) | $spl290),
  inference(superposition,[],[f14055,f429])).
fof(f15041,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))) | $spl318),
  inference(cnf_transformation,[],[f15041_D])).
fof(f15041_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f15014,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl307),
  inference(cnf_transformation,[],[f15014_D])).
fof(f15014_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl307),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl307])])).
fof(f14978,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))) | $spl290),
  inference(superposition,[],[f14055,f474])).
fof(f15039,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))),s('type.num.num','const.num.|48|_0')))) | $spl317),
  inference(cnf_transformation,[],[f15039_D])).
fof(f15039_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl317),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl317])])).
fof(f14977,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))),s('type.num.num','const.num.|48|_0')))) | $spl290),
  inference(superposition,[],[f14055,f464])).
fof(f15037,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))),s('type.num.num',X18)))) ) | $spl314),
  inference(cnf_transformation,[],[f15037_D])).
fof(f15037_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))),s('type.num.num',X18)))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f14974,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))),s('type.num.num',X18)))) ) | $spl290),
  inference(superposition,[],[f14055,f429])).
fof(f15036,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))) ) | $spl312),
  inference(cnf_transformation,[],[f15036_D])).
fof(f15036_D,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f14973,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))))) ) | $spl290),
  inference(superposition,[],[f14055,f429])).
fof(f15032,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))) | $spl311),
  inference(cnf_transformation,[],[f15032_D])).
fof(f15032_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))) <=> ~$spl311),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl311])])).
fof(f14966,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4)))) | $spl290),
  inference(superposition,[],[f14055,f433])).
fof(f15024,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl290),
  inference(forward_demodulation,[],[f15023,f480])).
fof(f15023,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl290),
  inference(forward_demodulation,[],[f15022,f478])).
fof(f15022,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl290),
  inference(forward_demodulation,[],[f14953,f479])).
fof(f14953,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl290),
  inference(superposition,[],[f14055,f429])).
fof(f15019,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl290),
  inference(forward_demodulation,[],[f15018,f480])).
fof(f15018,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl290),
  inference(forward_demodulation,[],[f15017,f478])).
fof(f15017,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl290),
  inference(forward_demodulation,[],[f14952,f479])).
fof(f14952,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl290),
  inference(superposition,[],[f14055,f429])).
fof(f776,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl44),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f15016,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl290),
  inference(forward_demodulation,[],[f14947,f480])).
fof(f14947,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl290),
  inference(superposition,[],[f14055,f474])).
fof(f15015,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) ) | $spl290),
  inference(forward_demodulation,[],[f14943,f480])).
fof(f14943,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) ) | $spl290),
  inference(superposition,[],[f14055,f429])).
fof(f766,plain,(
  ( ! [X13] : (s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | $spl36),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  ( ! [X13] : (s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f15012,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl290),
  inference(forward_demodulation,[],[f14942,f480])).
fof(f14942,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl290),
  inference(superposition,[],[f14055,f429])).
fof(f15009,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) | $spl305),
  inference(cnf_transformation,[],[f15009_D])).
fof(f15009_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) <=> ~$spl305),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl305])])).
fof(f14924,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl290),
  inference(superposition,[],[f14055,f476])).
fof(f15005,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl301),
  inference(cnf_transformation,[],[f15005_D])).
fof(f15005_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl301),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl301])])).
fof(f15010,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl290),
  inference(forward_demodulation,[],[f14923,f460])).
fof(f14923,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl290),
  inference(superposition,[],[f14055,f476])).
fof(f14922,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl290),
  inference(superposition,[],[f14055,f475])).
fof(f15001,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl290),
  inference(forward_demodulation,[],[f14921,f460])).
fof(f14921,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK4))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl290),
  inference(superposition,[],[f14055,f475])).
fof(f14337,plain,(
  $false | ($spl61 | $spl290)),
  inference(subsumption_resolution,[],[f14063,f511])).
fof(f511,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X6,X7] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.LESS_EQ_ADD')).
fof(f14063,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ($spl61 | $spl290)),
  inference(backward_demodulation,[],[f13832,f802])).
fof(f13832,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK5) | $spl290),
  inference(cnf_transformation,[],[f13832_D])).
fof(f13832_D,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK5) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f13840,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK5) | s('type.num.num',sK8(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK4))))) | ($spl24 | $spl148)),
  inference(forward_demodulation,[],[f13578,f4374])).
fof(f13578,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK4))))) | $spl24),
  inference(superposition,[],[f743,f474])).
fof(f13837,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK4)),s('type.num.num','const.num.|48|_0')))) | ($spl24 | $spl148)),
  inference(forward_demodulation,[],[f13577,f4374])).
fof(f13577,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK4)),s('type.num.num','const.num.|48|_0')))) | $spl24),
  inference(superposition,[],[f743,f464])).
fof(f13833,plain,(
  ( ! [X13] : (s('type.num.num',sK4) = s('type.num.num',sK5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK4)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK4)),s('type.num.num',X13)))) ) | ($spl24 | $spl148)),
  inference(forward_demodulation,[],[f13574,f4374])).
fof(f13574,plain,(
  ( ! [X13] : (s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK4)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK4)),s('type.num.num',X13)))) ) | $spl24),
  inference(superposition,[],[f743,f429])).
fof(f13829,plain,(
  ( ! [X12] : (s('type.num.num',sK4) = s('type.num.num',sK5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK4))))) ) | ($spl24 | $spl148)),
  inference(forward_demodulation,[],[f13573,f4374])).
fof(f13573,plain,(
  ( ! [X12] : (s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK4))))) ) | $spl24),
  inference(superposition,[],[f743,f429])).
fof(f13568,plain,(
  $false | $spl282),
  inference(subsumption_resolution,[],[f13498,f13121])).
fof(f13498,plain,(
  ( ! [X233,X234] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X233),s('type.num.num',X234))))) ) | $spl282),
  inference(resolution,[],[f13121,f513])).
fof(f13567,plain,(
  $false | $spl282),
  inference(subsumption_resolution,[],[f13468,f13121])).
fof(f13468,plain,(
  ( ! [X187,X186] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X186),s('type.num.num',X187))))) ) | $spl282),
  inference(resolution,[],[f13121,f513])).
fof(f13558,plain,(
  $false | $spl282),
  inference(subsumption_resolution,[],[f13450,f13121])).
fof(f13450,plain,(
  ~p(s(bool,'T_0')) | $spl282),
  inference(resolution,[],[f13121,f596])).
fof(f596,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f387])).
fof(f387,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f251])).
fof(f251,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f57])).
fof(f57,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.NOT_CLAUSES|split|2')).
fof(f13557,plain,(
  $false | $spl282),
  inference(subsumption_resolution,[],[f13449,f13121])).
fof(f13449,plain,(
  ~p(s(bool,'T_0')) | $spl282),
  inference(resolution,[],[f13121,f594])).
fof(f594,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f386])).
fof(f386,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f250])).
fof(f250,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f56])).
fof(f56,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.NOT_CLAUSES|split|1')).
fof(f13556,plain,(
  $false | $spl282),
  inference(subsumption_resolution,[],[f13441,f13121])).
fof(f13441,plain,(
  ( ! [X132] : (~p(s(bool,X132))) ) | $spl282),
  inference(resolution,[],[f13121,f543])).
fof(f543,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f354])).
fof(f354,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X17] : ((p(s(bool,X17)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X17)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.IMP_CLAUSES|split|4')).
fof(f13448,plain,(
  $false | $spl282),
  inference(resolution,[],[f13121,f592])).
fof(f592,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.arithmetic.SUB_LEFT_SUB',aHL_FALSITY)).
fof(f13470,plain,(
  $false | ($spl81 | $spl282)),
  inference(resolution,[],[f13121,f820])).
fof(f13471,plain,(
  $false | ($spl43 | $spl282)),
  inference(resolution,[],[f13121,f773])).
fof(f13472,plain,(
  $false | ($spl75 | $spl282)),
  inference(resolution,[],[f13121,f813])).
fof(f13474,plain,(
  $false | $spl282),
  inference(resolution,[],[f13121,f518])).
fof(f518,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X7] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num','const.num.|48|_0'))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.NOT_SUC_LESS_EQ_|48|')).
fof(f13475,plain,(
  $false | ($spl9 | $spl282)),
  inference(resolution,[],[f13121,f7232])).
fof(f7232,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl9),
  inference(subsumption_resolution,[],[f7167,f480])).
fof(f7167,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl9),
  inference(superposition,[],[f718,f464])).
fof(f13476,plain,(
  $false | ($spl53 | $spl282)),
  inference(resolution,[],[f13121,f791])).
fof(f13477,plain,(
  $false | ($spl61 | $spl282)),
  inference(resolution,[],[f13121,f802])).
fof(f13478,plain,(
  $false | ($spl1 | $spl282)),
  inference(resolution,[],[f13121,f708])).
fof(f13479,plain,(
  $false | ($spl33 | $spl282)),
  inference(resolution,[],[f13121,f758])).
fof(f13494,plain,(
  $false | $spl282),
  inference(resolution,[],[f13121,f606])).
fof(f606,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(rectify,[],[f127])).
fof(f127,axiom,(
  ! [X7] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X7))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.prim_rec.LESS_REFL')).
fof(f13501,plain,(
  $false | $spl282),
  inference(resolution,[],[f13121,f607])).
fof(f607,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f114])).
fof(f114,axiom,(
  ! [X7] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0'))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.prim_rec.NOT_LESS_|48|')).
fof(f13502,plain,(
  $false | ($spl1 | $spl282)),
  inference(resolution,[],[f13121,f833])).
fof(f13121,plain,(
  ( ! [X10] : (p(s(bool,X10))) ) | $spl282),
  inference(cnf_transformation,[],[f13121_D])).
fof(f13121_D,plain,(
  ( ! [X10] : (p(s(bool,X10))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f13122,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) | $spl284),
  inference(cnf_transformation,[],[f13122_D])).
fof(f13122_D,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f13380,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | p(s(bool,X10))) ) | $spl132),
  inference(subsumption_resolution,[],[f13363,f519])).
fof(f13363,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X10))) ) | $spl132),
  inference(resolution,[],[f2265,f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f13359,plain,(
  $false | ($spl45 | $spl69)),
  inference(subsumption_resolution,[],[f13353,f775])).
fof(f775,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl45),
  inference(cnf_transformation,[],[f775_D])).
fof(f775_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f13353,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl69),
  inference(trivial_inequality_removal,[],[f13352])).
fof(f13352,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl69),
  inference(superposition,[],[f807,f474])).
fof(f13358,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl287),
  inference(cnf_transformation,[],[f13358_D])).
fof(f13358_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f13329,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl69),
  inference(superposition,[],[f807,f474])).
fof(f13325,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl69),
  inference(superposition,[],[f807,f429])).
fof(f13324,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl69),
  inference(superposition,[],[f807,f429])).
fof(f13310,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | ~p(s(bool,X10))) ) | $spl130),
  inference(subsumption_resolution,[],[f13294,f519])).
fof(f13294,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9)) | ~p(s(bool,X10))) ) | $spl130),
  inference(resolution,[],[f2228,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f311])).
fof(f13290,plain,(
  $false | $spl284),
  inference(subsumption_resolution,[],[f13226,f13122])).
fof(f13226,plain,(
  ( ! [X182,X181] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X181),s('type.num.num',X182))))) ) | $spl284),
  inference(resolution,[],[f13122,f512])).
fof(f13289,plain,(
  $false | $spl284),
  inference(subsumption_resolution,[],[f13202,f13122])).
fof(f13202,plain,(
  ( ! [X136,X137] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X136),s('type.num.num',X137))))) ) | $spl284),
  inference(resolution,[],[f13122,f512])).
fof(f13288,plain,(
  $false | $spl284),
  inference(subsumption_resolution,[],[f13193,f13122])).
fof(f13193,plain,(
  p(s(bool,'F_0')) | $spl284),
  inference(resolution,[],[f13122,f595])).
fof(f595,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f387])).
fof(f13287,plain,(
  $false | $spl284),
  inference(subsumption_resolution,[],[f13192,f13122])).
fof(f13192,plain,(
  p(s(bool,'F_0')) | $spl284),
  inference(resolution,[],[f13122,f593])).
fof(f593,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f386])).
fof(f13286,plain,(
  $false | $spl284),
  inference(subsumption_resolution,[],[f13187,f13122])).
fof(f13187,plain,(
  ( ! [X129] : (p(s(bool,X129))) ) | $spl284),
  inference(resolution,[],[f13122,f554])).
fof(f554,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f362])).
fof(f362,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X17] : ((p(s(bool,X17)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.IMP_CLAUSES|split|1')).
fof(f13285,plain,(
  $false | $spl284),
  inference(subsumption_resolution,[],[f13185,f13122])).
fof(f13185,plain,(
  p(s(bool,'F_0')) | $spl284),
  inference(resolution,[],[f13122,f551])).
fof(f551,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f360])).
fof(f360,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X17] : ((p(s(bool,'F_0')) => p(s(bool,X17))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.IMP_CLAUSES|split|2')).
fof(f13284,plain,(
  $false | $spl284),
  inference(subsumption_resolution,[],[f13183,f13122])).
fof(f13183,plain,(
  ( ! [X126] : (p(s(bool,X126))) ) | $spl284),
  inference(resolution,[],[f13122,f548])).
fof(f548,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X17] : ((p(s(bool,X17)) => p(s(bool,X17))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.IMP_CLAUSES|split|3')).
fof(f13283,plain,(
  $false | $spl284),
  inference(subsumption_resolution,[],[f13182,f13122])).
fof(f13182,plain,(
  ( ! [X125] : (p(s(bool,X125))) ) | $spl284),
  inference(resolution,[],[f13122,f545])).
fof(f545,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X17] : ((p(s(bool,'T_0')) => p(s(bool,X17))) <=> p(s(bool,X17)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.IMP_CLAUSES|split|0')).
fof(f13194,plain,(
  $false | $spl284),
  inference(resolution,[],[f13122,f599])).
fof(f599,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f12])).
fof(f12,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.TRUTH')).
fof(f13195,plain,(
  $false | $spl284),
  inference(resolution,[],[f13122,f600])).
fof(f13196,plain,(
  $false | $spl284),
  inference(resolution,[],[f13122,f609])).
fof(f13205,plain,(
  $false | $spl284),
  inference(resolution,[],[f13122,f511])).
fof(f13207,plain,(
  $false | $spl284),
  inference(resolution,[],[f13122,f517])).
fof(f517,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X0))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X7,X6] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X7))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_LESS_EQ')).
fof(f13208,plain,(
  $false | $spl284),
  inference(resolution,[],[f13122,f519])).
fof(f13209,plain,(
  $false | ($spl56 | $spl284)),
  inference(resolution,[],[f13122,f797])).
fof(f13227,plain,(
  $false | $spl284),
  inference(resolution,[],[f13122,f601])).
fof(f601,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X8,X6] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_LESS_SUC')).
fof(f13229,plain,(
  $false | $spl284),
  inference(resolution,[],[f13122,f608])).
fof(f608,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.prim_rec.LESS_|48|')).
fof(f13230,plain,(
  $false | ($spl33 | $spl284)),
  inference(resolution,[],[f13122,f1248])).
fof(f1248,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl33),
  inference(resolution,[],[f758,f512])).
fof(f13231,plain,(
  $false | ($spl53 | $spl284)),
  inference(resolution,[],[f13122,f930])).
fof(f930,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | $spl53),
  inference(resolution,[],[f791,f512])).
fof(f13232,plain,(
  $false | ($spl1 | $spl284)),
  inference(resolution,[],[f13122,f834])).
fof(f13120,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | p(s(bool,X10))) ) | $spl128),
  inference(subsumption_resolution,[],[f13103,f519])).
fof(f13103,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl128),
  inference(resolution,[],[f2191,f392])).
fof(f13021,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) | $spl280),
  inference(cnf_transformation,[],[f13021_D])).
fof(f13021_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f807,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl69),
  inference(cnf_transformation,[],[f807_D])).
fof(f807_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f13008,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) | $spl45),
  inference(superposition,[],[f775,f474])).
fof(f13019,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))),s('type.num.num',X15)))) ) | $spl278),
  inference(cnf_transformation,[],[f13019_D])).
fof(f13019_D,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))),s('type.num.num',X15)))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f13004,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))),s('type.num.num',X15)))) ) | $spl45),
  inference(superposition,[],[f775,f429])).
fof(f13018,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num',X14) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) ) | $spl276),
  inference(cnf_transformation,[],[f13018_D])).
fof(f13018_D,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num',X14) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f13003,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num',X14) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) ) | $spl45),
  inference(superposition,[],[f775,f429])).
fof(f12618,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl267),
  inference(cnf_transformation,[],[f12618_D])).
fof(f12618_D,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl267),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl267])])).
fof(f13016,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) | ($spl45 | $spl148)),
  inference(forward_demodulation,[],[f12993,f4374])).
fof(f12993,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) | $spl45),
  inference(superposition,[],[f775,f474])).
fof(f13015,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num','const.num.|48|_0')))) | ($spl45 | $spl148)),
  inference(forward_demodulation,[],[f12992,f4374])).
fof(f12992,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num','const.num.|48|_0')))) | $spl45),
  inference(superposition,[],[f775,f464])).
fof(f13014,plain,(
  ( ! [X13] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num',X13)))) ) | ($spl45 | $spl148)),
  inference(forward_demodulation,[],[f12989,f4374])).
fof(f12989,plain,(
  ( ! [X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num',X13)))) ) | $spl45),
  inference(superposition,[],[f775,f429])).
fof(f13013,plain,(
  ( ! [X12] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | ($spl45 | $spl148)),
  inference(forward_demodulation,[],[f12988,f4374])).
fof(f12988,plain,(
  ( ! [X12] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | $spl45),
  inference(superposition,[],[f775,f429])).
fof(f12612,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | $spl263),
  inference(cnf_transformation,[],[f12612_D])).
fof(f12612_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) <=> ~$spl263),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl263])])).
fof(f13012,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl45),
  inference(forward_demodulation,[],[f12971,f478])).
fof(f12971,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl45),
  inference(superposition,[],[f775,f474])).
fof(f13011,plain,(
  ( ! [X13] : (s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl45),
  inference(forward_demodulation,[],[f12967,f478])).
fof(f12967,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl45),
  inference(superposition,[],[f775,f429])).
fof(f13010,plain,(
  ( ! [X12] : (s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl45),
  inference(forward_demodulation,[],[f12966,f478])).
fof(f12966,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl45),
  inference(superposition,[],[f775,f429])).
fof(f806,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl68),
  inference(cnf_transformation,[],[f806_D])).
fof(f806_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f12645,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) | $spl274),
  inference(cnf_transformation,[],[f12645_D])).
fof(f12645_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f12787,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl44),
  inference(forward_demodulation,[],[f12512,f776])).
fof(f12512,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) | $spl44),
  inference(superposition,[],[f474,f776])).
fof(f12643,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl44),
  inference(forward_demodulation,[],[f12368,f776])).
fof(f12368,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))))) | $spl44),
  inference(superposition,[],[f776,f474])).
fof(f12619,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl266),
  inference(cnf_transformation,[],[f12619_D])).
fof(f12619_D,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f12625,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) | ($spl44 | $spl148)),
  inference(forward_demodulation,[],[f12354,f4374])).
fof(f12354,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) | $spl44),
  inference(superposition,[],[f776,f474])).
fof(f12622,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num','const.num.|48|_0')))) | ($spl44 | $spl148)),
  inference(forward_demodulation,[],[f12353,f4374])).
fof(f12353,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num','const.num.|48|_0')))) | $spl44),
  inference(superposition,[],[f776,f464])).
fof(f12620,plain,(
  ( ! [X13] : (s('type.num.num',sK5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num',X13)))) ) | ($spl44 | $spl148)),
  inference(forward_demodulation,[],[f12350,f4374])).
fof(f12350,plain,(
  ( ! [X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),s('type.num.num',X13)))) ) | $spl44),
  inference(superposition,[],[f776,f429])).
fof(f12616,plain,(
  ( ! [X12] : (s('type.num.num',sK5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | ($spl44 | $spl148)),
  inference(forward_demodulation,[],[f12349,f4374])).
fof(f12349,plain,(
  ( ! [X12] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) ) | $spl44),
  inference(superposition,[],[f776,f429])).
fof(f12615,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl44),
  inference(forward_demodulation,[],[f12332,f478])).
fof(f12332,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl44),
  inference(superposition,[],[f776,f474])).
fof(f12614,plain,(
  ( ! [X13] : (s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl44),
  inference(forward_demodulation,[],[f12328,f478])).
fof(f12328,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl44),
  inference(superposition,[],[f776,f429])).
fof(f12611,plain,(
  ( ! [X12] : (s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl44),
  inference(forward_demodulation,[],[f12327,f478])).
fof(f12327,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl44),
  inference(superposition,[],[f776,f429])).
fof(f11698,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl240),
  inference(superposition,[],[f8804,f476])).
fof(f11696,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl240),
  inference(superposition,[],[f8804,f475])).
fof(f11590,plain,(
  $false | ($spl33 | $spl99)),
  inference(subsumption_resolution,[],[f1134,f1248])).
fof(f11551,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f420])).
fof(f11550,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f419])).
fof(f11549,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2))))) ) | $spl87),
  inference(superposition,[],[f825,f419])).
fof(f11548,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f417])).
fof(f11547,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f417])).
fof(f11328,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f420])).
fof(f11327,plain,(
  ( ! [X3] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f419])).
fof(f11326,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2))))) ) | $spl9),
  inference(superposition,[],[f718,f419])).
fof(f11325,plain,(
  ( ! [X1] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f417])).
fof(f11324,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f417])).
fof(f11142,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))))),
  inference(superposition,[],[f391,f474])).
fof(f11141,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f11138,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f11137,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) )),
  inference(superposition,[],[f391,f429])).
fof(f11130,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))),
  inference(superposition,[],[f391,f474])).
fof(f11126,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f391,f429])).
fof(f11125,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f809,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl71),
  inference(cnf_transformation,[],[f809_D])).
fof(f809_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f11116,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f430])).
fof(f11164,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))),
  inference(forward_demodulation,[],[f11099,f480])).
fof(f11099,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))),
  inference(superposition,[],[f391,f474])).
fof(f11163,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(forward_demodulation,[],[f11095,f480])).
fof(f11095,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(superposition,[],[f391,f429])).
fof(f11162,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) )),
  inference(forward_demodulation,[],[f11094,f480])).
fof(f11094,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f11160,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f11159,f480])).
fof(f11159,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f11076,f478])).
fof(f11076,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(superposition,[],[f391,f474])).
fof(f11158,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f11157,f480])).
fof(f11157,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f11072,f478])).
fof(f11072,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f11156,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f11155,f480])).
fof(f11155,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f11071,f478])).
fof(f11071,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(superposition,[],[f391,f429])).
fof(f11066,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f11154,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f11065,f460])).
fof(f11065,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f11064,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f11153,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f11063,f460])).
fof(f11063,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f10956,plain,(
  $false | ($spl56 | $spl81 | $spl86)),
  inference(subsumption_resolution,[],[f10955,f519])).
fof(f10955,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl56 | $spl81 | $spl86)),
  inference(forward_demodulation,[],[f10954,f824])).
fof(f824,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl86),
  inference(cnf_transformation,[],[f824_D])).
fof(f824_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f10954,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | ($spl56 | $spl81)),
  inference(subsumption_resolution,[],[f10892,f797])).
fof(f10892,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl81),
  inference(superposition,[],[f820,f441])).
fof(f10947,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl259),
  inference(cnf_transformation,[],[f10947_D])).
fof(f10947_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f10949,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl81),
  inference(forward_demodulation,[],[f10858,f478])).
fof(f10858,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl81),
  inference(superposition,[],[f820,f474])).
fof(f10948,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl81),
  inference(forward_demodulation,[],[f10854,f478])).
fof(f10854,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl81),
  inference(superposition,[],[f820,f429])).
fof(f10945,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl81),
  inference(forward_demodulation,[],[f10853,f478])).
fof(f10853,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl81),
  inference(superposition,[],[f820,f429])).
fof(f10405,plain,(
  $false | ($spl1 | $spl241)),
  inference(subsumption_resolution,[],[f10404,f708])).
fof(f10404,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl241),
  inference(subsumption_resolution,[],[f10347,f459])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X12] : ! [X20,X21] : s(X12,X21) = s(X12,'const.bool.COND_3'(s(bool,'F_0'),s(X12,X20),s(X12,X21)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.COND_CLAUSES|split|1')).
fof(f10347,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl241),
  inference(superposition,[],[f8805,f469])).
fof(f469,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X17] : (s(bool,'F_0') = s(bool,X17) <=> ~p(s(bool,X17)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.EQ_CLAUSES|split|2')).
fof(f10403,plain,(
  $false | ($spl1 | $spl241)),
  inference(subsumption_resolution,[],[f10402,f708])).
fof(f10402,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl241),
  inference(subsumption_resolution,[],[f10346,f459])).
fof(f10346,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl241),
  inference(superposition,[],[f8805,f467])).
fof(f467,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X17] : (s(bool,'F_0') = s(bool,X17) <=> ~p(s(bool,X17)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.EQ_CLAUSES|split|3')).
fof(f10084,plain,(
  $false | ($spl1 | $spl120 | $spl148)),
  inference(subsumption_resolution,[],[f9819,f1139])).
fof(f1139,plain,(
  s('type.num.num',sK5) != s('type.num.num',sK6) | $spl1),
  inference(resolution,[],[f834,f604])).
fof(f604,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1))),
  inference(ennf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => s('type.num.num',X0) != s('type.num.num',X1))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => ~s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) => ~s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.prim_rec.LESS_NOT_EQ')).
fof(f9819,plain,(
  s('type.num.num',sK5) = s('type.num.num',sK6) | ($spl120 | $spl148)),
  inference(backward_demodulation,[],[f4374,f1734])).
fof(f9739,plain,(
  ( ! [X62] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl50),
  inference(superposition,[],[f786,f602])).
fof(f9697,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X19)))) ) | $spl46),
  inference(subsumption_resolution,[],[f9641,f479])).
fof(f9641,plain,(
  ( ! [X19,X18] : (s('type.num.num',X18) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X19)))) ) | $spl46),
  inference(superposition,[],[f780,f429])).
fof(f9595,plain,(
  ( ! [X62] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl46),
  inference(superposition,[],[f780,f602])).
fof(f9554,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl253),
  inference(cnf_transformation,[],[f9554_D])).
fof(f9554_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f9487,plain,(
  ( ! [X62] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl30),
  inference(superposition,[],[f753,f602])).
fof(f753,plain,(
  ( ! [X13] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | $spl30),
  inference(cnf_transformation,[],[f753_D])).
fof(f753_D,plain,(
  ( ! [X13] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f9444,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X19)))) ) | $spl26),
  inference(subsumption_resolution,[],[f9386,f479])).
fof(f9386,plain,(
  ( ! [X19,X18] : (s('type.num.num',X18) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X19)))) ) | $spl26),
  inference(superposition,[],[f747,f429])).
fof(f9431,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK6))) | $spl249),
  inference(cnf_transformation,[],[f9431_D])).
fof(f9431_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK6))) <=> ~$spl249),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl249])])).
fof(f9340,plain,(
  ( ! [X62] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))),s('type.num.num',X62)))) ) | $spl26),
  inference(superposition,[],[f747,f602])).
fof(f8805,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl241),
  inference(cnf_transformation,[],[f8805_D])).
fof(f8805_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl241),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl241])])).
fof(f8761,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) | $spl86),
  inference(superposition,[],[f7933,f474])).
fof(f7933,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | $spl86),
  inference(backward_demodulation,[],[f824,f391])).
fof(f8760,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl86),
  inference(superposition,[],[f7933,f464])).
fof(f8757,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8756,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8749,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) | $spl86),
  inference(superposition,[],[f7933,f433])).
fof(f8816,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl245),
  inference(cnf_transformation,[],[f8816_D])).
fof(f8816_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f8747,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl86),
  inference(superposition,[],[f7933,f430])).
fof(f8812,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl86),
  inference(forward_demodulation,[],[f8736,f479])).
fof(f8736,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8809,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl86),
  inference(forward_demodulation,[],[f8735,f479])).
fof(f8735,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8808,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl86),
  inference(forward_demodulation,[],[f8730,f480])).
fof(f8730,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | $spl86),
  inference(superposition,[],[f7933,f474])).
fof(f773,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl43),
  inference(cnf_transformation,[],[f773_D])).
fof(f773_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f8807,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl86),
  inference(forward_demodulation,[],[f8729,f480])).
fof(f8729,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl86),
  inference(superposition,[],[f7933,f464])).
fof(f8806,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) ) | $spl86),
  inference(forward_demodulation,[],[f8726,f480])).
fof(f8726,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8803,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl86),
  inference(forward_demodulation,[],[f8725,f480])).
fof(f8725,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8799,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | $spl239),
  inference(cnf_transformation,[],[f8799_D])).
fof(f8799_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f8801,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl86),
  inference(forward_demodulation,[],[f8800,f478])).
fof(f8800,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl86),
  inference(forward_demodulation,[],[f8713,f449])).
fof(f8713,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f747,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) | $spl26),
  inference(cnf_transformation,[],[f747_D])).
fof(f747_D,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8797,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl86),
  inference(forward_demodulation,[],[f8796,f478])).
fof(f8796,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl86),
  inference(forward_demodulation,[],[f8712,f449])).
fof(f8712,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8795,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl86),
  inference(forward_demodulation,[],[f8707,f478])).
fof(f8707,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl86),
  inference(superposition,[],[f7933,f474])).
fof(f739,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl23),
  inference(cnf_transformation,[],[f739_D])).
fof(f739_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f8794,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl86),
  inference(forward_demodulation,[],[f8706,f478])).
fof(f8706,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl86),
  inference(superposition,[],[f7933,f464])).
fof(f8793,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl86),
  inference(forward_demodulation,[],[f8703,f478])).
fof(f8703,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8792,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl86),
  inference(forward_demodulation,[],[f8702,f478])).
fof(f8702,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl86),
  inference(superposition,[],[f7933,f429])).
fof(f8791,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | $spl237),
  inference(cnf_transformation,[],[f8791_D])).
fof(f8791_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f8697,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl86),
  inference(superposition,[],[f7933,f476])).
fof(f8789,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | $spl235),
  inference(cnf_transformation,[],[f8789_D])).
fof(f8789_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f8696,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl86),
  inference(superposition,[],[f7933,f476])).
fof(f8695,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl86),
  inference(superposition,[],[f7933,f475])).
fof(f8694,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl86),
  inference(superposition,[],[f7933,f475])).
fof(f8667,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f847,f429])).
fof(f8666,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f847,f429])).
fof(f8639,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f844,f429])).
fof(f8638,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f844,f429])).
fof(f8610,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f843,f429])).
fof(f8609,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f843,f429])).
fof(f8584,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f8583,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f8558,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f8557,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f8293,plain,(
  ( ! [X14] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X14)))))) ) | $spl228),
  inference(cnf_transformation,[],[f8293_D])).
fof(f8293_D,plain,(
  ( ! [X14] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X14)))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f8018,plain,(
  ( ! [X14] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X14))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f510,f824])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ! [X2] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))),
  inference(ennf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) => ! [X2] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X11,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X11)))) => ! [X10] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num',X9))),s('type.num.num',X10))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.LESS_EQ_SUB_LESS')).
fof(f8292,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl226),
  inference(cnf_transformation,[],[f8292_D])).
fof(f8292_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f8290,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | $spl86),
  inference(forward_demodulation,[],[f8015,f478])).
fof(f8015,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | $spl86),
  inference(superposition,[],[f443,f824])).
fof(f443,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) => s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))) => s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X7))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.SUB_ADD')).
fof(f8289,plain,(
  ( ! [X12] : (s('type.num.num',X12) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',sK4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl224),
  inference(cnf_transformation,[],[f8289_D])).
fof(f8289_D,plain,(
  ( ! [X12] : (s('type.num.num',X12) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',sK4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f8288,plain,(
  ( ! [X12] : (s('type.num.num',X12) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',sK4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(forward_demodulation,[],[f8014,f477])).
fof(f477,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.ADD_|48|')).
fof(f8014,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',sK4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f442,f824])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) => ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X10,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))) => ! [X11] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X9))),s('type.num.num',X10))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X9),s('type.num.num',X10))))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.arithmetic.LESS_EQ_ADD_SUB')).
fof(f8287,plain,(
  ( ! [X11] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X11)))))) ) | $spl222),
  inference(cnf_transformation,[],[f8287_D])).
fof(f8287_D,plain,(
  ( ! [X11] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X11)))))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f8013,plain,(
  ( ! [X11] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X11))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f440,f824])).
fof(f8284,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl220),
  inference(cnf_transformation,[],[f8284_D])).
fof(f8284_D,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f8007,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num',sK4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f421,f824])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f325])).
fof(f8263,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl214),
  inference(cnf_transformation,[],[f8263_D])).
fof(f8263_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f8006,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f420,f824])).
fof(f8262,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2)))) | s('type.num.num',sK4) != s('type.num.num',X2)) ) | $spl212),
  inference(cnf_transformation,[],[f8262_D])).
fof(f8262_D,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2)))) | s('type.num.num',sK4) != s('type.num.num',X2)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f8005,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X7)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f419,f824])).
fof(f8004,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X6) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X6))))) ) | $spl86),
  inference(superposition,[],[f419,f824])).
fof(f8283,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X4)))) | s('type.num.num',sK4) = s('type.num.num',X4)) ) | $spl218),
  inference(cnf_transformation,[],[f8283_D])).
fof(f8283_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X4)))) | s('type.num.num',sK4) = s('type.num.num',X4)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f8003,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) = s('type.num.num',X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f418,f824])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X1) = s('type.num.num',X2) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f324])).
fof(f8002,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) = s('type.num.num',X4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X4))))) ) | $spl86),
  inference(superposition,[],[f418,f824])).
fof(f8261,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) | $spl210),
  inference(cnf_transformation,[],[f8261_D])).
fof(f8261_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f8001,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X3))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f417,f824])).
fof(f8000,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X2))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f417,f824])).
fof(f8282,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) | $spl216),
  inference(cnf_transformation,[],[f8282_D])).
fof(f8282_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f7999,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f416,f824])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X1) = s('type.num.num',X2) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f323])).
fof(f7998,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f416,f824])).
fof(f7987,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f824,f420])).
fof(f7986,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f824,f419])).
fof(f7985,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2))))) ) | $spl86),
  inference(superposition,[],[f824,f419])).
fof(f7984,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f824,f417])).
fof(f7983,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl86),
  inference(superposition,[],[f824,f417])).
fof(f7570,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl196),
  inference(cnf_transformation,[],[f7570_D])).
fof(f7570_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f8245,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl86),
  inference(forward_demodulation,[],[f7959,f479])).
fof(f7959,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl86),
  inference(superposition,[],[f824,f429])).
fof(f8244,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl86),
  inference(forward_demodulation,[],[f7958,f479])).
fof(f7958,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl86),
  inference(superposition,[],[f824,f429])).
fof(f7565,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl194),
  inference(cnf_transformation,[],[f7565_D])).
fof(f7565_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f8240,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl86),
  inference(forward_demodulation,[],[f7948,f480])).
fof(f7948,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl86),
  inference(superposition,[],[f824,f429])).
fof(f7915,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f7914,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f7890,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f7889,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f798,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl57),
  inference(cnf_transformation,[],[f798_D])).
fof(f798_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f7680,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X1)) ) | $spl208),
  inference(cnf_transformation,[],[f7680_D])).
fof(f7680_D,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X1))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X1)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f7515,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X0))))) | s('type.num.num',sK6) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK5))))) ) | $spl87),
  inference(superposition,[],[f825,f417])).
fof(f7516,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X1))))) | s('type.num.num',sK6) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl87),
  inference(superposition,[],[f825,f417])).
fof(f7679,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X3)))) | s('type.num.num',sK5) != s('type.num.num',X3)) ) | $spl206),
  inference(cnf_transformation,[],[f7679_D])).
fof(f7679_D,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X3)))) | s('type.num.num',sK5) != s('type.num.num',X3)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f7517,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | s('type.num.num',sK5) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2))))) ) | $spl87),
  inference(superposition,[],[f825,f419])).
fof(f7518,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK6))))) | s('type.num.num',sK5) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl87),
  inference(superposition,[],[f825,f419])).
fof(f7678,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) | $spl204),
  inference(cnf_transformation,[],[f7678_D])).
fof(f7678_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7519,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl87),
  inference(superposition,[],[f825,f420])).
fof(f7671,plain,(
  $false | ($spl1 | $spl57)),
  inference(subsumption_resolution,[],[f7586,f798])).
fof(f7586,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl1),
  inference(resolution,[],[f833,f512])).
fof(f7538,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f420])).
fof(f7537,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f419])).
fof(f7536,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2))))) ) | $spl87),
  inference(superposition,[],[f825,f419])).
fof(f7535,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f417])).
fof(f7534,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl87),
  inference(superposition,[],[f825,f417])).
fof(f7571,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl197),
  inference(cnf_transformation,[],[f7571_D])).
fof(f7571_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f7572,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl87),
  inference(forward_demodulation,[],[f7510,f479])).
fof(f7510,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl87),
  inference(superposition,[],[f825,f429])).
fof(f7569,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl87),
  inference(forward_demodulation,[],[f7509,f479])).
fof(f7509,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl87),
  inference(superposition,[],[f825,f429])).
fof(f7564,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl87),
  inference(forward_demodulation,[],[f7499,f480])).
fof(f7499,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl87),
  inference(superposition,[],[f825,f429])).
fof(f7208,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f417])).
fof(f7209,plain,(
  ( ! [X1] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X1))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f417])).
fof(f7210,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X2))))) ) | $spl9),
  inference(superposition,[],[f718,f419])).
fof(f7211,plain,(
  ( ! [X3] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',sK4) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f419])).
fof(f7212,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num',X4) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4))))) ) | $spl9),
  inference(superposition,[],[f718,f420])).
fof(f7269,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) | $spl185),
  inference(cnf_transformation,[],[f7269_D])).
fof(f7269_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f7263,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | $spl177),
  inference(cnf_transformation,[],[f7263_D])).
fof(f7263_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f7215,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) | $spl9),
  inference(superposition,[],[f718,f433])).
fof(f7267,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) ) | $spl182),
  inference(cnf_transformation,[],[f7267_D])).
fof(f7267_D,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f7222,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) ) | $spl9),
  inference(superposition,[],[f718,f429])).
fof(f7266,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num',X18)))) ) | $spl180),
  inference(cnf_transformation,[],[f7266_D])).
fof(f7266_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num',X18)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f7223,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num',X18)))) ) | $spl9),
  inference(superposition,[],[f718,f429])).
fof(f7265,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | $spl179),
  inference(cnf_transformation,[],[f7265_D])).
fof(f7265_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f7226,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | $spl9),
  inference(superposition,[],[f718,f464])).
fof(f7261,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) | $spl174),
  inference(cnf_transformation,[],[f7261_D])).
fof(f7261_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7227,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) | $spl9),
  inference(superposition,[],[f718,f474])).
fof(f709,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl0),
  inference(cnf_transformation,[],[f709_D])).
fof(f709_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f710,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,X9))) ) | $spl2),
  inference(cnf_transformation,[],[f710_D])).
fof(f710_D,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,X9))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f5909,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) )),
  inference(superposition,[],[f391,f436])).
fof(f436,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X20,X21] : ((p(s(bool,X20)) => p(s(bool,X21))) => ((p(s(bool,X21)) => p(s(bool,X20))) => s(bool,X20) = s(bool,X21)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.IMP_ANTISYM_AX')).
fof(f5910,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | p(s(bool,X10))) )),
  inference(superposition,[],[f391,f436])).
fof(f708,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl1),
  inference(cnf_transformation,[],[f708_D])).
fof(f708_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f711,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,X11))) ) | $spl4),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,X11))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5911,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f391,f439])).
fof(f439,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f271])).
fof(f5912,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) )),
  inference(superposition,[],[f391,f439])).
fof(f714,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl7),
  inference(cnf_transformation,[],[f714_D])).
fof(f714_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f7259,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f5913,f459])).
fof(f5913,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f467])).
fof(f7258,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f5914,f459])).
fof(f5914,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f469])).
fof(f831,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | s('type.num.num',sK4) != s('type.num.num',X17)) ) | $spl94),
  inference(cnf_transformation,[],[f831_D])).
fof(f831_D,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | s('type.num.num',sK4) != s('type.num.num',X17)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5997,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s('type.num.num',sK4) != s('type.num.num',X17) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f451])).
fof(f451,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | ~p(s(bool,X2)) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X12] : ! [X19,X40,X18,X41,X42,X43] : ((s(bool,X19) = s(bool,X40) & (p(s(bool,X40)) => s(X12,X18) = s(X12,X41)) & (~p(s(bool,X40)) => s(X12,X42) = s(X12,X43))) => s(X12,'const.bool.COND_3'(s(bool,X19),s(X12,X18),s(X12,X42))) = s(X12,'const.bool.COND_3'(s(bool,X40),s(X12,X41),s(X12,X43))))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.bool.COND_CONG')).
fof(f832,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) ) | $spl96),
  inference(cnf_transformation,[],[f832_D])).
fof(f832_D,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f5999,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f452])).
fof(f452,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | p(s(bool,X2)) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f277])).
fof(f6990,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl61),
  inference(forward_demodulation,[],[f6942,f478])).
fof(f6942,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl61),
  inference(superposition,[],[f802,f429])).
fof(f6992,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl61),
  inference(forward_demodulation,[],[f6943,f478])).
fof(f6943,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl61),
  inference(superposition,[],[f802,f429])).
fof(f6994,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl61),
  inference(forward_demodulation,[],[f6946,f478])).
fof(f6946,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl61),
  inference(superposition,[],[f802,f464])).
fof(f6996,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl61),
  inference(forward_demodulation,[],[f6947,f478])).
fof(f6947,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4)))) | $spl61),
  inference(superposition,[],[f802,f474])).
fof(f7246,plain,(
  $false | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f7245,f709])).
fof(f7245,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl9),
  inference(subsumption_resolution,[],[f7196,f479])).
fof(f7196,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl9),
  inference(superposition,[],[f718,f433])).
fof(f7244,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) | $spl172),
  inference(cnf_transformation,[],[f7244_D])).
fof(f7244_D,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7193,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl9),
  inference(superposition,[],[f718,f420])).
fof(f7243,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2)))) | s('type.num.num',sK5) != s('type.num.num',X2)) ) | $spl170),
  inference(cnf_transformation,[],[f7243_D])).
fof(f7243_D,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2)))) | s('type.num.num',sK5) != s('type.num.num',X2)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7192,plain,(
  ( ! [X3] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK6))))) | s('type.num.num',sK5) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl9),
  inference(superposition,[],[f718,f419])).
fof(f7191,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK6))))) | s('type.num.num',sK5) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X2))))) ) | $spl9),
  inference(superposition,[],[f718,f419])).
fof(f7242,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X0)) ) | $spl168),
  inference(cnf_transformation,[],[f7242_D])).
fof(f7242_D,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X0)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f7190,plain,(
  ( ! [X1] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X1))))) | s('type.num.num',sK6) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl9),
  inference(superposition,[],[f718,f417])).
fof(f7189,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X0))))) | s('type.num.num',sK6) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK5))))) ) | $spl9),
  inference(superposition,[],[f718,f417])).
fof(f7240,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl167),
  inference(cnf_transformation,[],[f7240_D])).
fof(f7240_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f7241,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl9),
  inference(forward_demodulation,[],[f7184,f479])).
fof(f7184,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl9),
  inference(superposition,[],[f718,f429])).
fof(f7238,plain,(
  ( ! [X12] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl9),
  inference(forward_demodulation,[],[f7183,f479])).
fof(f7183,plain,(
  ( ! [X12] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl9),
  inference(superposition,[],[f718,f429])).
fof(f7131,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl33),
  inference(superposition,[],[f1261,f429])).
fof(f1261,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | $spl33),
  inference(resolution,[],[f758,f407])).
fof(f7130,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl33),
  inference(superposition,[],[f1261,f429])).
fof(f7094,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl33),
  inference(superposition,[],[f1258,f429])).
fof(f1258,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X16,X17)) ) | $spl33),
  inference(resolution,[],[f758,f403])).
fof(f7093,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl33),
  inference(superposition,[],[f1258,f429])).
fof(f7055,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl33),
  inference(superposition,[],[f1257,f429])).
fof(f1257,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X15)) ) | $spl33),
  inference(resolution,[],[f758,f402])).
fof(f7054,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl33),
  inference(superposition,[],[f1257,f429])).
fof(f7019,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl33),
  inference(superposition,[],[f1255,f429])).
fof(f1255,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl33),
  inference(resolution,[],[f758,f399])).
fof(f7018,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl33),
  inference(superposition,[],[f1255,f429])).
fof(f7001,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | $spl165),
  inference(cnf_transformation,[],[f7001_D])).
fof(f7001_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f7003,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl61),
  inference(forward_demodulation,[],[f7002,f478])).
fof(f7002,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl61),
  inference(forward_demodulation,[],[f6953,f449])).
fof(f6953,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl61),
  inference(superposition,[],[f802,f429])).
fof(f6999,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK4)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl61),
  inference(forward_demodulation,[],[f6998,f478])).
fof(f6998,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl61),
  inference(forward_demodulation,[],[f6952,f449])).
fof(f6952,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl61),
  inference(superposition,[],[f802,f429])).
fof(f6848,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl33),
  inference(superposition,[],[f1254,f429])).
fof(f1254,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X9)) ) | $spl33),
  inference(resolution,[],[f758,f398])).
fof(f6847,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl33),
  inference(superposition,[],[f1254,f429])).
fof(f6812,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl33),
  inference(superposition,[],[f1252,f429])).
fof(f1252,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X4,X5)) ) | $spl33),
  inference(resolution,[],[f758,f395])).
fof(f6811,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl33),
  inference(superposition,[],[f1252,f429])).
fof(f6777,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl33),
  inference(superposition,[],[f1251,f429])).
fof(f1251,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) ) | $spl33),
  inference(resolution,[],[f758,f393])).
fof(f6776,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl33),
  inference(superposition,[],[f1251,f429])).
fof(f6739,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f943,f429])).
fof(f6738,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f943,f429])).
fof(f6701,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f940,f429])).
fof(f6700,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f940,f429])).
fof(f6235,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK5)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK5))))) | $spl54),
  inference(superposition,[],[f796,f474])).
fof(f6234,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK5)),s('type.num.num','const.num.|48|_0')))) | $spl54),
  inference(superposition,[],[f796,f464])).
fof(f6231,plain,(
  ( ! [X13] : (s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK5)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK5)),s('type.num.num',X13)))) ) | $spl54),
  inference(superposition,[],[f796,f429])).
fof(f6230,plain,(
  ( ! [X12] : (s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK5))))) ) | $spl54),
  inference(superposition,[],[f796,f429])).
fof(f6199,plain,(
  ( ! [X35,X34] : (sP1(X34,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f939,f474])).
fof(f6195,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f939,f429])).
fof(f6194,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f939,f429])).
fof(f6163,plain,(
  ( ! [X35,X34] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f937,f474])).
fof(f6159,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f937,f429])).
fof(f6158,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f937,f429])).
fof(f6127,plain,(
  ( ! [X35,X34] : (sP2(X34,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f936,f474])).
fof(f6123,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f936,f429])).
fof(f6122,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f936,f429])).
fof(f6091,plain,(
  ( ! [X35,X34] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f934,f474])).
fof(f6087,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f934,f429])).
fof(f6086,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f934,f429])).
fof(f6056,plain,(
  ( ! [X35,X34] : (sP3(X34,X35,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f933,f474])).
fof(f6052,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f933,f429])).
fof(f6051,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f933,f429])).
fof(f5996,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))))),
  inference(superposition,[],[f391,f474])).
fof(f5995,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f5992,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f5991,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) )),
  inference(superposition,[],[f391,f429])).
fof(f5984,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))),
  inference(superposition,[],[f391,f474])).
fof(f5983,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f5980,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f391,f429])).
fof(f5979,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f5972,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),
  inference(superposition,[],[f391,f433])).
fof(f5970,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f430])).
fof(f5969,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f420])).
fof(f5968,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f419])).
fof(f5967,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2))))) )),
  inference(superposition,[],[f391,f419])).
fof(f5966,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f5965,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))))) | s('type.num.num',sK5) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f800,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))) | $spl59),
  inference(cnf_transformation,[],[f800_D])).
fof(f800_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f5964,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))),
  inference(superposition,[],[f391,f441])).
fof(f782,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl49),
  inference(cnf_transformation,[],[f782_D])).
fof(f782_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f6038,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f6037,f479])).
fof(f6037,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f6036,f480])).
fof(f6036,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f5963,f479])).
fof(f5963,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(superposition,[],[f391,f474])).
fof(f6035,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f6034,f479])).
fof(f6034,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f6033,f480])).
fof(f6033,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f5959,f479])).
fof(f5959,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f6032,plain,(
  ( ! [X12] : (s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f6031,f479])).
fof(f6031,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f6030,f480])).
fof(f6030,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f5958,f479])).
fof(f5958,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(superposition,[],[f391,f429])).
fof(f6029,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))),
  inference(forward_demodulation,[],[f5953,f480])).
fof(f5953,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))),
  inference(superposition,[],[f391,f474])).
fof(f6028,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f5952,f480])).
fof(f5952,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f6027,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(forward_demodulation,[],[f5949,f480])).
fof(f5949,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(superposition,[],[f391,f429])).
fof(f6026,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) )),
  inference(forward_demodulation,[],[f5948,f480])).
fof(f5948,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f749,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | $spl29),
  inference(cnf_transformation,[],[f749_D])).
fof(f749_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f6024,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f6023,f479])).
fof(f6023,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f6022,f478])).
fof(f6022,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f5936,f449])).
fof(f5936,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f6021,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f6020,f479])).
fof(f6020,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f6019,f478])).
fof(f6019,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f5935,f449])).
fof(f5935,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f429])).
fof(f6018,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f6017,f480])).
fof(f6017,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f5930,f478])).
fof(f5930,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(superposition,[],[f391,f474])).
fof(f6016,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f6015,f480])).
fof(f6015,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f5929,f478])).
fof(f5929,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f6014,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f6013,f480])).
fof(f6013,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f5926,f478])).
fof(f5926,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f6012,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f6011,f480])).
fof(f6011,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f5925,f478])).
fof(f5925,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(superposition,[],[f391,f429])).
fof(f5920,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f6010,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f5919,f460])).
fof(f5919,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f5918,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f6009,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f5917,f460])).
fof(f5917,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f790,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | $spl52),
  inference(cnf_transformation,[],[f790_D])).
fof(f790_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5676,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl0),
  inference(superposition,[],[f709,f429])).
fof(f5675,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl0),
  inference(superposition,[],[f709,f429])).
fof(f5567,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f5566,f519])).
fof(f5566,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl52),
  inference(subsumption_resolution,[],[f5500,f460])).
fof(f5500,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl52),
  inference(superposition,[],[f5096,f473])).
fof(f5096,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl52),
  inference(forward_demodulation,[],[f5095,f479])).
fof(f5095,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl52),
  inference(forward_demodulation,[],[f5094,f480])).
fof(f5094,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl52),
  inference(forward_demodulation,[],[f4871,f479])).
fof(f4871,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | $spl52),
  inference(backward_demodulation,[],[f4736,f391])).
fof(f4736,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK5) | $spl52),
  inference(resolution,[],[f790,f464])).
fof(f5565,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f5564,f519])).
fof(f5564,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl52),
  inference(subsumption_resolution,[],[f5499,f460])).
fof(f5499,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl52),
  inference(superposition,[],[f5096,f471])).
fof(f791,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | $spl53),
  inference(cnf_transformation,[],[f791_D])).
fof(f791_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f4735,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f4734,f479])).
fof(f4734,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f4733,f480])).
fof(f4733,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f2912,f479])).
fof(f2912,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f811,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5) | $spl73),
  inference(cnf_transformation,[],[f811_D])).
fof(f811_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f2921,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5)),
  inference(superposition,[],[f391,f431])).
fof(f4670,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl0),
  inference(superposition,[],[f709,f429])).
fof(f4671,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl0),
  inference(superposition,[],[f709,f429])).
fof(f4709,plain,(
  ( ! [X35,X34] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f937,f474])).
fof(f4705,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f937,f429])).
fof(f4704,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f937,f429])).
fof(f4125,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK5)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK5))))) | $spl54),
  inference(superposition,[],[f796,f474])).
fof(f4124,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK5)),s('type.num.num','const.num.|48|_0')))) | $spl54),
  inference(superposition,[],[f796,f464])).
fof(f4121,plain,(
  ( ! [X13] : (s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK5)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK5)),s('type.num.num',X13)))) ) | $spl54),
  inference(superposition,[],[f796,f429])).
fof(f4120,plain,(
  ( ! [X12] : (s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK5))))) ) | $spl54),
  inference(superposition,[],[f796,f429])).
fof(f4115,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl144),
  inference(cnf_transformation,[],[f4115_D])).
fof(f4115_D,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f4104,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f843,f429])).
fof(f4103,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f843,f429])).
fof(f4098,plain,(
  ( ! [X35,X34] : (sP1(X34,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f843,f474])).
fof(f4094,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f843,f429])).
fof(f4093,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f843,f429])).
fof(f1475,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl116),
  inference(cnf_transformation,[],[f1475_D])).
fof(f1475_D,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4078,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f4077,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f4072,plain,(
  ( ! [X35,X34] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f841,f474])).
fof(f4068,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f4067,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f4046,plain,(
  ( ! [X35,X34] : (sP2(X34,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f840,f474])).
fof(f4042,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f4041,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f4020,plain,(
  ( ! [X35,X34] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f838,f474])).
fof(f4016,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f4015,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f3995,plain,(
  ( ! [X35,X34] : (sP3(X34,X35,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f837,f474])).
fof(f3991,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f3990,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f1741,plain,(
  s('type.num.num',sK8(sK6)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK6))))) | $spl126),
  inference(cnf_transformation,[],[f1741_D])).
fof(f1741_D,plain,(
  s('type.num.num',sK8(sK6)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK6))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1734,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl120),
  inference(cnf_transformation,[],[f1734_D])).
fof(f1734_D,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f3525,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK6)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK6))))) | $spl34),
  inference(superposition,[],[f763,f474])).
fof(f1739,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK6)),s('type.num.num','const.num.|48|_0')))) | $spl125),
  inference(cnf_transformation,[],[f1739_D])).
fof(f1739_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK6)),s('type.num.num','const.num.|48|_0')))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f3524,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK6)),s('type.num.num','const.num.|48|_0')))) | $spl34),
  inference(superposition,[],[f763,f464])).
fof(f1735,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK6)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK6)),s('type.num.num',X13)))) ) | $spl122),
  inference(cnf_transformation,[],[f1735_D])).
fof(f1735_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK6)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK6)),s('type.num.num',X13)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f3521,plain,(
  ( ! [X13] : (s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK6)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK6)),s('type.num.num',X13)))) ) | $spl34),
  inference(superposition,[],[f763,f429])).
fof(f1732,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK6))))) ) | $spl118),
  inference(cnf_transformation,[],[f1732_D])).
fof(f1732_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK6))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f3520,plain,(
  ( ! [X12] : (s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK6))))) ) | $spl34),
  inference(superposition,[],[f763,f429])).
fof(f1134,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl99),
  inference(cnf_transformation,[],[f1134_D])).
fof(f1134_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f3392,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5)))) | $spl1),
  inference(superposition,[],[f833,f474])).
fof(f3388,plain,(
  ( ! [X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl1),
  inference(superposition,[],[f833,f429])).
fof(f3387,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f833,f429])).
fof(f3195,plain,(
  $false | ($spl32 | $spl34)),
  inference(subsumption_resolution,[],[f3090,f1508])).
fof(f1508,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) | $spl34),
  inference(superposition,[],[f463,f763])).
fof(f463,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X7] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_LEFT_SUB','thm.num.NOT_SUC')).
fof(f3090,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK6) | $spl32),
  inference(resolution,[],[f757,f464])).
fof(f757,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | $spl32),
  inference(cnf_transformation,[],[f757_D])).
fof(f757_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2949,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f452])).
fof(f2947,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s('type.num.num',sK4) != s('type.num.num',X17) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f451])).
fof(f2946,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))))),
  inference(superposition,[],[f391,f474])).
fof(f2945,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f2942,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2941,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) )),
  inference(superposition,[],[f391,f429])).
fof(f2934,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))),
  inference(superposition,[],[f391,f474])).
fof(f2933,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f2930,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2929,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f2922,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),
  inference(superposition,[],[f391,f433])).
fof(f2920,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f430])).
fof(f2919,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f420])).
fof(f2918,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f419])).
fof(f2917,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2))))) )),
  inference(superposition,[],[f391,f419])).
fof(f2916,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f2915,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))))) | s('type.num.num',sK5) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f2914,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))),
  inference(superposition,[],[f391,f441])).
fof(f2992,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f2991,f479])).
fof(f2991,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f2990,f480])).
fof(f2990,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f2913,f479])).
fof(f2913,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(superposition,[],[f391,f474])).
fof(f2989,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f2988,f479])).
fof(f2988,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f2987,f480])).
fof(f2987,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f2909,f479])).
fof(f2909,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2986,plain,(
  ( ! [X12] : (s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f2985,f479])).
fof(f2985,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f2984,f480])).
fof(f2984,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f2908,f479])).
fof(f2908,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2983,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))),
  inference(forward_demodulation,[],[f2903,f480])).
fof(f2903,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))),
  inference(superposition,[],[f391,f474])).
fof(f2982,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f2902,f480])).
fof(f2902,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f2981,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(forward_demodulation,[],[f2899,f480])).
fof(f2899,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2980,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) )),
  inference(forward_demodulation,[],[f2898,f480])).
fof(f2898,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f763,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) | $spl34),
  inference(cnf_transformation,[],[f763_D])).
fof(f763_D,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2978,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))),
  inference(forward_demodulation,[],[f2977,f479])).
fof(f2977,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))),
  inference(forward_demodulation,[],[f2976,f478])).
fof(f2976,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))),
  inference(forward_demodulation,[],[f2890,f449])).
fof(f2890,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))),
  inference(superposition,[],[f391,f474])).
fof(f758,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | $spl33),
  inference(cnf_transformation,[],[f758_D])).
fof(f758_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f2975,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f2974,f479])).
fof(f2974,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f2973,f478])).
fof(f2973,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f2889,f449])).
fof(f2889,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f2972,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f2971,f479])).
fof(f2971,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f2970,f478])).
fof(f2970,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f2886,f449])).
fof(f2886,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2969,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f2968,f479])).
fof(f2968,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f2967,f478])).
fof(f2967,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f2885,f449])).
fof(f2885,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2966,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f2965,f480])).
fof(f2965,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f2880,f478])).
fof(f2880,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(superposition,[],[f391,f474])).
fof(f2964,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f2963,f480])).
fof(f2963,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f2879,f478])).
fof(f2879,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f2962,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f2961,f480])).
fof(f2961,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f2876,f478])).
fof(f2876,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2960,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f2959,f480])).
fof(f2959,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f2875,f478])).
fof(f2875,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(superposition,[],[f391,f429])).
fof(f2870,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f2958,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f2869,f460])).
fof(f2869,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f2868,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f2957,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f2867,f460])).
fof(f2867,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f2956,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f2866,f460])).
fof(f2866,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f473])).
fof(f2955,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f2865,f460])).
fof(f2865,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f471])).
fof(f2954,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f2864,f459])).
fof(f2864,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f469])).
fof(f2953,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f2863,f459])).
fof(f2863,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f467])).
fof(f2862,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) )),
  inference(superposition,[],[f391,f439])).
fof(f2861,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f391,f439])).
fof(f2860,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | p(s(bool,X10))) )),
  inference(superposition,[],[f391,f436])).
fof(f2859,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) )),
  inference(superposition,[],[f391,f436])).
fof(f2834,plain,(
  $false | ($spl35 | $spl137)),
  inference(subsumption_resolution,[],[f2300,f2295])).
fof(f2295,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK6) | $spl35),
  inference(trivial_inequality_removal,[],[f2282])).
fof(f2282,plain,(
  s('type.num.num',sK6) != s('type.num.num',sK6) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK6) | $spl35),
  inference(superposition,[],[f762,f474])).
fof(f762,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) | $spl35),
  inference(cnf_transformation,[],[f762_D])).
fof(f762_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f2530,plain,(
  $false | ($spl0 | $spl35 | $spl53)),
  inference(subsumption_resolution,[],[f2305,f791])).
fof(f2305,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | ($spl0 | $spl35)),
  inference(backward_demodulation,[],[f2295,f709])).
fof(f2302,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num',X14) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))))) ) | $spl140),
  inference(cnf_transformation,[],[f2302_D])).
fof(f2302_D,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num',X14) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f2300,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) | $spl137),
  inference(cnf_transformation,[],[f2300_D])).
fof(f2300_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f2289,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num',X14) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))))) ) | $spl35),
  inference(superposition,[],[f762,f429])).
fof(f2301,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))),s('type.num.num',X15)))) ) | $spl138),
  inference(cnf_transformation,[],[f2301_D])).
fof(f2301_D,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))),s('type.num.num',X15)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f2290,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))),s('type.num.num',X15)))) ) | $spl35),
  inference(superposition,[],[f762,f429])).
fof(f2298,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8(sK6))))))) | $spl134),
  inference(cnf_transformation,[],[f2298_D])).
fof(f2298_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8(sK6))))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f2294,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.num.SUC_1'(s('type.num.num',sK8(sK6))))))) | $spl35),
  inference(superposition,[],[f762,f474])).
fof(f2296,plain,(
  $false | ($spl33 | $spl35)),
  inference(subsumption_resolution,[],[f2295,f1246])).
fof(f1246,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) | $spl33),
  inference(resolution,[],[f758,f465])).
fof(f465,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f333])).
fof(f1733,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl121),
  inference(cnf_transformation,[],[f1733_D])).
fof(f1733_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f2281,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK6)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK6))))) | $spl35),
  inference(superposition,[],[f762,f474])).
fof(f2280,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK6)),s('type.num.num','const.num.|48|_0')))) | $spl35),
  inference(superposition,[],[f762,f464])).
fof(f2277,plain,(
  ( ! [X13] : (s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK6)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK6)),s('type.num.num',X13)))) ) | $spl35),
  inference(superposition,[],[f762,f429])).
fof(f2276,plain,(
  ( ! [X12] : (s('type.num.num',sK6) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK6))))) ) | $spl35),
  inference(superposition,[],[f762,f429])).
fof(f2248,plain,(
  ( ! [X35,X34] : (sP2(X34,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f936,f474])).
fof(f2244,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f936,f429])).
fof(f2243,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f936,f429])).
fof(f2211,plain,(
  ( ! [X35,X34] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f934,f474])).
fof(f2207,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f934,f429])).
fof(f2206,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f934,f429])).
fof(f2175,plain,(
  ( ! [X35,X34] : (sP3(X34,X35,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl53),
  inference(superposition,[],[f933,f474])).
fof(f2171,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl53),
  inference(superposition,[],[f933,f429])).
fof(f2170,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl53),
  inference(superposition,[],[f933,f429])).
fof(f1485,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK8(sK6)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK8(sK6))))) | $spl34),
  inference(superposition,[],[f763,f474])).
fof(f1484,plain,(
  s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8(sK6)),s('type.num.num','const.num.|48|_0')))) | $spl34),
  inference(superposition,[],[f763,f464])).
fof(f1481,plain,(
  ( ! [X13] : (s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK8(sK6)) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK8(sK6)),s('type.num.num',X13)))) ) | $spl34),
  inference(superposition,[],[f763,f429])).
fof(f1480,plain,(
  ( ! [X12] : (s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK8(sK6))))) ) | $spl34),
  inference(superposition,[],[f763,f429])).
fof(f1471,plain,(
  ( ! [X35,X34] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X34,X35) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))) ) | $spl1),
  inference(superposition,[],[f841,f474])).
fof(f1467,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f1466,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f1461,plain,(
  ( ! [X35,X34] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f841,f474])).
fof(f1457,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f1456,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f841,f429])).
fof(f1447,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl112),
  inference(cnf_transformation,[],[f1447_D])).
fof(f1447_D,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1443,plain,(
  ( ! [X35,X34] : (sP2(X34,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X35) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))) ) | $spl1),
  inference(superposition,[],[f840,f474])).
fof(f1439,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f1438,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f1433,plain,(
  ( ! [X35,X34] : (sP2(X34,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f840,f474])).
fof(f1429,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f1428,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f840,f429])).
fof(f1419,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl108),
  inference(cnf_transformation,[],[f1419_D])).
fof(f1419_D,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1415,plain,(
  ( ! [X35,X34] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X34,X35) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))) ) | $spl1),
  inference(superposition,[],[f838,f474])).
fof(f1411,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f1410,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f1405,plain,(
  ( ! [X35,X34] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X34,X35) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f838,f474])).
fof(f1401,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f1400,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f838,f429])).
fof(f1391,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) ) | $spl104),
  inference(cnf_transformation,[],[f1391_D])).
fof(f1391_D,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1388,plain,(
  ( ! [X35,X34] : (sP3(X34,X35,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))) ) | $spl1),
  inference(superposition,[],[f837,f474])).
fof(f1384,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f1383,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f1378,plain,(
  ( ! [X35,X34] : (sP3(X34,X35,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))) ) | $spl1),
  inference(superposition,[],[f837,f474])).
fof(f1374,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f1373,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f837,f429])).
fof(f1228,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | $spl100),
  inference(cnf_transformation,[],[f1228_D])).
fof(f1228_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1202,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) | $spl1),
  inference(superposition,[],[f834,f474])).
fof(f1201,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f834,f464])).
fof(f1198,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl1),
  inference(superposition,[],[f834,f429])).
fof(f1197,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f834,f429])).
fof(f1108,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5)))) | $spl1),
  inference(superposition,[],[f833,f474])).
fof(f1104,plain,(
  ( ! [X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl1),
  inference(superposition,[],[f833,f429])).
fof(f1103,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f833,f429])).
fof(f905,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6)))) | $spl1),
  inference(superposition,[],[f708,f474])).
fof(f904,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f708,f464])).
fof(f901,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl1),
  inference(superposition,[],[f708,f429])).
fof(f900,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f708,f429])).
fof(f704,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f452])).
fof(f702,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s('type.num.num',sK4) != s('type.num.num',X17) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f451])).
fof(f701,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))))),
  inference(superposition,[],[f391,f474])).
fof(f700,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f697,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f696,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))))) )),
  inference(superposition,[],[f391,f429])).
fof(f689,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))))),
  inference(superposition,[],[f391,f474])).
fof(f688,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f685,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f391,f429])).
fof(f684,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f677,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))),
  inference(superposition,[],[f391,f433])).
fof(f676,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5)),
  inference(superposition,[],[f391,f431])).
fof(f675,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f430])).
fof(f674,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f420])).
fof(f673,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',sK5))))) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f419])).
fof(f672,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',sK5))))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',X2))))) )),
  inference(superposition,[],[f391,f419])).
fof(f671,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X1))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f670,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X0))))) | s('type.num.num',sK5) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))) )),
  inference(superposition,[],[f391,f417])).
fof(f669,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))),
  inference(superposition,[],[f391,f441])).
fof(f794,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f793,f479])).
fof(f793,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f792,f480])).
fof(f792,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(forward_demodulation,[],[f668,f479])).
fof(f668,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK5))))),
  inference(superposition,[],[f391,f474])).
fof(f789,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f788,f479])).
fof(f788,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f787,f480])).
fof(f787,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f667,f479])).
fof(f667,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f785,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f784,f479])).
fof(f784,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f783,f480])).
fof(f783,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f664,f479])).
fof(f664,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f779,plain,(
  ( ! [X12] : (s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f778,f479])).
fof(f778,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f777,f480])).
fof(f777,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f663,f479])).
fof(f663,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(superposition,[],[f391,f429])).
fof(f774,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))),
  inference(forward_demodulation,[],[f658,f480])).
fof(f658,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8('const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))))))),
  inference(superposition,[],[f391,f474])).
fof(f771,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f657,f480])).
fof(f657,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f769,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(forward_demodulation,[],[f654,f480])).
fof(f654,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X14)))) )),
  inference(superposition,[],[f391,f429])).
fof(f765,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) )),
  inference(forward_demodulation,[],[f653,f480])).
fof(f653,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) )),
  inference(superposition,[],[f391,f429])).
fof(f761,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))),
  inference(forward_demodulation,[],[f760,f479])).
fof(f760,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))),
  inference(forward_demodulation,[],[f759,f478])).
fof(f759,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))),
  inference(forward_demodulation,[],[f645,f449])).
fof(f645,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num',sK6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK6))))),
  inference(superposition,[],[f391,f474])).
fof(f756,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f755,f479])).
fof(f755,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f754,f478])).
fof(f754,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f644,f449])).
fof(f644,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f752,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f751,f479])).
fof(f751,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f750,f478])).
fof(f750,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f641,f449])).
fof(f641,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f746,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f745,f479])).
fof(f745,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f744,f478])).
fof(f744,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f640,f449])).
fof(f640,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f391,f429])).
fof(f741,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f740,f480])).
fof(f740,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(forward_demodulation,[],[f635,f478])).
fof(f635,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(sK4))))),
  inference(superposition,[],[f391,f474])).
fof(f737,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f736,f480])).
fof(f736,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f634,f478])).
fof(f634,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f391,f464])).
fof(f734,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f733,f480])).
fof(f733,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f631,f478])).
fof(f631,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(superposition,[],[f391,f429])).
fof(f729,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f728,f480])).
fof(f728,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f630,f478])).
fof(f630,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(superposition,[],[f391,f429])).
fof(f625,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f727,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f624,f460])).
fof(f624,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f476])).
fof(f623,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f720,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f622,f460])).
fof(f622,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f391,f475])).
fof(f719,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f621,f460])).
fof(f621,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f473])).
fof(f716,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f620,f460])).
fof(f620,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f471])).
fof(f715,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f619,f459])).
fof(f619,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f469])).
fof(f712,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f618,f459])).
fof(f618,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f391,f467])).
fof(f617,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) )),
  inference(superposition,[],[f391,f439])).
fof(f616,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f391,f439])).
fof(f615,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | p(s(bool,X10))) )),
  inference(superposition,[],[f391,f436])).
fof(f614,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',sK5))))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) )),
  inference(superposition,[],[f391,f436])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 438
% Passive clauses: 9868
% Generated clauses: 30218
% Final active clauses: 263
% Final passive clauses: 5652
% Input formulas: 133
% Initial clauses: 218
% 
% Selected by SInE selection: 104
% SInE iterations: 1
% 
% Duplicate literals: 118
% Trivial inequalities: 154
% Fw subsumption resolutions: 1909
% Fw demodulations: 7434
% Bw demodulations: 69
% 
% Simple tautologies: 399
% Equational tautologies: 2
% Forward subsumptions: 9159
% Fw demodulations to eq. taut.: 60
% 
% Binary resolution: 3474
% Forward superposition: 9971
% Backward superposition: 6287
% 
% Split clauses: 652
% Split components: 1304
% SAT solver clauses: 714
% SAT solver unit clauses: 29
% SAT solver binary clauses: 211
% 
% Sat splits: 652
% Sat splitting refutations: 61
% 
% Memory used [KB]: 11513
% Time elapsed: 0.750 s
% ------------------------------
% ------------------------------
% Success in time 12.5 s
