Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f7972,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7971,f7923,f29,f49,f30,f51,f31,f53,f32,f55,f33,f57,f47,f34,f58,f56,f54,f52,f50,f48,f7967,f7965,f7963,f7969])).
fof(f7969,plain,(
  $false | ($spl0 | $spl11)),
  inference(trivial_inequality_removal,[],[f7968])).
fof(f7968,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) | ($spl0 | $spl11)),
  inference(forward_demodulation,[],[f58,f47])).
fof(f7963,plain,(
  $false | ($spl0 | $spl9)),
  inference(trivial_inequality_removal,[],[f7962])).
fof(f7962,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f56,f47])).
fof(f7965,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f7964])).
fof(f7964,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f50,f47])).
fof(f7967,plain,(
  $false | ($spl0 | $spl7)),
  inference(trivial_inequality_removal,[],[f7966])).
fof(f7966,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) | ($spl0 | $spl7)),
  inference(forward_demodulation,[],[f54,f47])).
fof(f48,plain,(
  s('type.ring.ring'(sK0),sK1) != s('type.ring.ring'(sK0),sK2) | $spl1),
  inference(cnf_transformation,[],[f48_D])).
fof(f48_D,plain,(
  s('type.ring.ring'(sK0),sK1) != s('type.ring.ring'(sK0),sK2) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f50,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) | $spl3),
  inference(cnf_transformation,[],[f50_D])).
fof(f50_D,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f52,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) | $spl5),
  inference(cnf_transformation,[],[f52_D])).
fof(f52_D,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f54,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) | $spl7),
  inference(cnf_transformation,[],[f54_D])).
fof(f54_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f56,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK2))) | $spl9),
  inference(cnf_transformation,[],[f56_D])).
fof(f56_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f58,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))) | $spl11),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f34,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK2))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) | s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) | s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) | s('type.ring.ring'(sK0),sK1) != s('type.ring.ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  (s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2) | (s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) & s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) & s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) & s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK2))) & s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))))) & (s('type.ring.ring'(sK0),sK1) != s('type.ring.ring'(sK0),sK2) | s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) | s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK2))) | s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f24])).
fof(f24,plain,(
  ? [X0,X1,X2] : ((s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))) & (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ? [X0,X1,X2] : ((s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))) & (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ? [X0,X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <~> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ~! [X0,X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <=> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <=> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.ring.ring'(X6),X7) = s('type.ring.ring'(X6),X8) <=> (s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X8))) & s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.ring.ring'(X6),X7) = s('type.ring.ring'(X6),X8) <=> (s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X8))) & s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X8)))))),
  file('thm.ring.ring_component_equality','thm.ring.ring_component_equality')).
fof(f47,plain,(
  s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2) | $spl0),
  inference(cnf_transformation,[],[f47_D])).
fof(f47_D,plain,(
  s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f57,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))) | $spl10),
  inference(cnf_transformation,[],[f57_D])).
fof(f57_D,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f33,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))) | s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f55,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK2))) | $spl8),
  inference(cnf_transformation,[],[f55_D])).
fof(f55_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f32,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK2))) | s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f53,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) | $spl6),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f31,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) | s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f51,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) | $spl4),
  inference(cnf_transformation,[],[f51_D])).
fof(f51_D,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f30,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) | s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f49,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) | $spl2),
  inference(cnf_transformation,[],[f49_D])).
fof(f49_D,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f29,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) | s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2)),
  inference(cnf_transformation,[],[f25])).
fof(f7923,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f7914,f48])).
fof(f7914,plain,(
  s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2) | ($spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f7913])).
fof(f7913,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) | s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) | s('type.ring.ring'(sK0),sK1) = s('type.ring.ring'(sK0),sK2) | ($spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(equality_resolution,[],[f2931])).
fof(f2931,plain,(
  ( ! [X0] : (s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),X0))) | s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),X0))) | s('type.ring.ring'(sK0),sK2) = s('type.ring.ring'(sK0),X0)) ) | ($spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f2930,f49])).
fof(f2930,plain,(
  ( ! [X0] : (s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),X0))) | s('type.ring.ring'(sK0),sK2) = s('type.ring.ring'(sK0),X0) | s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),X0)))) ) | ($spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f2929,f51])).
fof(f2929,plain,(
  ( ! [X0] : (s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),X0))) | s('type.ring.ring'(sK0),sK2) = s('type.ring.ring'(sK0),X0) | s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),X0))) | s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),X0)))) ) | ($spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f2928,f53])).
fof(f2928,plain,(
  ( ! [X0] : (s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),X0))) | s('type.ring.ring'(sK0),sK2) = s('type.ring.ring'(sK0),X0) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),X0))) | s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),X0))) | s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),X0)))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f2913,f57])).
fof(f2913,plain,(
  ( ! [X0] : (s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),sK1))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),X0))) | s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),sK2))) != s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),X0))) | s('type.ring.ring'(sK0),sK2) = s('type.ring.ring'(sK0),X0) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),sK2))) != s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),X0))) | s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK2))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),X0))) | s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),sK2))) != s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),X0)))) ) | $spl8),
  inference(superposition,[],[f43,f55])).
fof(f43,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))) | s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : ((s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))) & (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))) | s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : ((s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))) & ((s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))) | s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2)))),
  inference(nnf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <=> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ! [X0] : ! [X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <=> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X8,X7] : (s('type.ring.ring'(X6),X7) = s('type.ring.ring'(X6),X8) <=> (s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X8))) & s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X8))) & s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X8)))))),
  file('thm.ring.ring_component_equality','thm.ring.ring_63_prove_ltb')).
fof(f7971,plain,(
  $false | ($spl0 | $spl5)),
  inference(trivial_inequality_removal,[],[f7970])).
fof(f7970,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f52,f47])).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 57
% Passive clauses: 1518
% Generated clauses: 7701
% Final active clauses: 45
% Final passive clauses: 1399
% Input formulas: 8
% Initial clauses: 18
% 
% Duplicate literals: 596
% Trivial inequalities: 98
% Fw subsumption resolutions: 698
% Fw demodulations: 391
% 
% Simple tautologies: 161
% Equational tautologies: 93
% Forward subsumptions: 3155
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 25
% Forward superposition: 3058
% Backward superposition: 2198
% Self superposition: 298
% Equality factoring: 237
% Equality resolution: 13
% 
% Split clauses: 390
% Split components: 789
% SAT solver clauses: 1029
% SAT solver unit clauses: 266
% SAT solver binary clauses: 306
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 390
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1025
% TWLsolver calls for satisfiability: 40
% 
% Memory used [KB]: 4221
% Time elapsed: 0.359 s
% ------------------------------
% ------------------------------
% Success in time 0.374 s
