Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (88868)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 743
% Passive clauses: 7525
% Generated clauses: 20072
% Final active clauses: 730
% Final passive clauses: 6435
% Input formulas: 37
% Initial clauses: 163
% 
% Duplicate literals: 664
% Trivial inequalities: 90
% Fw subsumption resolutions: 2105
% Fw demodulations: 788
% Bw demodulations: 13
% 
% Simple tautologies: 243
% Equational tautologies: 61
% Forward subsumptions: 8005
% 
% Binary resolution: 8617
% Factoring: 1378
% Forward superposition: 3996
% Backward superposition: 5391
% Self superposition: 113
% Equality factoring: 28
% Equality resolution: 25
% 
% Split clauses: 591
% Split components: 1194
% SAT solver clauses: 953
% SAT solver unit clauses: 20
% SAT solver binary clauses: 110
% 
% Sat splits: 591
% 
% TWLsolver clauses: 947
% TWLsolver calls for satisfiability: 414
% 
% Memory used [KB]: 11769
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f8460,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8459,f217,f392,f2263,f2265,f393,f395,f262,f409,f520,f524,f216,f387,f236,f402,f8358,f8360,f1753,f1766,f214,f390,f232,f400,f7203,f7207,f7205,f510,f512,f215,f389,f2040,f2064,f260,f411,f7254,f984,f987,f985,f224,f399,f231,f401,f1004,f3848,f3872,f1795,f1800,f261,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) ) | $spl24),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f261,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) != s(X0,X3)) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X8,X9,X10] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) <=> (s(X6,X9) = s(X6,X10) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10))))))),
  file('thm.relation.RC_lifts_monotonicities','thm.relation.RC_DEF')).
fof(f1800,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))) ) | $spl642),
  inference(cnf_transformation,[],[f1800_D])).
fof(f1800_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))) ) <=> ~$spl642),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl642])])).
fof(f1795,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))) ) | $spl24),
  inference(equality_resolution,[],[f410])).
fof(f3872,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) ) | $spl1348),
  inference(cnf_transformation,[],[f3872_D])).
fof(f3872_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) ) <=> ~$spl1348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1348])])).
fof(f3848,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2)))) ) | ($spl12 | $spl642)),
  inference(resolution,[],[f1800,f400])).
fof(f1004,plain,(
  $false | ($spl9 | $spl380)),
  inference(resolution,[],[f985,f395])).
fof(f401,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) | $spl14),
  inference(cnf_transformation,[],[f401_D])).
fof(f401_D,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f231,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.relation.RC_lifts_monotonicities','thm.bool.EQ_CLAUSES|split|1')).
fof(f399,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f399_D])).
fof(f399_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f224,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.RC_lifts_monotonicities',aHL_BOOL_CASES)).
fof(f985,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl380),
  inference(cnf_transformation,[],[f985_D])).
fof(f985_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f987,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl383),
  inference(cnf_transformation,[],[f987_D])).
fof(f987_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl383),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl383])])).
fof(f984,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1))) ) | ($spl10 | $spl14)),
  inference(subsumption_resolution,[],[f954,f401])).
fof(f954,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f401,f399])).
fof(f7254,plain,(
  $false | ($spl200 | $spl383 | $spl1348 | $spl2318)),
  inference(subsumption_resolution,[],[f7253,f987])).
fof(f7253,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl200 | $spl1348 | $spl2318)),
  inference(forward_demodulation,[],[f7212,f3872])).
fof(f7212,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))) | ($spl200 | $spl2318)),
  inference(backward_demodulation,[],[f7207,f524])).
fof(f411,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) ) | $spl26),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f260,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f2064,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X1),s(X0,X2)))) | s(X0,X1) = s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X1))),s(X0,X2)))) ) | $spl796),
  inference(cnf_transformation,[],[f2064_D])).
fof(f2064_D,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X1),s(X0,X2)))) | s(X0,X1) = s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X1))),s(X0,X2)))) ) <=> ~$spl796),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl796])])).
fof(f2040,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,X2) | ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X3),s(X0,X1),s(X0,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X1))),s(X0,X2)))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f411,f400])).
fof(f389,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) | $spl2),
  inference(cnf_transformation,[],[f389_D])).
fof(f389_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f215,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))))),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X5,X6] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6)))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6))))))) & (p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17)))) & ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f141])).
fof(f141,plain,(
  ? [X0,X1,X2] : (! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X5))),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))) & ? [X3,X4] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X3),s(X0,X4)))) & ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))))),
  inference(rectify,[],[f96])).
fof(f96,plain,(
  ? [X0,X1,X2] : (! [X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) & ? [X5,X6] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) & ~p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2] : (! [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) => ! [X5,X6] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) => p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2] : (! [X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,X3))),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X4))))))) => ! [X5,X6] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,X5),s(X0,X6)))) => p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),X1),s(X0,X6))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))) => ! [X9,X10] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,X9))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))) => ! [X9,X10] : (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,X9),s(X6,X10)))) => p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X6,X6),X7),s(X6,X9))),s(X6,happ(s(fun(X6,X6),X7),s(X6,X10))))))))),
  file('thm.relation.RC_lifts_monotonicities','thm.relation.RC_lifts_monotonicities')).
fof(f512,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | $spl198),
  inference(cnf_transformation,[],[f512_D])).
fof(f512_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f510,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16),s(sK13,sK17))) | ($spl2 | $spl12)),
  inference(resolution,[],[f400,f389])).
fof(f7205,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17))) | $spl2316),
  inference(cnf_transformation,[],[f7205_D])).
fof(f7205_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17))) <=> ~$spl2316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2316])])).
fof(f7207,plain,(
  s(sK13,sK16) = s(sK13,sK17) | $spl2318),
  inference(cnf_transformation,[],[f7207_D])).
fof(f7207_D,plain,(
  s(sK13,sK16) = s(sK13,sK17) <=> ~$spl2318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2318])])).
fof(f7203,plain,(
  s(sK13,sK16) = s(sK13,sK17) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17))) | ($spl6 | $spl198 | $spl796)),
  inference(subsumption_resolution,[],[f7145,f392])).
fof(f7145,plain,(
  ~p(s(bool,'T_0')) | s(sK13,sK16) = s(sK13,sK17) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,sK16))),s(sK13,sK17))) | ($spl198 | $spl796)),
  inference(superposition,[],[f2064,f512])).
fof(f400,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f400_D])).
fof(f400_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f232,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f145])).
fof(f390,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6))))) ) | $spl4),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f214,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X5))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X6)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X5))),s(sK13,X6))))) )),
  inference(cnf_transformation,[],[f142])).
fof(f1766,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X0))),s(sK13,X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X0))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X1)))))) ) | $spl616),
  inference(cnf_transformation,[],[f1766_D])).
fof(f1766_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X0))),s(sK13,X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X0))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X1)))))) ) <=> ~$spl616),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl616])])).
fof(f1753,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,X0))),s(sK13,X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X0))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,X1)))))) ) | ($spl4 | $spl12)),
  inference(resolution,[],[f390,f400])).
fof(f8360,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | $spl2580),
  inference(cnf_transformation,[],[f8360_D])).
fof(f8360_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) <=> ~$spl2580),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2580])])).
fof(f8358,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | ($spl6 | $spl616 | $spl2316)),
  inference(subsumption_resolution,[],[f8239,f392])).
fof(f8239,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | ($spl616 | $spl2316)),
  inference(superposition,[],[f1766,f7205])).
fof(f402,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f402_D])).
fof(f402_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f236,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.relation.RC_lifts_monotonicities','thm.bool.EQ_CLAUSES|split|2')).
fof(f387,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | $spl1),
  inference(cnf_transformation,[],[f387_D])).
fof(f387_D,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f216,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))))),
  inference(cnf_transformation,[],[f142])).
fof(f524,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | $spl200),
  inference(cnf_transformation,[],[f524_D])).
fof(f524_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f520,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RC_3'(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17))))) | ($spl1 | $spl16)),
  inference(resolution,[],[f402,f387])).
fof(f409,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) | $spl22),
  inference(cnf_transformation,[],[f409_D])).
fof(f409_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f262,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f395,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f395_D])).
fof(f395_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f393,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f220,f217])).
fof(f220,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.relation.RC_lifts_monotonicities','thm.bool.NOT_CLAUSES|split|2')).
fof(f2265,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | $spl903),
  inference(cnf_transformation,[],[f2265_D])).
fof(f2265_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) <=> ~$spl903),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl903])])).
fof(f2263,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | ($spl9 | $spl22 | $spl200)),
  inference(subsumption_resolution,[],[f2217,f395])).
fof(f2217,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK15),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK16))))),s(sK13,happ(s(fun(sK13,sK13),sK14),s(sK13,sK17)))))) | ($spl22 | $spl200)),
  inference(superposition,[],[f409,f524])).
fof(f392,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f392_D])).
fof(f392_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f217,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.RC_lifts_monotonicities','thm.bool.TRUTH')).
fof(f8459,plain,(
  $false | ($spl6 | $spl903 | $spl2580)),
  inference(subsumption_resolution,[],[f8458,f392])).
fof(f8458,plain,(
  ~p(s(bool,'T_0')) | ($spl903 | $spl2580)),
  inference(backward_demodulation,[],[f8360,f2265])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 373
% Passive clauses: 1503
% Generated clauses: 8503
% Final active clauses: 318
% Final passive clauses: 884
% Input formulas: 37
% Initial clauses: 163
% 
% Duplicate literals: 204
% Trivial inequalities: 30
% Fw subsumption resolutions: 916
% Fw demodulations: 537
% Bw demodulations: 288
% Condensations: 233
% Global subsumptions: 13
% 
% Simple tautologies: 230
% Equational tautologies: 17
% Forward subsumptions: 3006
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 2014
% Factoring: 264
% Forward superposition: 1059
% Backward superposition: 1090
% Self superposition: 28
% Equality factoring: 26
% Equality resolution: 13
% 
% Split clauses: 70
% Split components: 145
% SAT solver clauses: 1855
% SAT solver unit clauses: 181
% SAT solver binary clauses: 72
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 70
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 1802
% TWLsolver calls for satisfiability: 1716
% 
% Memory used [KB]: 7803
% Time elapsed: 0.373 s
% ------------------------------
% ------------------------------
% Success in time 0.988 s
