Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (100613)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 109
% Passive clauses: 4036
% Generated clauses: 10675
% Final active clauses: 107
% Final passive clauses: 3763
% Input formulas: 41
% Initial clauses: 104
% 
% Duplicate literals: 376
% Trivial inequalities: 165
% Fw subsumption resolutions: 1652
% Fw demodulations: 73
% Bw demodulations: 2
% 
% Simple tautologies: 262
% Equational tautologies: 70
% Forward subsumptions: 3163
% 
% Binary resolution: 898
% Factoring: 4
% Forward superposition: 3010
% Backward superposition: 3854
% Self superposition: 206
% Equality factoring: 166
% Equality resolution: 20
% 
% Split clauses: 402
% Split components: 845
% SAT solver clauses: 969
% SAT solver unit clauses: 40
% SAT solver binary clauses: 386
% 
% Sat splits: 402
% 
% TWLsolver clauses: 969
% TWLsolver calls for satisfiability: 191
% 
% Memory used [KB]: 7803
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (100642)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 191
% Passive clauses: 1981
% Generated clauses: 14525
% Final active clauses: 122
% Final passive clauses: 1046
% Input formulas: 41
% Initial clauses: 104
% 
% Duplicate literals: 390
% Trivial inequalities: 106
% Fw subsumption resolutions: 1557
% Fw demodulations: 1753
% Bw demodulations: 769
% Condensations: 597
% Global subsumptions: 13
% 
% Simple tautologies: 121
% Equational tautologies: 32
% Forward subsumptions: 6017
% Bw demodulations to eq. taut.: 40
% 
% Binary resolution: 730
% Factoring: 3
% Forward superposition: 2829
% Backward superposition: 3536
% Self superposition: 72
% Equality factoring: 52
% Equality resolution: 7
% 
% Split clauses: 50
% Split components: 110
% SAT solver clauses: 1909
% SAT solver unit clauses: 117
% SAT solver binary clauses: 50
% SAT solver learnt clauses: 26
% 
% Sat splits: 50
% Sat splitting refutations: 53
% 
% TWLsolver clauses: 1857
% TWLsolver calls for satisfiability: 575
% 
% Memory used [KB]: 13176
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% (100659)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1049
% Passive clauses: 16377
% Generated clauses: 142465
% Final active clauses: 250
% Final passive clauses: 2165
% Input formulas: 41
% Initial clauses: 104
% Discarded non-redundant clauses: 4901
% 
% Duplicate literals: 3976
% Trivial inequalities: 457
% Fw subsumption resolutions: 10766
% Fw demodulations: 25538
% Bw demodulations: 4220
% Condensations: 5913
% Global subsumptions: 477
% 
% Simple tautologies: 679
% Equational tautologies: 465
% Forward subsumptions: 76294
% Fw demodulations to eq. taut.: 24
% Bw demodulations to eq. taut.: 426
% 
% Unit resulting resolution: 4465
% Factoring: 29
% Forward superposition: 46693
% Backward superposition: 33774
% Self superposition: 347
% Equality factoring: 95
% Equality resolution: 56
% 
% Split clauses: 230
% Split components: 473
% SAT solver clauses: 17905
% SAT solver unit clauses: 89
% SAT solver binary clauses: 3946
% 
% Sat splits: 230
% Sat splitting refutations: 588
% 
% Memory used [KB]: 112450
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quantHeuristics
% (100964)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 414
% Passive clauses: 17017
% Generated clauses: 72707
% Final active clauses: 334
% Final passive clauses: 15229
% Input formulas: 41
% Initial clauses: 85
% 
% Selected by SInE selection: 34
% SInE iterations: 1
% 
% Duplicate literals: 1981
% Trivial inequalities: 1586
% Fw subsumption resolutions: 13269
% Fw demodulations: 3385
% Bw demodulations: 63
% 
% Simple tautologies: 1949
% Equational tautologies: 50
% Forward subsumptions: 29656
% 
% Binary resolution: 12071
% Factoring: 9
% Forward superposition: 20433
% Backward superposition: 18435
% Self superposition: 207
% Equality factoring: 155
% Equality resolution: 89
% 
% Split clauses: 4543
% Split components: 9599
% SAT solver clauses: 4644
% SAT solver unit clauses: 34
% SAT solver binary clauses: 2405
% 
% Sat splits: 4543
% Sat splitting refutations: 43
% 
% Memory used [KB]: 32238
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.quantHeuristics
% (101068)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 82
% Passive clauses: 1929
% Generated clauses: 9717
% Final active clauses: 71
% Final passive clauses: 1739
% Input formulas: 41
% Initial clauses: 104
% 
% Duplicate literals: 423
% Trivial inequalities: 234
% Fw subsumption resolutions: 1186
% Fw demodulations: 181
% Bw demodulations: 86
% Global subsumptions: 10
% 
% Simple tautologies: 487
% Equational tautologies: 46
% Forward subsumptions: 5070
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 203
% Forward superposition: 3914
% Backward superposition: 3125
% Self superposition: 94
% Equality factoring: 50
% Equality resolution: 10
% 
% Split clauses: 17
% Split components: 35
% SAT solver clauses: 2006
% SAT solver unit clauses: 74
% SAT solver binary clauses: 276
% 
% Sat splits: 17
% Sat splitting refutations: 48
% 
% Memory used [KB]: 11385
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.quantHeuristics
% (101080)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 223
% Passive clauses: 4141
% Generated clauses: 13385
% Final active clauses: 41
% Final passive clauses: 2076
% Input formulas: 41
% Initial clauses: 104
% 
% Duplicate literals: 567
% Trivial inequalities: 228
% Fw subsumption resolutions: 820
% Bw subsumption resolutions: 1552
% Fw demodulations: 266
% Bw demodulations: 10
% Condensations: 716
% 
% Simple tautologies: 449
% Equational tautologies: 13
% Forward subsumptions: 6127
% Backward subsumptions: 522
% 
% Binary resolution: 1813
% Factoring: 4
% Forward superposition: 3766
% Backward superposition: 1615
% Self superposition: 32
% Equality factoring: 23
% Equality resolution: 22
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 4247
% SAT solver unit clauses: 123
% SAT solver binary clauses: 235
% 
% Sat splits: 2
% Sat splitting refutations: 101
% 
% Memory used [KB]: 11769
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f1210,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f185,f291,f311,f310,f330,f333,f334,f328,f386,f385,f382,f381,f380,f379,f376,f375,f372,f371,f370,f369,f367,f365,f358,f357,f356,f355,f396,f400,f402,f387,f388,f398,f657,f654,f653,f652,f648,f641,f635,f634,f633,f631,f628,f627,f626,f622,f618,f616,f615,f613,f612,f606,f605,f669,f604,f603,f602,f692,f659,f660,f693,f690,f1100])).
fof(f1100,plain,(
  $false | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f353,f692,f721])).
fof(f721,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,X3))),s(sK2,X2)))) | s(sK1,X3) = s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,sK12(X3,'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X2)),sK3,sK0,sK1))))) )),
  inference(superposition,[],[f314,f279])).
fof(f279,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X25,X26,X27] : ! [X28,X29,X30] : s(X25,happ(s(fun(X26,X25),'_dst_x1x0_2'(s(fun(X26,fun(X27,X25)),X28),s(X27,X29))),s(X26,X30))) = s(X25,happ(s(fun(X27,X25),happ(s(fun(X26,fun(X27,X25)),X28),s(X26,X30))),s(X27,X29)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','_dst_x1x0_ax')).
fof(f314,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0))),s(sK1,X1)))) | s(sK1,X1) = s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,sK12(X1,'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0)),sK3,sK0,sK1))))) )),
  inference(resolution,[],[f184,f258])).
fof(f258,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK13(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK13(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6,X8] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X8,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X20] : (p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X20)))) => ? [X21] : s(X6,X20) = s(X6,happ(s(fun(X8,X6),X22),s(X8,X21)))))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.quantHeuristics.GUESS_REWRITES|split|4')).
fof(f184,plain,(
  ( ! [X5] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X5))))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X5] : p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X5)))))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f99])).
fof(f99,plain,(
  ? [X0,X1,X2,X3,X4] : (! [X5] : p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X2,bool)),X4),s(X2,X5)))))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(fun(X2,bool),bool),'const.bool.?|33|_0'),s(fun(X1,fun(X2,bool)),X4)))))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1,X2,X3,X4] : (! [X5] : p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X2,bool)),X4),s(X2,X5)))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(fun(X2,bool),bool),'const.bool.?|33|_0'),s(fun(X1,fun(X2,bool)),X4)))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ~! [X0,X1,X2] : ! [X3,X4] : (! [X5] : p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X2,bool)),X4),s(X2,X5)))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(fun(X2,bool),bool),'const.bool.?|33|_0'),s(fun(X1,fun(X2,bool)),X4)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10] : (! [X11] : p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x1x0_2'(s(fun(X7,fun(X8,bool)),X10),s(X8,X11)))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1_2'(s(fun(fun(X8,bool),bool),'const.bool.?|33|_0'),s(fun(X7,fun(X8,bool)),X10)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10] : (! [X11] : p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x1x0_2'(s(fun(X7,fun(X8,bool)),X10),s(X8,X11)))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1_2'(s(fun(fun(X8,bool),bool),'const.bool.?|33|_0'),s(fun(X7,fun(X8,bool)),X10)))))))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1')).
fof(f353,plain,(
  ( ! [X0] : (s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X0))) != s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f291,f260])).
fof(f260,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,sK13(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f690,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl12),
  inference(cnf_transformation,[],[f690_D])).
fof(f690_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f693,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(global_subsumption,[],[f185,f184,f313,f311,f310,f306,f305,f330,f333,f334,f314,f402,f657,f654,f653,f652,f648,f641,f635,f634,f633,f628,f627,f626,f622,f616,f615,f613,f612,f606,f605,f604,f603,f602,f659,f660,f665,f666])).
fof(f666,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(superposition,[],[f402,f193])).
fof(f193,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1',aHL_BOOL_CASES)).
fof(f665,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))))) ) | $spl8),
  inference(resolution,[],[f402,f254])).
fof(f254,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f305,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0)))))) )),
  inference(unit_resulting_resolution,[],[f186,f184,f254])).
fof(f186,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1',aHL_TRUTH)).
fof(f306,plain,(
  ( ! [X0,X1] : (s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0))))) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X1)))))) )),
  inference(unit_resulting_resolution,[],[f184,f184,f254])).
fof(f313,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0)))))) )),
  inference(unit_resulting_resolution,[],[f192,f192,f184,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f284])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X15,X18,X11,X19] : ((s(bool,X15) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X11) = s(bool,X19))) => ((p(s(bool,X15)) => p(s(bool,X11))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.bool.IMP_CONG')).
fof(f192,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1',aHL_FALSITY)).
fof(f660,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(resolution,[],[f402,f198])).
fof(f198,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f659,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(resolution,[],[f402,f196])).
fof(f196,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f692,plain,(
  p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))),s(sK2,sK5(happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))),sK2))))) | $spl14),
  inference(cnf_transformation,[],[f692_D])).
fof(f692_D,plain,(
  p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))),s(sK2,sK5(happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))),sK2))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f602,plain,(
  p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))),s(sK2,sK5(happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))),sK2))))) | $spl8),
  inference(unit_resulting_resolution,[],[f402,f236])).
fof(f236,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0))))) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0))))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0))))) & s(X0,sK6(X1,X0)) != s(X0,sK7(X1,X0))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s(X0,X2) = s(X0,X3))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X8] : ! [X10] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X8,bool),X10)))) <=> (? [X15] : p(s(bool,happ(s(fun(X8,bool),X10),s(X8,X15)))) & ! [X15,X11] : ((p(s(bool,happ(s(fun(X8,bool),X10),s(X8,X15)))) & p(s(bool,happ(s(fun(X8,bool),X10),s(X8,X11))))) => s(X8,X11) = s(X8,X15))))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.bool.EXISTS_UNIQUE_THM')).
fof(f603,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f402,f196])).
fof(f604,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f402,f198])).
fof(f669,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl11),
  inference(cnf_transformation,[],[f669_D])).
fof(f669_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f605,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f402,f199])).
fof(f199,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f606,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f402,f201])).
fof(f201,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f612,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f187,f402,f254])).
fof(f187,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','thm.bool.TRUTH')).
fof(f613,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f186,f402,f254])).
fof(f615,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f187,f402,f254])).
fof(f616,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f186,f402,f254])).
fof(f618,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f192,f192,f398,f402,f281])).
fof(f281,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f622,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f402,f192,f235,f402,f283])).
fof(f283,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f235,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.?|33|_0'),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.?|33|_0'),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.?|33|_0'),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X8] : ! [X24] : s(bool,'const.bool.?|33|_1'(s(fun(X8,bool),X24))) = s(bool,happ(s(fun(fun(X8,bool),bool),'const.bool.?|33|_0'),s(fun(X8,bool),X24)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','ARITY_0_1_const.bool.?|33|')).
fof(f626,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f402,f235,f402,f283])).
fof(f627,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f187,f305,f402,f283])).
fof(f628,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f186,f305,f402,f283])).
fof(f631,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f192,f192,f398,f402,f285])).
fof(f285,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f633,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f192,f187,f402,f286])).
fof(f286,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f183])).
fof(f634,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f192,f186,f402,f286])).
fof(f635,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f192,f402,f402,f286])).
fof(f641,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f192,f402,f402,f286])).
fof(f648,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f402,f192,f235,f402,f287])).
fof(f287,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f652,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f402,f235,f402,f287])).
fof(f653,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f187,f305,f402,f287])).
fof(f654,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f186,f305,f402,f287])).
fof(f657,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1)))))) | $spl8),
  inference(unit_resulting_resolution,[],[f192,f192,f402,f289])).
fof(f398,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl6),
  inference(cnf_transformation,[],[f398_D])).
fof(f398_D,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f388,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(resolution,[],[f291,f202])).
fof(f202,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f126])).
fof(f387,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(resolution,[],[f291,f200])).
fof(f200,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f125])).
fof(f402,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))))) | $spl8),
  inference(cnf_transformation,[],[f402_D])).
fof(f402_D,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f400,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))))) | $spl1),
  inference(forward_demodulation,[],[f399,f235])).
fof(f399,plain,(
  p(s(bool,happ(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK2,bool),happ(s(fun(sK1,fun(sK2,bool)),sK4),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))))) | $spl1),
  inference(forward_demodulation,[],[f352,f278])).
fof(f278,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X25,X27,X26] : ! [X28,X29,X30] : s(X25,happ(s(fun(X27,X25),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))) = s(X25,happ(s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X27,X25),X28),s(fun(X26,X27),X29))),s(X26,X30)))),
  file('thm.quantHeuristics.GUESS_RULES_EXISTS_UNIQUE|split|1','_dst_x0x1_ax')).
fof(f352,plain,(
  p(s(bool,happ(s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))),s(sK1,sK13('_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)),sK3,sK0,sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f291,f259])).
fof(f259,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK13(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f396,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl5),
  inference(cnf_transformation,[],[f396_D])).
fof(f396_D,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f355,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f291,f195])).
fof(f195,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f123])).
fof(f356,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f291,f197])).
fof(f197,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f124])).
fof(f357,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f291,f200])).
fof(f358,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f291,f202])).
fof(f365,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f192,f291,f251])).
fof(f251,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f367,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f192,f291,f251])).
fof(f369,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f192,f187,f187,f291,f286])).
fof(f370,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f192,f186,f187,f291,f286])).
fof(f371,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f192,f187,f186,f291,f286])).
fof(f372,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f192,f186,f186,f291,f286])).
fof(f375,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f187,f291,f187,f291,f286])).
fof(f376,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f186,f291,f187,f291,f286])).
fof(f379,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f187,f291,f186,f291,f286])).
fof(f380,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f186,f291,f186,f291,f286])).
fof(f381,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f192,f187,f291,f289])).
fof(f382,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f192,f186,f291,f289])).
fof(f385,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f187,f291,f291,f289])).
fof(f386,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))) | $spl1),
  inference(unit_resulting_resolution,[],[f186,f291,f291,f289])).
fof(f328,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f328_D])).
fof(f328_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f334,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f294,f305])).
fof(f294,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0)))))) )),
  inference(unit_resulting_resolution,[],[f184,f199])).
fof(f333,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f295,f305])).
fof(f295,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x1x0_2'(s(fun(sK1,fun(sK2,bool)),sK4),s(sK2,X0)))))) )),
  inference(unit_resulting_resolution,[],[f184,f201])).
fof(f330,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(backward_demodulation,[],[f305,f313])).
fof(f310,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f192,f192,f187,f184,f286])).
fof(f311,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f192,f192,f186,f184,f286])).
fof(f291,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)))))) | $spl1),
  inference(cnf_transformation,[],[f291_D])).
fof(f291_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f185,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1_2'(s(fun(fun(sK2,bool),bool),'const.bool.?|33|_0'),s(fun(sK1,fun(sK2,bool)),sK4))))))),
  inference(cnf_transformation,[],[f120])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 81
% Passive clauses: 112
% Generated clauses: 914
% Final active clauses: 77
% Final passive clauses: 90
% Input formulas: 41
% Initial clauses: 104
% 
% Duplicate literals: 8
% Trivial inequalities: 6
% Fw demodulations: 155
% Bw demodulations: 9
% Global subsumptions: 2
% 
% Simple tautologies: 48
% Equational tautologies: 8
% Forward subsumptions: 508
% 
% Binary resolution: 115
% Unit resulting resolution: 372
% Forward superposition: 79
% Backward superposition: 160
% 
% SAT solver clauses: 277
% SAT solver unit clauses: 16
% SAT solver binary clauses: 189
% 
% Sat splitting refutations: 1
% 
% Memory used [KB]: 2430
% Time elapsed: 0.064 s
% ------------------------------
% ------------------------------
% Success in time 18.09 s
