Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.semi_ring
% (142891)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 201
% Passive clauses: 4152
% Generated clauses: 26685
% Final active clauses: 174
% Final passive clauses: 3698
% Input formulas: 132
% Initial clauses: 165
% 
% Duplicate literals: 650
% Trivial inequalities: 509
% Fw subsumption resolutions: 2140
% Fw demodulations: 4944
% Bw demodulations: 25
% 
% Simple tautologies: 484
% Equational tautologies: 43
% Forward subsumptions: 13307
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 390
% Factoring: 1
% Forward superposition: 8648
% Backward superposition: 8859
% Self superposition: 162
% Equality factoring: 73
% Equality resolution: 36
% 
% Split clauses: 243
% Split components: 486
% SAT solver clauses: 690
% SAT solver unit clauses: 39
% SAT solver binary clauses: 228
% 
% Sat splits: 243
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 690
% TWLsolver calls for satisfiability: 179
% 
% Memory used [KB]: 18933
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.semi_ring
% (142915)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 154
% Passive clauses: 1303
% Generated clauses: 19113
% Final active clauses: 132
% Final passive clauses: 915
% Input formulas: 132
% Initial clauses: 165
% 
% Duplicate literals: 626
% Trivial inequalities: 244
% Fw subsumption resolutions: 737
% Fw demodulations: 2555
% Bw demodulations: 228
% Condensations: 907
% Global subsumptions: 4
% 
% Simple tautologies: 308
% Equational tautologies: 40
% Forward subsumptions: 10193
% Bw demodulations to eq. taut.: 25
% 
% Binary resolution: 222
% Factoring: 1
% Forward superposition: 5865
% Backward superposition: 5985
% Self superposition: 136
% Equality factoring: 68
% Equality resolution: 21
% 
% Split clauses: 99
% Split components: 198
% SAT solver clauses: 2299
% SAT solver unit clauses: 264
% SAT solver binary clauses: 46
% SAT solver learnt clauses: 8
% 
% Sat splits: 99
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2283
% TWLsolver calls for satisfiability: 455
% 
% Memory used [KB]: 16758
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.semi_ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.semi_ring
% SZS output start Proof for thm.semi_ring
fof(f41075,plain,(
  $false),
  inference(subsumption_resolution,[],[f41074,f507])).
fof(f507,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X3),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X3),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3] : s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X3),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X3),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X3),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X3),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6] : ! [X17,X7,X8] : s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),X8),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),X7),s('type.semi_ring.semi_ring'(X6),X17))))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),X7),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),X8),s('type.semi_ring.semi_ring'(X6),X17)))))),
  file('thm.semi_ring.semi_ring_fupdcanon_comp|split|0','thm.semi_ring.semi_ring_39_prove_ltb')).
fof(f41074,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),sK2),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),sK1),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),sK1),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),sK2),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)))))))),
  inference(forward_demodulation,[],[f41073,f487])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))) = s('type.semi_ring.semi_ring'(X0),happ(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.semi_ring.semi_ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))) = s('type.semi_ring.semi_ring'(X0),happ(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.semi_ring.semi_ring'(X0),X2)))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0] : ! [X1,X2] : s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))) = s('type.semi_ring.semi_ring'(X0),happ(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.semi_ring.semi_ring'(X0),X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X11,X12] : s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),X11),s('type.semi_ring.semi_ring'(X6),X12))) = s('type.semi_ring.semi_ring'(X6),happ(s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X6,X6),X11))),s('type.semi_ring.semi_ring'(X6),X12)))),
  file('thm.semi_ring.semi_ring_fupdcanon_comp|split|0','ARITY_1_2_const.semi_ring.semi_ring_SR|49|_fupd')).
fof(f41073,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),sK1),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),sK2),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0))))))) != s('type.semi_ring.semi_ring'(sK0),happ(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),sK1),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)))))))),
  inference(forward_demodulation,[],[f41072,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))) = s('type.semi_ring.semi_ring'(X0),happ(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.semi_ring.semi_ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))) = s('type.semi_ring.semi_ring'(X0),happ(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.semi_ring.semi_ring'(X0),X2)))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0] : ! [X1,X2] : s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))) = s('type.semi_ring.semi_ring'(X0),happ(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.semi_ring.semi_ring'(X0),X2)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X11,X12] : s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),X11),s('type.semi_ring.semi_ring'(X6),X12))) = s('type.semi_ring.semi_ring'(X6),happ(s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X6,X6),X11))),s('type.semi_ring.semi_ring'(X6),X12)))),
  file('thm.semi_ring.semi_ring_fupdcanon_comp|split|0','ARITY_1_2_const.semi_ring.semi_ring_SR|48|_fupd')).
fof(f41072,plain,(
  s('type.semi_ring.semi_ring'(sK0),happ(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),sK1),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0))))))) != s('type.semi_ring.semi_ring'(sK0),happ(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),sK2),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)))))))),
  inference(forward_demodulation,[],[f41071,f24141])).
fof(f24141,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(X181,happ(s(fun('type.semi_ring.semi_ring'(X178),X181),X182),s('type.semi_ring.semi_ring'(X178),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X178,X178),X179),s('type.semi_ring.semi_ring'(X178),X180))))) = s(X181,'const.combin.o_3'(s(fun('type.semi_ring.semi_ring'(X178),X181),X182),s(fun('type.semi_ring.semi_ring'(X178),'type.semi_ring.semi_ring'(X178)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X178,X178),X179))),s('type.semi_ring.semi_ring'(X178),X180)))) )),
  inference(superposition,[],[f572,f486])).
fof(f572,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f334])).
fof(f334,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X9,X10,X6] : ! [X8,X7,X5] : s(X9,'const.combin.o_3'(s(fun(X10,X9),X8),s(fun(X6,X10),X7),s(X6,X5))) = s(X9,happ(s(fun(X10,X9),X8),s(X10,happ(s(fun(X6,X10),X7),s(X6,X5)))))),
  file('thm.semi_ring.semi_ring_fupdcanon_comp|split|0','thm.combin.o_DEF')).
fof(f41071,plain,(
  s('type.semi_ring.semi_ring'(sK0),happ(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),sK2),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0))))))) != s('type.semi_ring.semi_ring'(sK0),'const.combin.o_3'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)))))),
  inference(forward_demodulation,[],[f41070,f24142])).
fof(f24142,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(X186,happ(s(fun('type.semi_ring.semi_ring'(X183),X186),X187),s('type.semi_ring.semi_ring'(X183),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X183,X183),X184),s('type.semi_ring.semi_ring'(X183),X185))))) = s(X186,'const.combin.o_3'(s(fun('type.semi_ring.semi_ring'(X183),X186),X187),s(fun('type.semi_ring.semi_ring'(X183),'type.semi_ring.semi_ring'(X183)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X183,X183),X184))),s('type.semi_ring.semi_ring'(X183),X185)))) )),
  inference(superposition,[],[f572,f487])).
fof(f41070,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.combin.o_3'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0))))) != s('type.semi_ring.semi_ring'(sK0),'const.combin.o_3'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)))))),
  inference(forward_demodulation,[],[f41069,f574])).
fof(f574,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10,X6] : ! [X11,X12,X13] : s(X9,'const.combin.o_3'(s(fun(X10,X9),X11),s(fun(X6,X10),X12),s(X6,X13))) = s(X9,happ(s(fun(X6,X9),'const.combin.o_2'(s(fun(X10,X9),X11),s(fun(X6,X10),X12))),s(X6,X13)))),
  file('thm.semi_ring.semi_ring_fupdcanon_comp|split|0','ARITY_2_3_const.combin.o')).
fof(f41069,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.combin.o_3'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0))))) != s('type.semi_ring.semi_ring'(sK0),happ(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))))),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)))))),
  inference(forward_demodulation,[],[f40852,f574])).
fof(f40852,plain,(
  s('type.semi_ring.semi_ring'(sK0),happ(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))))),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0))))) != s('type.semi_ring.semi_ring'(sK0),happ(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))))),s('type.semi_ring.semi_ring'(sK0),sK10('const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1)))),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))),'type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)))))),
  inference(unit_resulting_resolution,[],[f428,f556])).
fof(f556,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK10(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK10(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f418])).
fof(f418,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X9,X6] : ! [X8,X7] : (s(fun(X6,X9),X7) = s(fun(X6,X9),X8) <=> ! [X14] : s(X9,happ(s(fun(X6,X9),X7),s(X6,X14))) = s(X9,happ(s(fun(X6,X9),X8),s(X6,X14))))),
  file('thm.semi_ring.semi_ring_fupdcanon_comp|split|0','thm.bool.FUN_EQ_THM')).
fof(f428,plain,(
  s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))))) != s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))))),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))))) != s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK0,sK0),sK1))),s(fun('type.semi_ring.semi_ring'(sK0),'type.semi_ring.semi_ring'(sK0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(sK0,sK0),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f376])).
fof(f376,plain,(
  ? [X0,X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))))) != s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ~! [X0] : ! [X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X0,X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X6,X6),X8))),s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X6,X6),X7))))) = s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X6,X6),X7))),s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X6,X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X6,X6),X8))),s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X6,X6),X7))))) = s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X6,X6),X7))),s(fun('type.semi_ring.semi_ring'(X6),'type.semi_ring.semi_ring'(X6)),'const.semi_ring.semi_ring_SR|49|_fupd_1'(s(fun(X6,X6),X8)))))),
  file('thm.semi_ring.semi_ring_fupdcanon_comp|split|0','thm.semi_ring.semi_ring_fupdcanon_comp|split|0')).
% SZS output end Proof for thm.semi_ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 200
% Passive clauses: 1868
% Generated clauses: 40427
% Final active clauses: 168
% Final passive clauses: 1550
% Input formulas: 132
% Initial clauses: 165
% 
% Duplicate literals: 1066
% Trivial inequalities: 286
% Fw subsumption resolutions: 1799
% Fw demodulations: 8658
% Bw demodulations: 148
% Condensations: 1317
% 
% Simple tautologies: 272
% Equational tautologies: 170
% Forward subsumptions: 24574
% Fw demodulations to eq. taut.: 146
% 
% Unit resulting resolution: 1358
% Factoring: 1
% Forward superposition: 12569
% Backward superposition: 12718
% Self superposition: 232
% Equality factoring: 85
% Equality resolution: 14
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 1939
% SAT solver unit clauses: 330
% SAT solver binary clauses: 963
% 
% Sat splits: 15
% Sat splitting refutations: 39
% 
% Memory used [KB]: 27760
% Time elapsed: 1.362 s
% ------------------------------
% ------------------------------
% Success in time 2.69 s
