Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (83231)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 65
% Passive clauses: 2041
% Generated clauses: 12877
% Final active clauses: 63
% Final passive clauses: 1920
% Input formulas: 25
% Initial clauses: 26
% 
% Duplicate literals: 815
% Trivial inequalities: 122
% Fw subsumption resolutions: 1212
% Fw demodulations: 319
% 
% Simple tautologies: 246
% Equational tautologies: 113
% Forward subsumptions: 6072
% 
% Binary resolution: 25
% Forward superposition: 5000
% Backward superposition: 4530
% Self superposition: 424
% Equality factoring: 338
% Equality resolution: 8
% 
% Split clauses: 547
% Split components: 1103
% SAT solver clauses: 1580
% SAT solver unit clauses: 376
% SAT solver binary clauses: 511
% 
% Sat splits: 547
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1576
% TWLsolver calls for satisfiability: 42
% 
% Memory used [KB]: 8955
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (83260)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 37
% Passive clauses: 229
% Generated clauses: 7625
% Final active clauses: 34
% Final passive clauses: 190
% Input formulas: 25
% Initial clauses: 26
% 
% Duplicate literals: 1304
% Trivial inequalities: 60
% Fw subsumption resolutions: 141
% Fw demodulations: 36
% Bw demodulations: 1
% Condensations: 2244
% 
% Simple tautologies: 122
% Equational tautologies: 77
% Forward subsumptions: 2108
% 
% Binary resolution: 16
% Forward superposition: 1713
% Backward superposition: 1511
% Self superposition: 268
% Equality factoring: 222
% Equality resolution: 5
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 601
% SAT solver unit clauses: 46
% SAT solver binary clauses: 66
% 
% Sat splits: 27
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 601
% TWLsolver calls for satisfiability: 57
% 
% Memory used [KB]: 6396
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f49860,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f49828])).
fof(f49828,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2)))),
  inference(superposition,[],[f77,f15402])).
fof(f15402,plain,(
  ( ! [X10,X8,X7,X9] : (s('type.canonical.canonical_sum'(X7),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X7),X10))) = s('type.canonical.canonical_sum'(X7),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X7),X8),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X7),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X7),X10)))))) )),
  inference(superposition,[],[f10553,f86])).
fof(f86,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X3),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X8,X9,X14] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X14),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8)))))),
  file('thm.ringNorm.canonical_sum_merge_def|split|6','thm.canonical.canonical_sum_merge_def|split|6')).
fof(f10553,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f10504,f91])).
fof(f91,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X27,X28,X29] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X27),s('type.canonical.canonical_sum'(X6),X28),s('type.canonical.canonical_sum'(X6),X29))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X27))),s('type.canonical.canonical_sum'(X6),X28))),s('type.canonical.canonical_sum'(X6),X29)))),
  file('thm.ringNorm.canonical_sum_merge_def|split|6','ARITY_1_3_const.ringNorm.r_canonical_sum_merge')).
fof(f10504,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f92,f81])).
fof(f81,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_merge_def|split|6','thm.ringNorm.r_canonical_sum_merge_def')).
fof(f92,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X27,X28,X29] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X27),s('type.canonical.canonical_sum'(X6),X28),s('type.canonical.canonical_sum'(X6),X29))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),X27))),s('type.canonical.canonical_sum'(X6),X28))),s('type.canonical.canonical_sum'(X6),X29)))),
  file('thm.ringNorm.canonical_sum_merge_def|split|6','ARITY_1_3_const.canonical.canonical_sum_merge')).
fof(f77,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2)))))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2)))))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8)))))),
  file('thm.ringNorm.canonical_sum_merge_def|split|6','thm.ringNorm.canonical_sum_merge_def|split|6')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 167
% Passive clauses: 4387
% Generated clauses: 55203
% Final active clauses: 87
% Final passive clauses: 750
% Input formulas: 25
% Initial clauses: 26
% Discarded non-redundant clauses: 511
% 
% Duplicate literals: 4356
% Trivial inequalities: 287
% Fw subsumption resolutions: 2311
% Fw demodulations: 3550
% Bw demodulations: 2615
% Condensations: 8763
% Global subsumptions: 36
% 
% Simple tautologies: 547
% Equational tautologies: 1074
% Forward subsumptions: 29754
% Bw demodulations to eq. taut.: 257
% 
% Unit resulting resolution: 2384
% Forward superposition: 13989
% Backward superposition: 10399
% Self superposition: 434
% Equality factoring: 278
% Equality resolution: 10
% 
% Split clauses: 174
% Split components: 355
% SAT solver clauses: 4805
% SAT solver unit clauses: 70
% SAT solver binary clauses: 396
% 
% Sat splits: 174
% Sat splitting refutations: 32
% 
% Memory used [KB]: 31598
% Time elapsed: 5.803 s
% ------------------------------
% ------------------------------
% Success in time 7.142 s
