Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (82721)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 138
% Passive clauses: 3179
% Generated clauses: 23856
% Final active clauses: 123
% Final passive clauses: 2760
% Input formulas: 28
% Initial clauses: 66
% 
% Duplicate literals: 948
% Trivial inequalities: 425
% Fw subsumption resolutions: 3870
% Fw demodulations: 3810
% Bw demodulations: 13
% 
% Simple tautologies: 498
% Equational tautologies: 52
% Forward subsumptions: 9761
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 566
% Factoring: 2
% Forward superposition: 6867
% Backward superposition: 6783
% Self superposition: 205
% Equality factoring: 133
% Equality resolution: 31
% 
% Split clauses: 560
% Split components: 1137
% SAT solver clauses: 1337
% SAT solver unit clauses: 21
% SAT solver binary clauses: 192
% 
% Sat splits: 560
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1332
% TWLsolver calls for satisfiability: 164
% 
% Memory used [KB]: 11769
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f7917,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1155,f118,f193,f227,f230,f126,f198,f221,f223,f1164,f1180,f302,f312,f171,f211,f7899,f1182,f125,f199,f240,f243,f165,f208,f260,f269,f270,f119,f194,f292,f176,f217,f689,f704,f122,f196,f174,f219,f1187,f1193,f316,f321,f168,f212,f115,f189,f187,f2127,f1205,f1228,f1464,f267,f1195,f1126,f1127,f907,f908,f121,f197,f294,f167,f209,f114,f188,f186])).
fof(f186,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | $spl0),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f188,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | $spl2),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f114,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))))),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3)) & (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f75])).
fof(f75,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <~> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3)))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X9)))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X9)))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  file('thm.pred_set.SURJ_EMPTY|split|0','thm.pred_set.SURJ_EMPTY|split|0')).
fof(f209,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl20),
  inference(cnf_transformation,[],[f209_D])).
fof(f209_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f167,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X9,X12] : (s(fun(X6,bool),X9) = s(fun(X6,bool),X12) <=> ! [X13] : s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X12))))),
  file('thm.pred_set.SURJ_EMPTY|split|0','thm.pred_set.EXTENSION')).
fof(f294,plain,(
  $false | ($spl7 | $spl54)),
  inference(resolution,[],[f267,f193])).
fof(f197,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) | $spl12),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f121,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.pred_set.SURJ_EMPTY|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f908,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X2,sK7(sK3,X0,X1,sK1,X2)),s(fun(X2,bool),X0)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X2,sK1),X1),s(fun(X2,bool),X0),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl238),
  inference(cnf_transformation,[],[f908_D])).
fof(f908_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X2,sK7(sK3,X0,X1,sK1,X2)),s(fun(X2,bool),X0)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X2,sK1),X1),s(fun(X2,bool),X0),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f907,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X2,sK1),X1),s(fun(X2,bool),X0),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(X2,sK7(sK3,X0,X1,sK1,X2)),s(fun(X2,bool),X0))))) ) | ($spl2 | $spl7 | $spl36 | $spl48)),
  inference(forward_demodulation,[],[f906,f188])).
fof(f906,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X2,sK1),X1),s(fun(X2,bool),X0),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(X2,sK7(sK3,X0,X1,sK1,X2)),s(fun(X2,bool),X0))))) ) | ($spl2 | $spl7 | $spl36 | $spl48)),
  inference(subsumption_resolution,[],[f905,f193])).
fof(f905,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X2,sK1),X1),s(fun(X2,bool),X0),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(X2,sK7(sK3,X0,X1,sK1,X2)),s(fun(X2,bool),X0))))) ) | ($spl2 | $spl36 | $spl48)),
  inference(forward_demodulation,[],[f892,f243])).
fof(f892,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK1,sK8(sK3,X0,X1,sK1,X2)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X2,sK1),X1),s(fun(X2,bool),X0),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(X2,sK7(sK3,X0,X1,sK1,X2)),s(fun(X2,bool),X0))))) ) | ($spl2 | $spl36)),
  inference(superposition,[],[f217,f188])).
fof(f1127,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,sK1),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) ) | $spl322),
  inference(cnf_transformation,[],[f1127_D])).
fof(f1127_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,sK1),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f1126,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,sK1),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl7 | $spl48 | $spl238)),
  inference(subsumption_resolution,[],[f1109,f193])).
fof(f1109,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,sK1),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl48 | $spl238)),
  inference(superposition,[],[f908,f243])).
fof(f1195,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,X11),s(fun(sK1,bool),sK3)))) ) | $spl354),
  inference(cnf_transformation,[],[f1195_D])).
fof(f1195_D,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,X11),s(fun(sK1,bool),sK3)))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f267,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) | $spl54),
  inference(cnf_transformation,[],[f267_D])).
fof(f267_D,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1464,plain,(
  ( ! [X45,X46] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,X45),s(fun(sK1,bool),sK3)))) ) | ($spl12 | $spl14 | $spl350)),
  inference(subsumption_resolution,[],[f1383,f197])).
fof(f1383,plain,(
  ( ! [X45,X46] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,X45),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,X46)) ) | ($spl14 | $spl350)),
  inference(superposition,[],[f198,f1193])).
fof(f1228,plain,(
  ( ! [X19] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4(X19,sK3,sK1)),s(fun(sK1,bool),X19))) | s(fun(sK1,bool),sK3) = s(fun(sK1,bool),X19)) ) | $spl356),
  inference(cnf_transformation,[],[f1228_D])).
fof(f1228_D,plain,(
  ( ! [X19] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4(X19,sK3,sK1)),s(fun(sK1,bool),X19))) | s(fun(sK1,bool),sK3) = s(fun(sK1,bool),X19)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f1205,plain,(
  ( ! [X19] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4(X19,sK3,sK1)),s(fun(sK1,bool),X19))) | s(fun(sK1,bool),sK3) = s(fun(sK1,bool),X19)) ) | ($spl20 | $spl354)),
  inference(superposition,[],[f209,f1195])).
fof(f2127,plain,(
  $false | ($spl3 | $spl48 | $spl356)),
  inference(subsumption_resolution,[],[f2126,f189])).
fof(f2126,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | ($spl48 | $spl356)),
  inference(trivial_inequality_removal,[],[f2119])).
fof(f2119,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | ($spl48 | $spl356)),
  inference(superposition,[],[f1228,f243])).
fof(f187,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) | $spl1),
  inference(cnf_transformation,[],[f187_D])).
fof(f187_D,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f189,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) | $spl3),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f115,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))))),
  inference(cnf_transformation,[],[f76])).
fof(f212,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | $spl26),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f168,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.SURJ_EMPTY|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f321,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl64),
  inference(cnf_transformation,[],[f321_D])).
fof(f321_D,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f316,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl10 | $spl26)),
  inference(resolution,[],[f212,f196])).
fof(f1193,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,'const.bool.IN_2'(s(sK1,X7),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,X8)) ) | $spl350),
  inference(cnf_transformation,[],[f1193_D])).
fof(f1193_D,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,'const.bool.IN_2'(s(sK1,X7),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f1187,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,'const.bool.IN_2'(s(sK1,X7),s(fun(sK1,bool),sK3))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl64 | $spl346)),
  inference(resolution,[],[f1182,f321])).
fof(f219,plain,(
  ( ! [X4,X2,X0,X8,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) ) | $spl40),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  ( ! [X4,X2,X0,X8,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f174,plain,(
  ( ! [X4,X2,X0,X8,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X8,X4,X3,X2,X1,X0)))))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,sK8(X4,X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7] : ! [X8,X9,X12] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X12)))) <=> (! [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,X13))),s(fun(X7,bool),X12))))) & ! [X13] : (p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X12)))) => ? [X14] : (p(s(bool,'const.bool.IN_2'(s(X6,X14),s(fun(X6,bool),X9)))) & s(X7,X13) = s(X7,happ(s(fun(X6,X7),X8),s(X6,X14)))))))),
  file('thm.pred_set.SURJ_EMPTY|split|0','thm.pred_set.SURJ_DEF')).
fof(f196,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f122,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f704,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),X48)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X49,X46),X50),s(fun(X49,bool),X51),s(fun(X46,bool),X48)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X49,sK6(X47,X48,X51,X50,X46,X49)),s(fun(X49,bool),X51)))) ) | $spl170),
  inference(cnf_transformation,[],[f704_D])).
fof(f704_D,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),X48)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X49,X46),X50),s(fun(X49,bool),X51),s(fun(X46,bool),X48)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X49,sK6(X47,X48,X51,X50,X46,X49)),s(fun(X49,bool),X51)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f689,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),X48)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X49,X46),X50),s(fun(X49,bool),X51),s(fun(X46,bool),X48)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X49,sK6(X47,X48,X51,X50,X46,X49)),s(fun(X49,bool),X51)))) ) | ($spl10 | $spl40)),
  inference(resolution,[],[f219,f196])).
fof(f217,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) ) | $spl36),
  inference(cnf_transformation,[],[f217_D])).
fof(f217_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f176,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f292,plain,(
  $false | ($spl7 | $spl36 | $spl58)),
  inference(subsumption_resolution,[],[f291,f193])).
fof(f291,plain,(
  p(s(bool,'F_0')) | ($spl7 | $spl36 | $spl58)),
  inference(forward_demodulation,[],[f290,f270])).
fof(f290,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) ) | ($spl7 | $spl36 | $spl58)),
  inference(subsumption_resolution,[],[f289,f193])).
fof(f289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) ) | ($spl7 | $spl36 | $spl58)),
  inference(forward_demodulation,[],[f288,f270])).
fof(f288,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) ) | ($spl7 | $spl36 | $spl58)),
  inference(subsumption_resolution,[],[f274,f193])).
fof(f274,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) ) | ($spl36 | $spl58)),
  inference(backward_demodulation,[],[f270,f217])).
fof(f194,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f119,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SURJ_EMPTY|split|0',aHL_BOOL_CASES)).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl58),
  inference(cnf_transformation,[],[f270_D])).
fof(f270_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f269,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl57),
  inference(cnf_transformation,[],[f269_D])).
fof(f269_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f260,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(equality_factoring,[],[f194])).
fof(f208,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl18),
  inference(cnf_transformation,[],[f208_D])).
fof(f208_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f165,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X13] : ~p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.SURJ_EMPTY|split|0','thm.pred_set.NOT_IN_EMPTY')).
fof(f243,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl48),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f240,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl14 | $spl18)),
  inference(resolution,[],[f208,f198])).
fof(f199,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f125,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.pred_set.SURJ_EMPTY|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f1182,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),sK3))))) ) | $spl346),
  inference(cnf_transformation,[],[f1182_D])).
fof(f1182_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),sK3))))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f7899,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),sK3))))) ) | ($spl16 | $spl48 | $spl57 | $spl170 | $spl342)),
  inference(subsumption_resolution,[],[f7898,f269])).
fof(f7898,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),sK3))))) ) | ($spl16 | $spl48 | $spl170 | $spl342)),
  inference(forward_demodulation,[],[f7897,f243])).
fof(f7897,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(X0,sK3,'const.pred_set.EMPTY_0',sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl16 | $spl170 | $spl342)),
  inference(condensation,[],[f7896])).
fof(f7896,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK1,X4),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(X4,sK3,'const.pred_set.EMPTY_0',sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl16 | $spl170 | $spl342)),
  inference(subsumption_resolution,[],[f7790,f199])).
fof(f7790,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK1,X4),s(fun(sK1,bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK6(X4,sK3,'const.pred_set.EMPTY_0',sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'F_0') = s(bool,X3)) ) | ($spl170 | $spl342)),
  inference(superposition,[],[f704,f1180])).
fof(f211,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | $spl24),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f171,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f72])).
fof(f312,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | $spl60),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f302,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | ($spl14 | $spl24)),
  inference(resolution,[],[f211,f198])).
fof(f1180,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,X4)) ) | $spl342),
  inference(cnf_transformation,[],[f1180_D])).
fof(f1180_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,X4)) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f1164,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK3))) | s(bool,'F_0') = s(bool,X4)) ) | ($spl0 | $spl60)),
  inference(resolution,[],[f186,f312])).
fof(f223,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | $spl45),
  inference(cnf_transformation,[],[f223_D])).
fof(f223_D,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f221,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f187,f188])).
fof(f198,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f126,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f79])).
fof(f230,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | $spl46),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f227,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK0,sK1),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | ($spl14 | $spl45)),
  inference(resolution,[],[f198,f223])).
fof(f193,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f118,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SURJ_EMPTY|split|0',aHL_FALSITY)).
fof(f1155,plain,(
  $false | ($spl7 | $spl46 | $spl322)),
  inference(subsumption_resolution,[],[f1143,f193])).
fof(f1143,plain,(
  p(s(bool,'F_0')) | ($spl46 | $spl322)),
  inference(superposition,[],[f1127,f230])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 146
% Passive clauses: 691
% Generated clauses: 8077
% Final active clauses: 79
% Final passive clauses: 295
% Input formulas: 28
% Initial clauses: 66
% 
% Duplicate literals: 235
% Trivial inequalities: 68
% Fw subsumption resolutions: 787
% Fw demodulations: 968
% Bw demodulations: 272
% Condensations: 355
% Global subsumptions: 9
% 
% Simple tautologies: 88
% Equational tautologies: 39
% Forward subsumptions: 3736
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 427
% Factoring: 1
% Forward superposition: 1733
% Backward superposition: 2004
% Self superposition: 79
% Equality factoring: 44
% Equality resolution: 13
% 
% Split clauses: 73
% Split components: 148
% SAT solver clauses: 1134
% SAT solver unit clauses: 76
% SAT solver binary clauses: 34
% SAT solver learnt clauses: 20
% SAT solver learnt literals: 1
% 
% Sat splits: 73
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 1096
% TWLsolver calls for satisfiability: 393
% 
% Memory used [KB]: 7036
% Time elapsed: 0.328 s
% ------------------------------
% ------------------------------
% Success in time 0.942 s
