Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rich_list
% (91365)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1449
% Passive clauses: 7176
% Generated clauses: 41158
% Final active clauses: 1416
% Final passive clauses: 5233
% Input formulas: 516
% Initial clauses: 945
% 
% Duplicate literals: 184
% Trivial inequalities: 287
% Fw subsumption resolutions: 1657
% Fw demodulations: 6469
% Bw demodulations: 23
% 
% Simple tautologies: 381
% Equational tautologies: 20
% Forward subsumptions: 24810
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 7374
% Factoring: 501
% Forward superposition: 13071
% Backward superposition: 10417
% Self superposition: 44
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 325
% Split components: 652
% SAT solver clauses: 502
% SAT solver unit clauses: 41
% SAT solver binary clauses: 102
% 
% Sat splits: 325
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 502
% TWLsolver calls for satisfiability: 411
% 
% Memory used [KB]: 20084
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% (91371)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 617
% Passive clauses: 3649
% Generated clauses: 23775
% Final active clauses: 603
% Final passive clauses: 2817
% Input formulas: 516
% Initial clauses: 945
% 
% Duplicate literals: 110
% Trivial inequalities: 76
% Fw subsumption resolutions: 872
% Fw demodulations: 3181
% Bw demodulations: 223
% Condensations: 64
% Global subsumptions: 26
% 
% Simple tautologies: 359
% Equational tautologies: 12
% Forward subsumptions: 11264
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 2838
% Factoring: 128
% Forward superposition: 7089
% Backward superposition: 4433
% Self superposition: 75
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 28
% Split components: 56
% SAT solver clauses: 5097
% SAT solver unit clauses: 1194
% SAT solver binary clauses: 281
% SAT solver learnt clauses: 52
% SAT solver learnt literals: 16
% 
% Sat splits: 28
% 
% TWLsolver clauses: 4993
% TWLsolver calls for satisfiability: 5289
% 
% Memory used [KB]: 22771
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rich_list
% (91380)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1970
% Passive clauses: 21785
% Generated clauses: 312872
% Final active clauses: 1614
% Final passive clauses: 4248
% Input formulas: 516
% Initial clauses: 945
% Discarded non-redundant clauses: 158242
% 
% Duplicate literals: 15
% Trivial inequalities: 316
% Fw subsumption resolutions: 4238
% Fw demodulations: 45944
% Bw demodulations: 6703
% Condensations: 68
% Global subsumptions: 8
% 
% Simple tautologies: 285
% Equational tautologies: 117
% Forward subsumptions: 104030
% Fw demodulations to eq. taut.: 2489
% Bw demodulations to eq. taut.: 74
% 
% Unit resulting resolution: 167348
% Factoring: 3
% Forward superposition: 49127
% Backward superposition: 33009
% Self superposition: 274
% Equality factoring: 2
% Equality resolution: 48
% 
% Split clauses: 68
% Split components: 137
% SAT solver clauses: 25463
% SAT solver unit clauses: 11331
% SAT solver binary clauses: 11443
% 
% Sat splits: 68
% Sat splitting refutations: 3532
% 
% Memory used [KB]: 196883
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f8831,plain,(
  $false),
  inference(subsumption_resolution,[],[f8830,f1404])).
fof(f1404,plain,(
  s('type.list.list'(sK14),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.list.list'(sK14),sK15))) != s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0')))),
  inference(forward_demodulation,[],[f992,f1063])).
fof(f1063,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f588])).
fof(f588,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f471])).
fof(f471,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))),
  file('thm.rich_list.LASTN_|49|','thm.arithmetic.NUMERAL_DEF')).
fof(f992,plain,(
  s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0'))) != s('type.list.list'(sK14),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'(sK14),sK15)))),
  inference(cnf_transformation,[],[f834])).
fof(f834,plain,(
  s('type.list.list'(sK14),'const.list.NIL_0') != s('type.list.list'(sK14),sK15) & s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0'))) != s('type.list.list'(sK14),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'(sK14),sK15)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f753])).
fof(f753,plain,(
  ? [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1) & s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) != s('type.list.list'(X0),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'(X0),X1))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1) => s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'(X0),X1))))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ~! [X0] : ! [X1] : (~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) => s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'(X0),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X7) => s('type.list.list'(X6),'const.list.CONS_2'(s(X6,'const.list.LAST_1'(s('type.list.list'(X6),X7))),s('type.list.list'(X6),'const.list.NIL_0'))) = s('type.list.list'(X6),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'(X6),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X7) => s('type.list.list'(X6),'const.list.CONS_2'(s(X6,'const.list.LAST_1'(s('type.list.list'(X6),X7))),s('type.list.list'(X6),'const.list.NIL_0'))) = s('type.list.list'(X6),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'(X6),X7))))),
  file('thm.rich_list.LASTN_|49|','thm.rich_list.LASTN_|49|')).
fof(f8830,plain,(
  s('type.list.list'(sK14),'const.rich_list.LASTN_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0')))),
  inference(forward_demodulation,[],[f8829,f1082])).
fof(f1082,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(cnf_transformation,[],[f357])).
fof(f357,axiom,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.rich_list.LASTN_|49|','thm.numeral.numeral_suc|split|0')).
fof(f8829,plain,(
  s('type.list.list'(sK14),'const.rich_list.LASTN_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0')))),
  inference(forward_demodulation,[],[f8828,f1087])).
fof(f1087,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  inference(cnf_transformation,[],[f441])).
fof(f441,axiom,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  file('thm.rich_list.LASTN_|49|','thm.arithmetic.ALT_ZERO')).
fof(f8828,plain,(
  s('type.list.list'(sK14),'const.rich_list.LASTN_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0')))),
  inference(forward_demodulation,[],[f8827,f1394])).
fof(f1394,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f752])).
fof(f752,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.rich_list.LASTN_|49|','thm.list.LENGTH|split|0')).
fof(f8827,plain,(
  s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0'))) = s('type.list.list'(sK14),'const.rich_list.LASTN_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK14),'const.list.NIL_0'))))),s('type.list.list'(sK14),sK15)))),
  inference(forward_demodulation,[],[f8457,f1138])).
fof(f1138,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f647])).
fof(f647,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f646])).
fof(f646,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X23,X18] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X18))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X23),s('type.list.list'(X6),X18)))))),
  file('thm.rich_list.LASTN_|49|','thm.list.LENGTH|split|1')).
fof(f8457,plain,(
  s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0'))) = s('type.list.list'(sK14),'const.rich_list.LASTN_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0'))))),s('type.list.list'(sK14),sK15)))),
  inference(superposition,[],[f1004,f1432])).
fof(f1432,plain,(
  s('type.list.list'(sK14),sK15) = s('type.list.list'(sK14),'const.list.APPEND_2'(s('type.list.list'(sK14),'const.list.FRONT_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0')))))),
  inference(trivial_inequality_removal,[],[f1415])).
fof(f1415,plain,(
  s('type.list.list'(sK14),'const.list.NIL_0') != s('type.list.list'(sK14),'const.list.NIL_0') | s('type.list.list'(sK14),sK15) = s('type.list.list'(sK14),'const.list.APPEND_2'(s('type.list.list'(sK14),'const.list.FRONT_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,'const.list.LAST_1'(s('type.list.list'(sK14),sK15))),s('type.list.list'(sK14),'const.list.NIL_0')))))),
  inference(superposition,[],[f991,f1157])).
fof(f1157,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f783])).
fof(f783,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))))),
  inference(ennf_transformation,[],[f671])).
fof(f671,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1) => s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))))),
  inference(flattening,[],[f670])).
fof(f670,plain,(
  ! [X0] : ! [X1] : (~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) => s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.LAST_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7] : (~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X7) => s('type.list.list'(X6),X7) = s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),X7))),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,'const.list.LAST_1'(s('type.list.list'(X6),X7))),s('type.list.list'(X6),'const.list.NIL_0'))))))),
  file('thm.rich_list.LASTN_|49|','thm.list.APPEND_FRONT_LAST')).
fof(f991,plain,(
  s('type.list.list'(sK14),'const.list.NIL_0') != s('type.list.list'(sK14),sK15)),
  inference(cnf_transformation,[],[f834])).
fof(f1004,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.rich_list.LASTN_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.rich_list.LASTN_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f542])).
fof(f542,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.rich_list.LASTN_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : ! [X28,X27] : s('type.list.list'(X6),X28) = s('type.list.list'(X6),'const.rich_list.LASTN_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X28))),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X27),s('type.list.list'(X6),X28)))))),
  file('thm.rich_list.LASTN_|49|','thm.rich_list.LASTN_LENGTH_APPEND')).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 415
% Passive clauses: 2649
% Generated clauses: 7389
% Final active clauses: 401
% Final passive clauses: 1803
% Input formulas: 516
% Initial clauses: 404
% 
% Selected by SInE selection: 165
% SInE iterations: 1
% 
% Duplicate literals: 109
% Trivial inequalities: 162
% Fw subsumption resolutions: 357
% Fw demodulations: 993
% 
% Simple tautologies: 226
% Equational tautologies: 3
% Forward subsumptions: 2067
% 
% Binary resolution: 1722
% Forward superposition: 3215
% Backward superposition: 660
% Equality resolution: 1
% 
% Split clauses: 420
% Split components: 848
% SAT solver clauses: 429
% SAT solver unit clauses: 7
% SAT solver binary clauses: 277
% 
% Sat splits: 420
% Sat splitting refutations: 8
% 
% Memory used [KB]: 6268
% Time elapsed: 0.180 s
% ------------------------------
% ------------------------------
% Success in time 11.951 s
