Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (22900)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 103
% Passive clauses: 3916
% Generated clauses: 12806
% Final active clauses: 93
% Final passive clauses: 3410
% Input formulas: 37
% Initial clauses: 76
% 
% Duplicate literals: 546
% Trivial inequalities: 214
% Fw subsumption resolutions: 1628
% Fw demodulations: 394
% Bw demodulations: 5
% 
% Simple tautologies: 289
% Equational tautologies: 89
% Forward subsumptions: 4828
% 
% Binary resolution: 294
% Factoring: 3
% Forward superposition: 3839
% Backward superposition: 5185
% Self superposition: 295
% Equality factoring: 214
% Equality resolution: 23
% 
% Split clauses: 447
% Split components: 960
% SAT solver clauses: 1053
% SAT solver unit clauses: 36
% SAT solver binary clauses: 362
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 447
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1007
% TWLsolver calls for satisfiability: 126
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f7629,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7628,f6195,f6196,f195,f255,f138,f221,f219,f6417,f7577,f6409,f6410,f184,f248,f173,f237,f172,f236,f137,f220,f218])).
fof(f218,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2))) | $spl0),
  inference(cnf_transformation,[],[f218_D])).
fof(f218_D,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f220,plain,(
  s('type.ringNorm.polynom'(sK0),sK1) = s('type.ringNorm.polynom'(sK0),sK2) | $spl2),
  inference(cnf_transformation,[],[f220_D])).
fof(f220_D,plain,(
  s('type.ringNorm.polynom'(sK0),sK1) = s('type.ringNorm.polynom'(sK0),sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f137,plain,(
  s('type.ringNorm.polynom'(sK0),sK1) = s('type.ringNorm.polynom'(sK0),sK2) | s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2)))),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2))) | s('type.ringNorm.polynom'(sK0),sK1) = s('type.ringNorm.polynom'(sK0),sK2)) & (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2))) | s('type.ringNorm.polynom'(sK0),sK1) != s('type.ringNorm.polynom'(sK0),sK2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f111])).
fof(f111,plain,(
  ? [X0,X1,X2] : ((s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X2))) | s('type.ringNorm.polynom'(X0),X1) = s('type.ringNorm.polynom'(X0),X2)) & (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X2))) | s('type.ringNorm.polynom'(X0),X1) != s('type.ringNorm.polynom'(X0),X2)))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ? [X0,X1,X2] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X2))) <~> s('type.ringNorm.polynom'(X0),X1) = s('type.ringNorm.polynom'(X0),X2))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X2))) <=> s('type.ringNorm.polynom'(X0),X1) = s('type.ringNorm.polynom'(X0),X2))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X2))) <=> s('type.ringNorm.polynom'(X0),X1) = s('type.ringNorm.polynom'(X0),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X7))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X8))) <=> s('type.ringNorm.polynom'(X6),X7) = s('type.ringNorm.polynom'(X6),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X7))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X8))) <=> s('type.ringNorm.polynom'(X6),X7) = s('type.ringNorm.polynom'(X6),X8))),
  file('thm.ringNorm.polynom_|49||49||split|4','thm.ringNorm.polynom_|49||49||split|4')).
fof(f236,plain,(
  ( ! [X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))) ) | $spl18),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  ( ! [X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f172,plain,(
  ( ! [X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X1))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X9] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(X6),X9)))),
  file('thm.ringNorm.polynom_|49||49||split|4','ARITY_0_1_const.ringNorm.Popp')).
fof(f237,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(cnf_transformation,[],[f237_D])).
fof(f237_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f173,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X7,X10] : s(X6,X7) = s(X6,'const.ind_type.FCONS_3'(s(X6,X7),s(fun('type.num.num',X6),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.ringNorm.polynom_|49||49||split|4','thm.ind_type.FCONS|split|0')).
fof(f248,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) | $spl42),
  inference(cnf_transformation,[],[f248_D])).
fof(f248_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f184,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X9,X23,X24] : s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X23),s('type.num.num',X24))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X9))),s(fun('type.num.num',X6),X23))),s('type.num.num',X24)))),
  file('thm.ringNorm.polynom_|49||49||split|4','ARITY_0_3_const.ind_type.FCONS')).
fof(f6410,plain,(
  ( ! [X144,X145] : (s('type.ringNorm.polynom'(sK0),'const.ind_type.FCONS_3'(s('type.ringNorm.polynom'(sK0),sK1),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X144),s('type.num.num',X145))) = s('type.ringNorm.polynom'(sK0),'const.ind_type.FCONS_3'(s('type.ringNorm.polynom'(sK0),sK2),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X144),s('type.num.num',X145)))) ) | $spl1012),
  inference(cnf_transformation,[],[f6410_D])).
fof(f6410_D,plain,(
  ( ! [X144,X145] : (s('type.ringNorm.polynom'(sK0),'const.ind_type.FCONS_3'(s('type.ringNorm.polynom'(sK0),sK1),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X144),s('type.num.num',X145))) = s('type.ringNorm.polynom'(sK0),'const.ind_type.FCONS_3'(s('type.ringNorm.polynom'(sK0),sK2),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X144),s('type.num.num',X145)))) ) <=> ~$spl1012),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1012])])).
fof(f6409,plain,(
  ( ! [X144,X145] : (s('type.ringNorm.polynom'(sK0),'const.ind_type.FCONS_3'(s('type.ringNorm.polynom'(sK0),sK1),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X144),s('type.num.num',X145))) = s('type.ringNorm.polynom'(sK0),'const.ind_type.FCONS_3'(s('type.ringNorm.polynom'(sK0),sK2),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X144),s('type.num.num',X145)))) ) | ($spl42 | $spl918)),
  inference(forward_demodulation,[],[f6329,f248])).
fof(f6329,plain,(
  ( ! [X144,X145] : (s('type.ringNorm.polynom'(sK0),'const.ind_type.FCONS_3'(s('type.ringNorm.polynom'(sK0),sK2),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X144),s('type.num.num',X145))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.num.num','type.ringNorm.polynom'(sK0)),happ(s(fun(fun('type.num.num','type.ringNorm.polynom'(sK0)),fun('type.num.num','type.ringNorm.polynom'(sK0))),happ(s(fun('type.ringNorm.polynom'(sK0),fun(fun('type.num.num','type.ringNorm.polynom'(sK0)),fun('type.num.num','type.ringNorm.polynom'(sK0)))),'const.ind_type.FCONS_0'),s('type.ringNorm.polynom'(sK0),sK1))),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X144))),s('type.num.num',X145)))) ) | ($spl42 | $spl918)),
  inference(superposition,[],[f248,f6196])).
fof(f7577,plain,(
  $false | ($spl3 | $spl20 | $spl1012)),
  inference(subsumption_resolution,[],[f7576,f221])).
fof(f7576,plain,(
  s('type.ringNorm.polynom'(sK0),sK1) = s('type.ringNorm.polynom'(sK0),sK2) | ($spl20 | $spl1012)),
  inference(forward_demodulation,[],[f7549,f237])).
fof(f7549,plain,(
  ( ! [X2] : (s('type.ringNorm.polynom'(sK0),sK2) = s('type.ringNorm.polynom'(sK0),'const.ind_type.FCONS_3'(s('type.ringNorm.polynom'(sK0),sK1),s(fun('type.num.num','type.ringNorm.polynom'(sK0)),X2),s('type.num.num','const.num.|48|_0')))) ) | ($spl20 | $spl1012)),
  inference(superposition,[],[f6410,f237])).
fof(f6417,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2))) | ($spl18 | $spl918)),
  inference(forward_demodulation,[],[f6346,f236])).
fof(f6346,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.ringNorm.polynom'(sK0),'type.ringNorm.polynom'(sK0)),'const.ringNorm.Popp_0'),s('type.ringNorm.polynom'(sK0),sK1))) | ($spl18 | $spl918)),
  inference(superposition,[],[f236,f6196])).
fof(f219,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2))) | $spl1),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f221,plain,(
  s('type.ringNorm.polynom'(sK0),sK1) != s('type.ringNorm.polynom'(sK0),sK2) | $spl3),
  inference(cnf_transformation,[],[f221_D])).
fof(f221_D,plain,(
  s('type.ringNorm.polynom'(sK0),sK1) != s('type.ringNorm.polynom'(sK0),sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f138,plain,(
  s('type.ringNorm.polynom'(sK0),sK1) != s('type.ringNorm.polynom'(sK0),sK2) | s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK2)))),
  inference(cnf_transformation,[],[f112])).
fof(f255,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) ) | $spl56),
  inference(cnf_transformation,[],[f255_D])).
fof(f255_D,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f195,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X6] : ! [X7,X10,X12,X13,X14,X15] : s(X11,happ(s(fun('type.ringNorm.polynom'(X6),X11),X15),s('type.ringNorm.polynom'(X6),X7))) = s(X11,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X7))),s(fun('type.quote.index',X11),X10),s(fun(X6,X11),X12),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X11)),X13),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X11)),X14),s(fun('type.ringNorm.polynom'(X6),X11),X15)))),
  file('thm.ringNorm.polynom_|49||49||split|4','thm.ringNorm.polynom_case_def|split|4')).
fof(f6196,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK2)))) ) | $spl918),
  inference(cnf_transformation,[],[f6196_D])).
fof(f6196_D,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK2)))) ) <=> ~$spl918),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl918])])).
fof(f6195,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK2)))) ) | ($spl0 | $spl56)),
  inference(forward_demodulation,[],[f6064,f255])).
fof(f6064,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),X1),s('type.ringNorm.polynom'(sK0),sK2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X5),s(fun('type.ringNorm.polynom'(sK0),X0),X1)))) ) | ($spl0 | $spl56)),
  inference(superposition,[],[f255,f218])).
fof(f7628,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f7627])).
fof(f7627,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK1))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f219,f220])).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 75
% Passive clauses: 602
% Generated clauses: 7425
% Final active clauses: 60
% Final passive clauses: 481
% Input formulas: 37
% Initial clauses: 76
% 
% Duplicate literals: 349
% Trivial inequalities: 104
% Fw subsumption resolutions: 445
% Fw demodulations: 586
% Bw demodulations: 48
% Condensations: 699
% Global subsumptions: 5
% 
% Simple tautologies: 147
% Equational tautologies: 37
% Forward subsumptions: 3264
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 54
% Factoring: 1
% Forward superposition: 1971
% Backward superposition: 2257
% Self superposition: 131
% Equality factoring: 72
% Equality resolution: 14
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 1256
% SAT solver unit clauses: 77
% SAT solver binary clauses: 62
% SAT solver learnt clauses: 12
% SAT solver learnt literals: 1
% 
% Sat splits: 15
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1234
% TWLsolver calls for satisfiability: 248
% 
% Memory used [KB]: 10362
% Time elapsed: 0.374 s
% ------------------------------
% ------------------------------
% Success in time 1.01 s
