Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (65204)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 2716
% Generated clauses: 20259
% Final active clauses: 137
% Final passive clauses: 2257
% Input formulas: 25
% Initial clauses: 66
% 
% Duplicate literals: 814
% Trivial inequalities: 364
% Fw subsumption resolutions: 3261
% Fw demodulations: 3158
% Bw demodulations: 15
% 
% Simple tautologies: 439
% Equational tautologies: 45
% Forward subsumptions: 8499
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 629
% Factoring: 4
% Forward superposition: 6530
% Backward superposition: 5922
% Self superposition: 190
% Equality factoring: 101
% Equality resolution: 31
% 
% Split clauses: 541
% Split components: 1103
% SAT solver clauses: 1235
% SAT solver unit clauses: 23
% SAT solver binary clauses: 240
% 
% Sat splits: 541
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1230
% TWLsolver calls for satisfiability: 198
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (65217)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 98
% Passive clauses: 532
% Generated clauses: 9133
% Final active clauses: 76
% Final passive clauses: 378
% Input formulas: 25
% Initial clauses: 66
% 
% Duplicate literals: 781
% Trivial inequalities: 155
% Fw subsumption resolutions: 922
% Fw demodulations: 361
% Bw demodulations: 57
% Condensations: 1334
% Global subsumptions: 9
% 
% Simple tautologies: 220
% Equational tautologies: 57
% Forward subsumptions: 3424
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 182
% Factoring: 1
% Forward superposition: 1710
% Backward superposition: 3052
% Self superposition: 172
% Equality factoring: 139
% Equality resolution: 23
% 
% Split clauses: 157
% Split components: 314
% SAT solver clauses: 1059
% SAT solver unit clauses: 77
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 21
% SAT solver learnt literals: 5
% 
% Sat splits: 157
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1020
% TWLsolver calls for satisfiability: 402
% 
% Memory used [KB]: 7036
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (65246)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 622
% Passive clauses: 6264
% Generated clauses: 165394
% Final active clauses: 164
% Final passive clauses: 848
% Input formulas: 25
% Initial clauses: 66
% 
% Duplicate literals: 7745
% Trivial inequalities: 1346
% Fw subsumption resolutions: 15934
% Fw demodulations: 21512
% Bw demodulations: 3276
% Condensations: 17930
% Global subsumptions: 821
% 
% Simple tautologies: 1417
% Equational tautologies: 1866
% Forward subsumptions: 88926
% Fw demodulations to eq. taut.: 136
% Bw demodulations to eq. taut.: 238
% 
% Unit resulting resolution: 8174
% Factoring: 18
% Forward superposition: 45100
% Backward superposition: 33892
% Self superposition: 858
% Equality factoring: 437
% Equality resolution: 96
% 
% Split clauses: 289
% Split components: 589
% SAT solver clauses: 8631
% SAT solver unit clauses: 52
% SAT solver binary clauses: 558
% 
% Sat splits: 289
% Sat splitting refutations: 966
% 
% Memory used [KB]: 75094
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f1640,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f110,f179,f177,f109,f178,f176,f363,f365,f1139,f1140,f1153,f1141,f1142,f1154,f1143,f1144,f1152,f1145,f1146,f1151,f1147,f1158,f1160,f1148,f1156,f1162,f1212,f1216,f1215,f1214,f1285,f1286,f1287,f1283,f1281,f1288,f1213,f1638,f1639])).
fof(f1639,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1632,f173])).
fof(f173,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7] : ! [X13] : ~p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.pred_set.NOT_IN_EMPTY')).
fof(f1632,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK4,'const.pred_set.EMPTY_0',sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f1624])).
fof(f1624,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK4,'const.pred_set.EMPTY_0',sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(superposition,[],[f1605,f120])).
fof(f120,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f1605,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK4,'const.pred_set.EMPTY_0',sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl3)),
  inference(equality_resolution,[],[f1353])).
fof(f1353,plain,(
  ( ! [X0] : (s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),X0) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK4,X0,sK2)),s(fun(sK2,bool),X0)))) ) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f368,f436])).
fof(f436,plain,(
  ( ! [X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X32),s(fun(sK2,bool),sK4)))) ) | $spl0),
  inference(resolution,[],[f424,f118])).
fof(f118,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f424,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X0),s(fun(sK2,bool),sK4))))) ) | $spl0),
  inference(subsumption_resolution,[],[f422,f173])).
fof(f422,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK2,sK1),sK3),s(sK2,X0))),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK2,X0),s(fun(sK2,bool),sK4))))) ) | $spl0),
  inference(resolution,[],[f374,f161])).
fof(f161,plain,(
  ( ! [X4,X2,X0,X10,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | ~sP0(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,sK8(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK8(X4,X3,X2,X1,X0)))) & s(X0,sK7(X4,X3,X2,X1,X0)) != s(X0,sK8(X4,X3,X2,X1,X0))) | sP0(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))) & s(X0,X5) != s(X0,X6)) | sP0(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f104])).
fof(f104,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP0(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6)) | sP0(X1,X0,X2,X3,X4)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP0(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6))) | sP0(X1,X0,X2,X3,X4)))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X1,X0,X2,X3,X4] : (sP0(X1,X0,X2,X3,X4) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f374,plain,(
  sP0(sK2,sK1,sK3,sK4,'const.pred_set.EMPTY_0') | $spl0),
  inference(resolution,[],[f176,f171])).
fof(f171,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X1,X0,X2,X3,X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | sP0(X1,X0,X2,X3,X4)) & (~sP0(X1,X0,X2,X3,X4) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> sP0(X1,X0,X2,X3,X4))),
  inference(definition_folding,[],[f68,f70])).
fof(f68,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) | (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X7] : ! [X8,X9,X12] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),X12)))) <=> (! [X13] : (p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X7,X6),X8),s(X7,X13))),s(fun(X6,bool),X12))))) & ! [X13,X14] : ((p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X9)))) & p(s(bool,'const.bool.IN_2'(s(X7,X14),s(fun(X7,bool),X9))))) => (s(X6,happ(s(fun(X7,X6),X8),s(X7,X13))) = s(X6,happ(s(fun(X7,X6),X8),s(X7,X14))) => s(X7,X13) = s(X7,X14)))))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.pred_set.INJ_DEF')).
fof(f368,plain,(
  ( ! [X0] : (s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),X0) | s(bool,'const.bool.IN_2'(s(sK2,sK5(sK4,X0,sK2)),s(fun(sK2,bool),sK4))) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK4,X0,sK2)),s(fun(sK2,bool),X0)))) ) | $spl3),
  inference(superposition,[],[f179,f116])).
fof(f116,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7] : ! [X9,X12] : (s(fun(X7,bool),X9) = s(fun(X7,bool),X12) <=> ! [X13] : s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X9))) = s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X12))))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.pred_set.EXTENSION')).
fof(f1638,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1633,f173])).
fof(f1633,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK4,'const.pred_set.EMPTY_0',sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f1623])).
fof(f1623,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK4,'const.pred_set.EMPTY_0',sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(superposition,[],[f1605,f118])).
fof(f1213,plain,(
  ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl19),
  inference(cnf_transformation,[],[f1213_D])).
fof(f1213_D,plain,(
  ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1288,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) | $spl22),
  inference(cnf_transformation,[],[f1288_D])).
fof(f1288_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1281,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))))) ) | $spl13),
  inference(subsumption_resolution,[],[f1268,f121])).
fof(f121,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f1268,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | ~p(s(bool,X2))) ) | $spl13),
  inference(superposition,[],[f1158,f114])).
fof(f114,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f1283,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))))) ) | $spl13),
  inference(subsumption_resolution,[],[f1269,f121])).
fof(f1269,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))))) ) | $spl13),
  inference(superposition,[],[f1158,f114])).
fof(f1287,plain,(
  $false | ($spl11 | $spl13)),
  inference(subsumption_resolution,[],[f1278,f1156])).
fof(f1278,plain,(
  s(bool,'F_0') = s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl13),
  inference(trivial_inequality_removal,[],[f1274])).
fof(f1274,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl13),
  inference(superposition,[],[f1158,f125])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INJ_EMPTY|split|1',aHL_BOOL_CASES)).
fof(f1286,plain,(
  $false | ($spl11 | $spl13)),
  inference(subsumption_resolution,[],[f1279,f1210])).
fof(f1210,plain,(
  p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl11),
  inference(trivial_inequality_removal,[],[f1202])).
fof(f1202,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl11),
  inference(superposition,[],[f1156,f120])).
fof(f1279,plain,(
  ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl13),
  inference(trivial_inequality_removal,[],[f1273])).
fof(f1273,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl13),
  inference(superposition,[],[f1158,f124])).
fof(f124,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f1285,plain,(
  $false | ($spl11 | $spl13)),
  inference(subsumption_resolution,[],[f1280,f1210])).
fof(f1280,plain,(
  ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl13),
  inference(trivial_inequality_removal,[],[f1272])).
fof(f1272,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl13),
  inference(superposition,[],[f1158,f122])).
fof(f122,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f79])).
fof(f1214,plain,(
  p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl18),
  inference(cnf_transformation,[],[f1214_D])).
fof(f1214_D,plain,(
  p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1215,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f1215_D])).
fof(f1215_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1216,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))))) ) | $spl11),
  inference(subsumption_resolution,[],[f1198,f117])).
fof(f117,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f1198,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | p(s(bool,X1))) ) | $spl11),
  inference(superposition,[],[f1156,f111])).
fof(f111,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f60])).
fof(f1212,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))))) ) | $spl11),
  inference(subsumption_resolution,[],[f1197,f117])).
fof(f1197,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))))) ) | $spl11),
  inference(superposition,[],[f1156,f111])).
fof(f1162,plain,(
  s(bool,'T_0') = s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl16),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1156,plain,(
  s(bool,'F_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl11),
  inference(cnf_transformation,[],[f1156_D])).
fof(f1156_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1148,plain,(
  s(bool,'F_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | s(bool,'T_0') = s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl3),
  inference(superposition,[],[f1136,f125])).
fof(f1136,plain,(
  s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) != s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl3),
  inference(equality_resolution,[],[f370])).
fof(f370,plain,(
  ( ! [X2] : (s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),X2) | s(bool,f(s(sK2,sK9(sK4,X2,bool,sK2)))) != s(bool,g(s(sK2,sK9(sK4,X2,bool,sK2))))) ) | $spl3),
  inference(superposition,[],[f179,f174])).
fof(f174,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.INJ_EMPTY|split|1',aHL_EXT)).
fof(f1160,plain,(
  s(bool,'F_0') = s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl14),
  inference(cnf_transformation,[],[f1160_D])).
fof(f1160_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1158,plain,(
  s(bool,'T_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl13),
  inference(cnf_transformation,[],[f1158_D])).
fof(f1158_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1147,plain,(
  s(bool,'T_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | s(bool,'F_0') = s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | $spl3),
  inference(superposition,[],[f1136,f125])).
fof(f1151,plain,(
  ~p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl5),
  inference(cnf_transformation,[],[f1151_D])).
fof(f1151_D,plain,(
  ~p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f1146,plain,(
  s(bool,'T_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | ~p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl3),
  inference(superposition,[],[f1136,f124])).
fof(f1145,plain,(
  s(bool,'T_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | ~p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl3),
  inference(superposition,[],[f1136,f122])).
fof(f1152,plain,(
  p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl4),
  inference(cnf_transformation,[],[f1152_D])).
fof(f1152_D,plain,(
  p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1144,plain,(
  s(bool,'F_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl3),
  inference(superposition,[],[f1136,f120])).
fof(f1143,plain,(
  s(bool,'F_0') != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | $spl3),
  inference(superposition,[],[f1136,f118])).
fof(f1154,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | ~p(s(bool,X2))) ) | $spl8),
  inference(cnf_transformation,[],[f1154_D])).
fof(f1154_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | ~p(s(bool,X2))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1142,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | ~p(s(bool,X3)) | ~p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))))) ) | $spl3),
  inference(superposition,[],[f1136,f114])).
fof(f1141,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | ~p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | ~p(s(bool,X2))) ) | $spl3),
  inference(superposition,[],[f1136,f114])).
fof(f1153,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | p(s(bool,X0))) ) | $spl6),
  inference(cnf_transformation,[],[f1153_D])).
fof(f1153_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | p(s(bool,X0))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1140,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2))))) | p(s(bool,X1))) ) | $spl3),
  inference(superposition,[],[f1136,f111])).
fof(f1139,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))) | p(s(bool,X0)) | p(s(bool,g(s(sK2,sK9(sK4,'const.pred_set.EMPTY_0',bool,sK2)))))) ) | $spl3),
  inference(superposition,[],[f1136,f111])).
fof(f365,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f364,f173])).
fof(f364,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK8('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK3,sK1,sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f355,f173])).
fof(f355,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK8('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK3,sK1,sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK6('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK3,sK1,sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(resolution,[],[f306,f165])).
fof(f165,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK8(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f306,plain,(
  ~sP0(sK2,sK1,sK3,'const.pred_set.EMPTY_0','const.pred_set.EMPTY_0') | ($spl1 | $spl2)),
  inference(resolution,[],[f283,f172])).
fof(f172,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | ~sP0(X1,X0,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f107])).
fof(f283,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f177,f178])).
fof(f363,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f362,f173])).
fof(f362,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK3,sK1,sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f353,f173])).
fof(f353,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK3,sK1,sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK6('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK3,sK1,sK2)),s(fun(sK2,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(resolution,[],[f306,f163])).
fof(f163,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f176,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),sK4),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | $spl0),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),sK4),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f178,plain,(
  s(fun(sK2,bool),'const.pred_set.EMPTY_0') = s(fun(sK2,bool),sK4) | $spl2),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  s(fun(sK2,bool),'const.pred_set.EMPTY_0') = s(fun(sK2,bool),sK4) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f109,plain,(
  s(fun(sK2,bool),'const.pred_set.EMPTY_0') = s(fun(sK2,bool),sK4) | p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),sK4),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),sK4),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK2,bool),'const.pred_set.EMPTY_0') = s(fun(sK2,bool),sK4)) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),sK4),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),sK4))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <~> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  file('thm.pred_set.INJ_EMPTY|split|1','thm.pred_set.INJ_EMPTY|split|1')).
fof(f177,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),sK4),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | $spl1),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),sK4),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f179,plain,(
  s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),sK4) | $spl3),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),sK4) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f110,plain,(
  s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),sK4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,sK1),sK3),s(fun(sK2,bool),sK4),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))),
  inference(cnf_transformation,[],[f73])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 76
% Passive clauses: 385
% Generated clauses: 1469
% Final active clauses: 59
% Final passive clauses: 219
% Input formulas: 25
% Initial clauses: 66
% 
% Selected by SInE selection: 25
% SInE iterations: 1
% 
% Duplicate literals: 28
% Trivial inequalities: 33
% Fw subsumption resolutions: 169
% Fw demodulations: 236
% Bw demodulations: 9
% 
% Simple tautologies: 63
% Forward subsumptions: 593
% 
% Binary resolution: 146
% Forward superposition: 518
% Backward superposition: 301
% Equality resolution: 4
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 24
% SAT solver binary clauses: 9
% 
% Sat splits: 16
% Sat splitting refutations: 7
% 
% Memory used [KB]: 1663
% Time elapsed: 0.034 s
% ------------------------------
% ------------------------------
% Success in time 11.74 s
