Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.probability
% (57843)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 112
% Passive clauses: 2615
% Generated clauses: 10931
% Final active clauses: 102
% Final passive clauses: 2395
% Input formulas: 19
% Initial clauses: 38
% 
% Duplicate literals: 538
% Trivial inequalities: 142
% Fw subsumption resolutions: 1304
% Fw demodulations: 916
% Bw demodulations: 6
% 
% Simple tautologies: 238
% Equational tautologies: 36
% Forward subsumptions: 3932
% Fw demodulations to eq. taut.: 9
% 
% Binary resolution: 223
% Forward superposition: 3255
% Backward superposition: 4099
% Self superposition: 219
% Equality factoring: 108
% Equality resolution: 22
% 
% Split clauses: 621
% Split components: 1299
% SAT solver clauses: 917
% SAT solver unit clauses: 36
% SAT solver binary clauses: 205
% 
% Sat splits: 621
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 916
% TWLsolver calls for satisfiability: 133
% 
% Memory used [KB]: 14711
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.probability
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.probability
% SZS output start Proof for thm.probability
fof(f6411,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f6400,f73,f127,f205,f209,f70,f119,f80,f131,f6005,f6006,f71,f125,f142,f144,f68,f123,f76,f129,f510,f519,f84,f134,f726,f747,f69,f121,f85,f135,f1685,f1687,f86,f136,f96,f138])).
fof(f138,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,X0),X4),s(fun(X2,X1),X5),s(fun('type.pair.prod'(X0,X1),bool),X6))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,bool),'const.pred_set.INTER_2'(s(fun(X2,bool),'const.util_prob.PREIMAGE_2'(s(fun(X2,'type.pair.prod'(X0,X1)),'_dst_x0x1x1_3'(s(fun(X0,fun(X1,'type.pair.prod'(X0,X1))),'const.pair.,_0'),s(fun(X2,X0),X4),s(fun(X2,X1),X5))),s(fun('type.pair.prod'(X0,X1),bool),X6))),s(fun(X2,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3)))))))) ) | $spl30),
  inference(cnf_transformation,[],[f138_D])).
fof(f138_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,X0),X4),s(fun(X2,X1),X5),s(fun('type.pair.prod'(X0,X1),bool),X6))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,bool),'const.pred_set.INTER_2'(s(fun(X2,bool),'const.util_prob.PREIMAGE_2'(s(fun(X2,'type.pair.prod'(X0,X1)),'_dst_x0x1x1_3'(s(fun(X0,fun(X1,'type.pair.prod'(X0,X1))),'const.pair.,_0'),s(fun(X2,X0),X4),s(fun(X2,X1),X5))),s(fun('type.pair.prod'(X0,X1),bool),X6))),s(fun(X2,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3)))))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f96,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,X0),X4),s(fun(X2,X1),X5),s(fun('type.pair.prod'(X0,X1),bool),X6))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,bool),'const.pred_set.INTER_2'(s(fun(X2,bool),'const.util_prob.PREIMAGE_2'(s(fun(X2,'type.pair.prod'(X0,X1)),'_dst_x0x1x1_3'(s(fun(X0,fun(X1,'type.pair.prod'(X0,X1))),'const.pair.,_0'),s(fun(X2,X0),X4),s(fun(X2,X1),X5))),s(fun('type.pair.prod'(X0,X1),bool),X6))),s(fun(X2,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3)))))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,X0),X4),s(fun(X2,X1),X5),s(fun('type.pair.prod'(X0,X1),bool),X6))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,bool),'const.pred_set.INTER_2'(s(fun(X2,bool),'const.util_prob.PREIMAGE_2'(s(fun(X2,'type.pair.prod'(X0,X1)),'_dst_x0x1x1_3'(s(fun(X0,fun(X1,'type.pair.prod'(X0,X1))),'const.pair.,_0'),s(fun(X2,X0),X4),s(fun(X2,X1),X5))),s(fun('type.pair.prod'(X0,X1),bool),X6))),s(fun(X2,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3)))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6] : s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,X0),X4),s(fun(X2,X1),X5),s(fun('type.pair.prod'(X0,X1),bool),X6))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3),s(fun(X2,bool),'const.pred_set.INTER_2'(s(fun(X2,bool),'const.util_prob.PREIMAGE_2'(s(fun(X2,'type.pair.prod'(X0,X1)),'_dst_x0x1x1_3'(s(fun(X0,fun(X1,'type.pair.prod'(X0,X1))),'const.pair.,_0'),s(fun(X2,X0),X4),s(fun(X2,X1),X5))),s(fun('type.pair.prod'(X0,X1),bool),X6))),s(fun(X2,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X2,bool),'type.pair.prod'(fun(fun(X2,bool),bool),fun(fun(X2,bool),'type.realax.real'))),X3)))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X5] : s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X8,bool),'type.pair.prod'(fun(fun(X8,bool),bool),fun(fun(X8,bool),'type.realax.real'))),X9),s(fun(X8,X6),X10),s(fun(X8,X7),X11),s(fun('type.pair.prod'(X6,X7),bool),X5))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X8,bool),'type.pair.prod'(fun(fun(X8,bool),bool),fun(fun(X8,bool),'type.realax.real'))),X9),s(fun(X8,bool),'const.pred_set.INTER_2'(s(fun(X8,bool),'const.util_prob.PREIMAGE_2'(s(fun(X8,'type.pair.prod'(X6,X7)),'_dst_x0x1x1_3'(s(fun(X6,fun(X7,'type.pair.prod'(X6,X7))),'const.pair.,_0'),s(fun(X8,X6),X10),s(fun(X8,X7),X11))),s(fun('type.pair.prod'(X6,X7),bool),X5))),s(fun(X8,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X8,bool),'type.pair.prod'(fun(fun(X8,bool),bool),fun(fun(X8,bool),'type.realax.real'))),X9)))))))),
  file('thm.probability.joint_distribution_pos','thm.probability.joint_distribution_def')).
fof(f136,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) ) | $spl26),
  inference(cnf_transformation,[],[f136_D])).
fof(f136_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2)))))))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2)))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X9,X21] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X21),s(fun(fun(X6,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9))))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9),s(fun(X6,bool),X21)))))))),
  file('thm.probability.joint_distribution_pos','thm.probability.PROB_POSITIVE')).
fof(f1687,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,'type.pair.prod'(X2,X3)),'_dst_x0x1x1_3'(s(fun(X2,fun(X3,'type.pair.prod'(X2,X3))),'const.pair.,_0'),s(fun(X0,X2),X4),s(fun(X0,X3),X5))),s(fun('type.pair.prod'(X2,X3),bool),X6))),s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,X2),X4),s(fun(X0,X3),X5),s(fun('type.pair.prod'(X2,X3),bool),X6)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) ) | $spl236),
  inference(cnf_transformation,[],[f1687_D])).
fof(f1687_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,'type.pair.prod'(X2,X3)),'_dst_x0x1x1_3'(s(fun(X2,fun(X3,'type.pair.prod'(X2,X3))),'const.pair.,_0'),s(fun(X0,X2),X4),s(fun(X0,X3),X5))),s(fun('type.pair.prod'(X2,X3),bool),X6))),s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,X2),X4),s(fun(X0,X3),X5),s(fun('type.pair.prod'(X2,X3),bool),X6)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f1685,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,X2),X4),s(fun(X0,X3),X5),s(fun('type.pair.prod'(X2,X3),bool),X6)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,'type.pair.prod'(X2,X3)),'_dst_x0x1x1_3'(s(fun(X2,fun(X3,'type.pair.prod'(X2,X3))),'const.pair.,_0'),s(fun(X0,X2),X4),s(fun(X0,X3),X5))),s(fun('type.pair.prod'(X2,X3),bool),X6))),s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) ) | ($spl26 | $spl30)),
  inference(superposition,[],[f136,f138])).
fof(f135,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))) ) | $spl24),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f85,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X22,X23] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X23),s(fun(X6,bool),X22))) = s(bool,'const.bool.IN_2'(s(fun(X6,bool),X23),s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),X22)))))),
  file('thm.probability.joint_distribution_pos','thm.pred_set.IN_POW')).
fof(f121,plain,(
  s(fun(fun(sK0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))) = s(fun(fun(sK0,bool),bool),'const.pred_set.POW_1'(s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) | $spl2),
  inference(cnf_transformation,[],[f121_D])).
fof(f121_D,plain,(
  s(fun(fun(sK0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))) = s(fun(fun(sK0,bool),bool),'const.pred_set.POW_1'(s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f69,plain,(
  s(fun(fun(sK0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))) = s(fun(fun(sK0,bool),bool),'const.pred_set.POW_1'(s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3)))))),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3)))) & s(fun(fun(sK0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))) = s(fun(fun(sK0,bool),bool),'const.pred_set.POW_1'(s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,sK1),sK4),s(fun(sK0,sK2),sK5),s(fun('type.pair.prod'(sK1,sK2),bool),sK6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f49])).
fof(f49,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3)))) & s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3))) = s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3),s(fun(X0,X1),X4),s(fun(X0,X2),X5),s(fun('type.pair.prod'(X1,X2),bool),X6)))))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3)))) & s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3))) = s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3)))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3),s(fun(X0,X1),X4),s(fun(X0,X2),X5),s(fun('type.pair.prod'(X1,X2),bool),X6)))))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3)))) & s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3))) = s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3),s(fun(X0,X1),X4),s(fun(X0,X2),X5),s(fun('type.pair.prod'(X1,X2),bool),X6)))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3)))) & s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3))) = s(fun(fun(X0,bool),bool),'const.pred_set.POW_1'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X3),s(fun(X0,X1),X4),s(fun(X0,X2),X5),s(fun('type.pair.prod'(X1,X2),bool),X6)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11,X12] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9)))) & s(fun(fun(X6,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9))) = s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9),s(fun(X6,X7),X10),s(fun(X6,X8),X11),s(fun('type.pair.prod'(X7,X8),bool),X12)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X12] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9)))) & s(fun(fun(X6,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9))) = s(fun(fun(X6,bool),bool),'const.pred_set.POW_1'(s(fun(X6,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X9),s(fun(X6,X7),X10),s(fun(X6,X8),X11),s(fun('type.pair.prod'(X7,X8),bool),X12)))))))),
  file('thm.probability.joint_distribution_pos','thm.probability.joint_distribution_pos')).
fof(f747,plain,(
  ( ! [X0] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3)))))) ) | $spl114),
  inference(cnf_transformation,[],[f747_D])).
fof(f747_D,plain,(
  ( ! [X0] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3)))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f726,plain,(
  ( ! [X0] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3)))))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f135,f121])).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))),s(fun(X0,bool),X1))))) ) | $spl22),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))),s(fun(X0,bool),X1))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f84,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X21,X15] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X15),s(fun(X6,bool),X21))),s(fun(X6,bool),X21))))),
  file('thm.probability.joint_distribution_pos','thm.pred_set.INTER_SUBSET|split|1')).
fof(f519,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X4,bool),'const.pred_set.INTER_2'(s(fun(X4,bool),X5),s(fun(X4,bool),X6))),s(fun(X4,bool),X6)))) ) | $spl76),
  inference(cnf_transformation,[],[f519_D])).
fof(f519_D,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X4,bool),'const.pred_set.INTER_2'(s(fun(X4,bool),X5),s(fun(X4,bool),X6))),s(fun(X4,bool),X6)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f510,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X4,bool),'const.pred_set.INTER_2'(s(fun(X4,bool),X5),s(fun(X4,bool),X6))),s(fun(X4,bool),X6)))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f134,f129])).
fof(f129,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f129_D])).
fof(f129_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f76,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.probability.joint_distribution_pos','thm.bool.EQ_CLAUSES|split|0')).
fof(f123,plain,(
  p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3)))) | $spl4),
  inference(cnf_transformation,[],[f123_D])).
fof(f123_D,plain,(
  p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f68,plain,(
  p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))),
  inference(cnf_transformation,[],[f57])).
fof(f144,plain,(
  s(bool,'T_0') = s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))) | $spl36),
  inference(cnf_transformation,[],[f144_D])).
fof(f144_D,plain,(
  s(bool,'T_0') = s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f142,plain,(
  s(bool,'T_0') = s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))) | ($spl4 | $spl12)),
  inference(resolution,[],[f129,f123])).
fof(f125,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f125_D])).
fof(f125_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f71,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.probability.joint_distribution_pos',aHL_TRUTH)).
fof(f6006,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,X0),X2),s(fun(sK0,X1),X3),s(fun('type.pair.prod'(X0,X1),bool),X4))))))) ) | $spl1268),
  inference(cnf_transformation,[],[f6006_D])).
fof(f6006_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,X0),X2),s(fun(sK0,X1),X3),s(fun('type.pair.prod'(X0,X1),bool),X4))))))) ) <=> ~$spl1268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1268])])).
fof(f6005,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,X0),X2),s(fun(sK0,X1),X3),s(fun('type.pair.prod'(X0,X1),bool),X4))))))) ) | ($spl6 | $spl36 | $spl76 | $spl114 | $spl236)),
  inference(subsumption_resolution,[],[f6004,f125])).
fof(f6004,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,X0),X2),s(fun(sK0,X1),X3),s(fun('type.pair.prod'(X0,X1),bool),X4))))))) ) | ($spl6 | $spl36 | $spl76 | $spl114 | $spl236)),
  inference(forward_demodulation,[],[f6003,f144])).
fof(f6003,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,X0),X2),s(fun(sK0,X1),X3),s(fun('type.pair.prod'(X0,X1),bool),X4)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) ) | ($spl6 | $spl76 | $spl114 | $spl236)),
  inference(subsumption_resolution,[],[f6002,f125])).
fof(f6002,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,X0),X2),s(fun(sK0,X1),X3),s(fun('type.pair.prod'(X0,X1),bool),X4)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) ) | ($spl76 | $spl114 | $spl236)),
  inference(forward_demodulation,[],[f5959,f519])).
fof(f5959,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,'type.pair.prod'(X0,X1)),'_dst_x0x1x1_3'(s(fun(X0,fun(X1,'type.pair.prod'(X0,X1))),'const.pair.,_0'),s(fun(sK0,X0),X2),s(fun(sK0,X1),X3))),s(fun('type.pair.prod'(X0,X1),bool),X4))),s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))),s(fun(sK0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,X0),X2),s(fun(sK0,X1),X3),s(fun('type.pair.prod'(X0,X1),bool),X4)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3))))) ) | ($spl114 | $spl236)),
  inference(superposition,[],[f1687,f747])).
fof(f131,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f131_D])).
fof(f131_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f80,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.probability.joint_distribution_pos','thm.bool.EQ_CLAUSES|split|3')).
fof(f119,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,sK1),sK4),s(fun(sK0,sK2),sK5),s(fun('type.pair.prod'(sK1,sK2),bool),sK6)))))) | $spl1),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,sK1),sK4),s(fun(sK0,sK2),sK5),s(fun('type.pair.prod'(sK1,sK2),bool),sK6)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f70,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,sK1),sK4),s(fun(sK0,sK2),sK5),s(fun('type.pair.prod'(sK1,sK2),bool),sK6))))))),
  inference(cnf_transformation,[],[f57])).
fof(f209,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,sK1),sK4),s(fun(sK0,sK2),sK5),s(fun('type.pair.prod'(sK1,sK2),bool),sK6))))) | $spl50),
  inference(cnf_transformation,[],[f209_D])).
fof(f209_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,sK1),sK4),s(fun(sK0,sK2),sK5),s(fun('type.pair.prod'(sK1,sK2),bool),sK6))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f205,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.probability.joint_distribution_4'(s('type.pair.prod'(fun(sK0,bool),'type.pair.prod'(fun(fun(sK0,bool),bool),fun(fun(sK0,bool),'type.realax.real'))),sK3),s(fun(sK0,sK1),sK4),s(fun(sK0,sK2),sK5),s(fun('type.pair.prod'(sK1,sK2),bool),sK6))))) | ($spl1 | $spl16)),
  inference(resolution,[],[f119,f131])).
fof(f127,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f73,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.probability.joint_distribution_pos',aHL_FALSITY)).
fof(f6400,plain,(
  $false | ($spl9 | $spl50 | $spl1268)),
  inference(subsumption_resolution,[],[f6387,f127])).
fof(f6387,plain,(
  p(s(bool,'F_0')) | ($spl50 | $spl1268)),
  inference(superposition,[],[f6006,f209])).
% SZS output end Proof for thm.probability
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 80
% Passive clauses: 668
% Generated clauses: 6279
% Final active clauses: 68
% Final passive clauses: 553
% Input formulas: 19
% Initial clauses: 38
% 
% Duplicate literals: 426
% Trivial inequalities: 42
% Fw subsumption resolutions: 323
% Fw demodulations: 705
% Bw demodulations: 41
% Condensations: 370
% 
% Simple tautologies: 105
% Equational tautologies: 19
% Forward subsumptions: 2711
% Fw demodulations to eq. taut.: 9
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 67
% Forward superposition: 2110
% Backward superposition: 1322
% Self superposition: 112
% Equality factoring: 32
% Equality resolution: 6
% 
% Split clauses: 36
% Split components: 72
% SAT solver clauses: 920
% SAT solver unit clauses: 39
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 1
% 
% Sat splits: 36
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 919
% TWLsolver calls for satisfiability: 210
% 
% Memory used [KB]: 15607
% Time elapsed: 0.542 s
% ------------------------------
% ------------------------------
% Success in time 1.163 s
