Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.semi_ring
% (80778)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 162
% Passive clauses: 3258
% Generated clauses: 21212
% Final active clauses: 135
% Final passive clauses: 2907
% Input formulas: 68
% Initial clauses: 81
% 
% Duplicate literals: 795
% Trivial inequalities: 118
% Fw subsumption resolutions: 1683
% Fw demodulations: 3278
% Bw demodulations: 24
% 
% Simple tautologies: 164
% Equational tautologies: 51
% Forward subsumptions: 11407
% 
% Binary resolution: 438
% Factoring: 2
% Forward superposition: 7820
% Backward superposition: 6640
% Self superposition: 215
% Equality factoring: 96
% Equality resolution: 15
% 
% Split clauses: 220
% Split components: 440
% SAT solver clauses: 642
% SAT solver unit clauses: 26
% SAT solver binary clauses: 121
% 
% Sat splits: 220
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 642
% TWLsolver calls for satisfiability: 144
% 
% Memory used [KB]: 20340
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.semi_ring
% (80818)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 119
% Passive clauses: 1129
% Generated clauses: 13703
% Final active clauses: 99
% Final passive clauses: 895
% Input formulas: 68
% Initial clauses: 81
% 
% Duplicate literals: 432
% Trivial inequalities: 83
% Fw subsumption resolutions: 705
% Fw demodulations: 1596
% Bw demodulations: 134
% Condensations: 964
% 
% Simple tautologies: 140
% Equational tautologies: 40
% Forward subsumptions: 6673
% 
% Binary resolution: 271
% Factoring: 1
% Forward superposition: 4449
% Backward superposition: 3570
% Self superposition: 165
% Equality factoring: 76
% Equality resolution: 11
% 
% Split clauses: 48
% Split components: 96
% SAT solver clauses: 1919
% SAT solver unit clauses: 129
% SAT solver binary clauses: 40
% 
% Sat splits: 48
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1919
% TWLsolver calls for satisfiability: 270
% 
% Memory used [KB]: 17654
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.semi_ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.semi_ring
% SZS output start Proof for thm.semi_ring
fof(f52475,plain,(
  $false),
  inference(subsumption_resolution,[],[f52474,f219])).
fof(f219,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7',aHL_TRUTH)).
fof(f52474,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f52470,f4055])).
fof(f4055,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f3990,f305])).
fof(f305,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f221,f222])).
fof(f222,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7',aHL_BOOL_CASES)).
fof(f221,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7',aHL_FALSITY)).
fof(f3990,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f3865,f1947])).
fof(f1947,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1868,f278])).
fof(f278,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7,X11] : ! [X10,X9,X16] : s(X6,'const.combin.o_3'(s(fun(X7,X6),X10),s(fun(X11,X7),X9),s(X11,X16))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,happ(s(fun(X11,X7),X9),s(X11,X16)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7','thm.combin.o_THM')).
fof(f1868,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f223,f276])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X41,X42,X43] : (p(s(bool,happ(s(fun(X41,bool),happ(s(fun(X41,fun(X41,bool)),'=_0'),s(X41,X42))),s(X41,X43)))) <=> s(X41,X42) = s(X41,X43))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7',arity_0_HL_EQ)).
fof(f223,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X31] : ! [X38] : s(X31,X38) = s(X31,happ(s(fun(X31,X31),'_i_0_0'),s(X31,X38)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7','_i0_ax')).
fof(f3865,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))) = s(X2,'const.combin.o_3'(s(fun(X0,X2),X3),s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f278,f223])).
fof(f52470,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))))),
  inference(backward_demodulation,[],[f52366,f14742])).
fof(f14742,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f218,f275])).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f217])).
fof(f218,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f192])).
fof(f192,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))))) != s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))))) = s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))))) = s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8))))) = s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8))))) = s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7','thm.semi_ring.semi_ring_fupdfupds_comp|split|7')).
fof(f52366,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(fun(X11,'type.semi_ring.semi_ring'(X8)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X8),'type.semi_ring.semi_ring'(X8)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X9))),s(fun(X11,'type.semi_ring.semi_ring'(X8)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X8),'type.semi_ring.semi_ring'(X8)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X10))),s(fun(X11,'type.semi_ring.semi_ring'(X8)),X12))))) = s(fun(X11,'type.semi_ring.semi_ring'(X8)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X8),'type.semi_ring.semi_ring'(X8)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),'const.combin.o_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X9),s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X10))))),s(fun(X11,'type.semi_ring.semi_ring'(X8)),X12)))) )),
  inference(superposition,[],[f282,f239])).
fof(f239,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))) = s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SRM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7','thm.semi_ring.semi_ring_fupdfupds_comp|split|6')).
fof(f282,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7,X17,X11] : ! [X10,X9,X8] : s(fun(X17,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X17,X7),'const.combin.o_2'(s(fun(X11,X7),X9),s(fun(X17,X11),X8))))) = s(fun(X17,X6),'const.combin.o_2'(s(fun(X11,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X11,X7),X9))),s(fun(X17,X11),X8)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|7','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.semi_ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 305
% Passive clauses: 5033
% Generated clauses: 59081
% Final active clauses: 128
% Final passive clauses: 1790
% Input formulas: 68
% Initial clauses: 81
% 
% Duplicate literals: 1390
% Trivial inequalities: 153
% Fw subsumption resolutions: 2650
% Fw demodulations: 9217
% Bw demodulations: 2705
% Condensations: 2334
% Global subsumptions: 14
% 
% Simple tautologies: 173
% Equational tautologies: 283
% Forward subsumptions: 36943
% Fw demodulations to eq. taut.: 39
% Bw demodulations to eq. taut.: 79
% 
% Unit resulting resolution: 1657
% Factoring: 1
% Forward superposition: 17026
% Backward superposition: 14151
% Self superposition: 379
% Equality factoring: 144
% Equality resolution: 13
% 
% Split clauses: 94
% Split components: 190
% SAT solver clauses: 5902
% SAT solver unit clauses: 1089
% SAT solver binary clauses: 2573
% 
% Sat splits: 94
% Sat splitting refutations: 665
% 
% Memory used [KB]: 52195
% Time elapsed: 2.924 s
% ------------------------------
% ------------------------------
% Success in time 4.26 s
