Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% (71076)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 764
% Passive clauses: 8429
% Generated clauses: 20190
% Final active clauses: 737
% Final passive clauses: 7053
% Input formulas: 53
% Initial clauses: 210
% 
% Duplicate literals: 558
% Trivial inequalities: 142
% Fw subsumption resolutions: 2004
% Fw demodulations: 500
% Bw demodulations: 22
% 
% Simple tautologies: 261
% Equational tautologies: 63
% Forward subsumptions: 7654
% 
% Binary resolution: 10141
% Factoring: 1406
% Forward superposition: 3140
% Backward superposition: 1751
% Self superposition: 30
% Equality factoring: 28
% Equality resolution: 34
% 
% Split clauses: 584
% Split components: 1175
% SAT solver clauses: 1034
% SAT solver unit clauses: 30
% SAT solver binary clauses: 198
% 
% Sat splits: 584
% 
% TWLsolver clauses: 1034
% TWLsolver calls for satisfiability: 404
% 
% Memory used [KB]: 7803
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (71107)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 398
% Passive clauses: 2207
% Generated clauses: 10004
% Final active clauses: 386
% Final passive clauses: 1785
% Input formulas: 53
% Initial clauses: 210
% 
% Duplicate literals: 256
% Trivial inequalities: 20
% Fw subsumption resolutions: 840
% Fw demodulations: 217
% Bw demodulations: 27
% Condensations: 239
% Global subsumptions: 22
% 
% Simple tautologies: 258
% Equational tautologies: 23
% Forward subsumptions: 3228
% 
% Binary resolution: 3446
% Factoring: 354
% Forward superposition: 1108
% Backward superposition: 966
% Self superposition: 38
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 28
% Split components: 61
% SAT solver clauses: 2870
% SAT solver unit clauses: 216
% SAT solver binary clauses: 94
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 11
% 
% Sat splits: 28
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 2781
% TWLsolver calls for satisfiability: 1760
% 
% Memory used [KB]: 11641
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f26827,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f274,f497,f495,f273,f496,f494,f501,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1728,f1647,f1730,f1651,f1732,f1654,f1734,f1657,f1736,f1660,f1738,f1663,f1740,f1666,f1742,f1669,f1744,f1672,f1746,f1675,f1748,f1678,f1750,f1681,f1752,f1684,f1754,f1687,f1756,f1690,f1758,f1693,f1760,f1696,f1762,f1699,f1785,f1786,f1788,f1790,f1792,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f3987,f3990,f3988,f3992,f3993,f3995,f3997,f3999,f4003,f4005,f4007,f4009,f4013,f4015,f4017,f4019,f4021,f4023,f4033,f4035,f4045,f4047,f4056,f4065,f4081,f4108,f4150,f4152,f4181,f4184,f4190,f4208,f4213,f4231,f4234,f4237,f4242,f4245,f4263,f4305,f4314,f4325,f4334,f4336,f4346,f4354,f4364,f4372,f4374,f4376,f4378,f4380,f5543,f5544,f5545,f5546,f5547,f4405,f4407,f4409,f4411,f4413,f4415,f4442,f4444,f4446,f4448,f4464,f4480,f4501,f4503,f4507,f4516,f4518,f4520,f4522,f4528,f4530,f4552,f4554,f4558,f4567,f4569,f4571,f4573,f4579,f4581,f4592,f4600,f4610,f4618,f4620,f4643,f4645,f4658,f4660,f4683,f4685,f4698,f4700,f4725,f4727,f4740,f4742,f4755,f4771,f4784,f4800,f4818,f4836,f4857,f4859,f4864,f4889,f4891,f4893,f4895,f4905,f4907,f4909,f4921,f4923,f4935,f4937,f4949,f4951,f4963,f4965,f4985,f5002,f5023,f5025,f5030,f5038,f5040,f5042,f5044,f5051,f5053,f5093,f5095,f5097,f5099,f5101,f5103,f5129,f5131,f5133,f5135,f5161,f5163,f5165,f5167,f5183,f5185,f5201,f5203,f5205,f5218,f5231,f5244,f5246,f5259,f5261,f5275,f5292,f5311,f5328,f5341,f5343,f5345,f5347,f5349,f5351,f5363,f5365,f5367,f5369,f5371,f5373,f5380,f5386,f5392,f5398,f5400,f5406,f5408,f5417,f5419,f5421,f5427,f5429,f5435,f5440,f5442,f5447,f5449,f5454,f5456,f5460,f5462,f5466,f5468,f5471,f5473,f5480,f5482,f5487,f5489,f5491,f5494,f5496,f5503,f5505,f5510,f5512,f5550,f5551,f5552,f5553,f5554,f5555,f5556,f5557,f5558,f5559,f5560,f5561,f5562,f5563,f5564,f5565,f5566,f5567,f5568,f5569,f5570,f5571,f5572,f5573,f5574,f5575,f5576,f5577,f5578,f5579,f5580,f5581,f5582,f5583,f5584,f5585,f5586,f5587,f5588,f5589,f5590,f5591,f5592,f5593,f5594,f5595,f5596,f5597,f5598,f5599,f5600,f5601,f5602,f5603,f5604,f5605,f5606,f5607,f5608,f5609,f5610,f5611,f5612,f5613,f5614,f5615,f5616,f5617,f5618,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f5631,f5632,f5633,f5634,f5635,f5636,f5637,f5638,f5639,f5640,f5641,f5642,f5643,f5644,f5645,f5646,f5647,f5648,f5649,f5650,f5651,f5652,f5653,f5654,f5655,f5656,f5657,f5658,f5659,f5660,f5661,f5662,f5663,f5664,f5665,f5666,f5667,f5668,f5669,f5670,f5671,f5672,f5673,f5674,f5675,f5676,f5677,f5678,f5679,f5680,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5690,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f5700,f5701,f5702,f5703,f5704,f5705,f5706,f5707,f5708,f5709,f5710,f5711,f5712,f5713,f5714,f5715,f5716,f5717,f5718,f5719,f5720,f5721,f5722,f5723,f5724,f5725,f5726,f5727,f5728,f5729,f5730,f5731,f5732,f5733,f5734,f5735,f5736,f5737,f5738,f5739,f5740,f5741,f5742,f5743,f5744,f5745,f5746,f5747,f5748,f5749,f5750,f5751,f5752,f5753,f5754,f5755,f5756,f5757,f5758,f5759,f5760,f5761,f5762,f5763,f5764,f5765,f5766,f5767,f5768,f5769,f5770,f5771,f5772,f5773,f5774,f5775,f5776,f5777,f5778,f5779,f5780,f5781,f5782,f5783,f5784,f5785,f5786,f5787,f5788,f5789,f5790,f5791,f5792,f5793,f5794,f5795,f5796,f5797,f5798,f5799,f5800,f5801,f5802,f5803,f5804,f5805,f5806,f5808,f5809,f5812,f5815,f5817,f5819,f5821,f5823,f5824,f5825,f5826,f5829,f5830,f5831,f5832,f5833,f5834,f5839,f5840,f5841,f5842,f5843,f5844,f5845,f5846,f5847,f5848,f5849,f5850,f5851,f5852,f5853,f5854,f5855,f5856,f5857,f5858,f5859,f5860,f5861,f5862,f5863,f5864,f5865,f5866,f5867,f5868,f5869,f5870,f5871,f5872,f5873,f5874,f5875,f5876,f5877,f5878,f5879,f5880,f5881,f5882,f5883,f5884,f5885,f5886,f5887,f5888,f5889,f5890,f5891,f5892,f5893,f5894,f5895,f5896,f5897,f5898,f5899,f5900,f5901,f5902,f5903,f5904,f5905,f5906,f5907,f5908,f5909,f5910,f5911,f5912,f5913,f5914,f5915,f5916,f5917,f5918,f5919,f5920,f5921,f5966,f5968,f5967,f5970,f5971,f6474,f6477,f8007,f7996,f8040,f7990,f7981,f8032,f7880,f8358,f8357,f8506,f8507,f8522,f8526,f9716,f9717,f9719,f9720,f9722,f9723,f7879,f9725,f8088,f2337,f9726,f9727,f6475,f2339,f9816,f9817,f9876,f9877,f12769,f12771,f12770,f1648,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f2307,f10221,f7874,f2310,f12779,f7622,f12775,f7617,f12777,f8157,f12780,f2312,f2313,f2330,f2315,f2329,f2328,f12774,f17698,f17699,f17700,f17702,f17704,f17706,f17709,f17711,f17778,f17785,f17787,f15130,f12776,f17791,f17792,f17793,f17927,f17701,f17928,f17703,f17929,f17705,f17930,f17707,f17932,f17710,f17933,f17712,f17997,f17779,f18003,f17786,f18004,f17788,f18192,f18197,f18193,f18296,f18301,f18297,f19481,f19482,f19516,f19483,f19487,f19517,f2308,f19526,f19604,f19606,f19608,f19610,f19611,f19823,f19824,f19826,f19827,f19829,f19832,f19835,f19836,f19837,f19838,f19839,f19840,f19841,f19842,f19843,f19844,f19845,f19846,f19847,f19848,f19849,f19850,f19851,f19852,f19853,f19854,f19855,f19856,f19857,f19858,f19859,f19860,f19862,f19864,f19866,f19867,f19868,f19869,f19870,f19871,f19872,f19874,f19876,f19878,f19880,f19882,f24956,f24958,f24957,f25604,f25609,f25605,f26326,f26331,f3989,f26327,f26808,f26817])).
fof(f26817,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f26816,f494])).
fof(f26816,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | $spl3),
  inference(forward_demodulation,[],[f26815,f355])).
fof(f355,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X19,X20] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X19)))),
  file('thm.rat.RAT_EQ_RMUL','thm.rat.RAT_MUL_COMM')).
fof(f26815,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) | $spl3),
  inference(forward_demodulation,[],[f26814,f355])).
fof(f26814,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) | $spl3),
  inference(forward_demodulation,[],[f26813,f356])).
fof(f356,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X24,X25] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X24))),s('type.rat.rat',X25)))),
  file('thm.rat.RAT_EQ_RMUL','ARITY_1_2_const.rat.rat_mul')).
fof(f26813,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',sK12))) | $spl3),
  inference(forward_demodulation,[],[f26771,f356])).
fof(f26771,plain,(
  s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',sK12))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f497,f272,f2171])).
fof(f2171,plain,(
  ( ! [X2,X3,X1] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X1))),s('type.rat.rat',X2))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X1))),s('type.rat.rat',X3))) | s('type.rat.rat',X2) = s('type.rat.rat',X3) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2126,f275])).
fof(f275,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_EQ_RMUL',aHL_TRUTH)).
fof(f2126,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X1))),s('type.rat.rat',X2))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X1))),s('type.rat.rat',X3))) | s('type.rat.rat',X2) = s('type.rat.rat',X3) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f398,f872])).
fof(f872,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f870,f498])).
fof(f498,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f278,f275])).
fof(f278,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.rat.RAT_EQ_RMUL','thm.bool.NOT_CLAUSES|split|2')).
fof(f870,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))) )),
  inference(superposition,[],[f347,f283])).
fof(f283,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_EQ_RMUL',aHL_BOOL_CASES)).
fof(f347,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0] : ((s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))))))) & (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))) | s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : (~s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : (~s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X6)))))))),
  file('thm.rat.RAT_EQ_RMUL','thm.rat.RAT_MUL_ONE_ONE')).
fof(f398,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & ((s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK16(X2,X1,X0)))) & s(X1,sK15(X2,X1,X0)) != s(X1,sK16(X2,X1,X0))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(rectify,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4)))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X14,X12] : ! [X16] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X12,X14),X16)))) <=> ! [X17,X18] : (s(X14,happ(s(fun(X12,X14),X16),s(X12,X17))) = s(X14,happ(s(fun(X12,X14),X16),s(X12,X18))) => s(X12,X17) = s(X12,X18)))),
  file('thm.rat.RAT_EQ_RMUL','thm.bool.ONE_ONE_THM')).
fof(f272,plain,(
  s('type.rat.rat',sK14) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  s('type.rat.rat',sK14) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) & (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | s('type.rat.rat',sK12) = s('type.rat.rat',sK13)) & (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | s('type.rat.rat',sK12) != s('type.rat.rat',sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f173])).
fof(f173,plain,(
  ? [X0,X1,X2] : (s('type.rat.rat',X2) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) & (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X0) != s('type.rat.rat',X1)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ? [X0,X1,X2] : (s('type.rat.rat',X2) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ((s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X0) != s('type.rat.rat',X1))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ? [X0,X1,X2] : (s('type.rat.rat',X2) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) & (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) <~> s('type.rat.rat',X0) = s('type.rat.rat',X1)))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ~! [X0,X1,X2] : (s('type.rat.rat',X2) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) <=> s('type.rat.rat',X0) = s('type.rat.rat',X1)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ~! [X0,X1,X2] : (~s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) <=> s('type.rat.rat',X0) = s('type.rat.rat',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (~s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X8))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (~s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X8))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7)))),
  file('thm.rat.RAT_EQ_RMUL','thm.rat.RAT_EQ_RMUL')).
fof(f26808,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f26807,f494])).
fof(f26807,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | $spl3),
  inference(forward_demodulation,[],[f26806,f355])).
fof(f26806,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) | $spl3),
  inference(forward_demodulation,[],[f26805,f355])).
fof(f26805,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK12))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) | $spl3),
  inference(forward_demodulation,[],[f26804,f356])).
fof(f26804,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK14),s('type.rat.rat',sK13))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',sK12))) | $spl3),
  inference(forward_demodulation,[],[f26774,f356])).
fof(f26774,plain,(
  s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',sK12))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',sK13))) | $spl3),
  inference(unit_resulting_resolution,[],[f497,f272,f2171])).
fof(f26327,plain,(
  ( ! [X28,X24,X27,X25] : (s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,X24,bool)))) | s(bool,'T_0') = s(bool,X28) | s(bool,X27) = s(bool,X28) | s(bool,X27) = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK15(X25,X24,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X24,bool),X25)))) ) | $spl140),
  inference(cnf_transformation,[],[f26327_D])).
fof(f26327_D,plain,(
  ( ! [X28,X24,X27,X25] : (s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,X24,bool)))) | s(bool,'T_0') = s(bool,X28) | s(bool,X27) = s(bool,X28) | s(bool,X27) = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK15(X25,X24,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X24,bool),X25)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f3989,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) | $spl60),
  inference(cnf_transformation,[],[f3989_D])).
fof(f3989_D,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f26331,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,happ(s(fun(X44,bool),X45),s(X44,sK16(X45,X44,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X44,bool),X45))) | s(bool,X47) = s(bool,happ(s(fun(X44,bool),X45),s(X44,sK15(X45,X44,bool)))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f26310,f294])).
fof(f294,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.rat.RAT_EQ_RMUL','thm.bool.EQ_CLAUSES|split|0')).
fof(f26310,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,happ(s(fun(X44,bool),X45),s(X44,sK16(X45,X44,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X44,bool),X45))) | s(bool,X47) = s(bool,happ(s(fun(X44,bool),X45),s(X44,sK15(X45,X44,bool)))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f2373,f1281])).
fof(f1281,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f608,f608])).
fof(f608,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f283,f283])).
fof(f2373,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,X8,bool))))) | s(bool,'F_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,X8,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X8,bool),X9)))) )),
  inference(superposition,[],[f299,f1067])).
fof(f1067,plain,(
  ( ! [X2,X0,X1] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK15(X2,X0,X1)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK16(X2,X0,X1)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X0,X1),X2)))) )),
  inference(subsumption_resolution,[],[f1066,f498])).
fof(f1066,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK15(X2,X0,X1)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK16(X2,X0,X1)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X0,X1),X2)))) )),
  inference(superposition,[],[f399,f283])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK16(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f299,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.rat.RAT_EQ_RMUL','thm.bool.EQ_CLAUSES|split|2')).
fof(f26326,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X26)) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,X24,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X24,bool),X25))) | s(bool,X27) = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK15(X25,X24,bool)))) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f26306,f294])).
fof(f26306,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X26)) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,X24,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X24,bool),X25))) | s(bool,X27) = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK15(X25,X24,bool)))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f2373,f1281])).
fof(f25605,plain,(
  ( ! [X28,X24,X27,X25] : (s(bool,'T_0') != s(bool,happ(s(fun(X24,bool),X25),s(X24,sK15(X25,X24,bool)))) | s(bool,'T_0') = s(bool,X28) | s(bool,X27) = s(bool,X28) | s(bool,X27) = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,X24,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X24,bool),X25)))) ) | $spl138),
  inference(cnf_transformation,[],[f25605_D])).
fof(f25605_D,plain,(
  ( ! [X28,X24,X27,X25] : (s(bool,'T_0') != s(bool,happ(s(fun(X24,bool),X25),s(X24,sK15(X25,X24,bool)))) | s(bool,'T_0') = s(bool,X28) | s(bool,X27) = s(bool,X28) | s(bool,X27) = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,X24,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X24,bool),X25)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f25609,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,X46)) | s(bool,'T_0') != s(bool,happ(s(fun(X44,bool),X45),s(X44,sK15(X45,X44,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X44,bool),X45))) | s(bool,X47) = s(bool,happ(s(fun(X44,bool),X45),s(X44,sK16(X45,X44,bool)))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f25588,f294])).
fof(f25588,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,X46)) | s(bool,'T_0') != s(bool,happ(s(fun(X44,bool),X45),s(X44,sK15(X45,X44,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X44,bool),X45))) | s(bool,X47) = s(bool,happ(s(fun(X44,bool),X45),s(X44,sK16(X45,X44,bool)))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f2370,f1281])).
fof(f2370,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK16(X3,X2,bool))))) | s(bool,'T_0') != s(bool,happ(s(fun(X2,bool),X3),s(X2,sK15(X3,X2,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f294,f1067])).
fof(f25604,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X26)) | s(bool,'T_0') != s(bool,happ(s(fun(X24,bool),X25),s(X24,sK15(X25,X24,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X24,bool),X25))) | s(bool,X27) = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,X24,bool)))) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f25584,f294])).
fof(f25584,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X26)) | s(bool,'T_0') != s(bool,happ(s(fun(X24,bool),X25),s(X24,sK15(X25,X24,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X24,bool),X25))) | s(bool,X27) = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,X24,bool)))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f2370,f1281])).
fof(f24957,plain,(
  ( ! [X19,X17,X20] : (s(bool,'F_0') = s(bool,f(s(bool,X17))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X20) | s(bool,X17) = s(bool,X19)) ) | $spl136),
  inference(cnf_transformation,[],[f24957_D])).
fof(f24957_D,plain,(
  ( ! [X19,X17,X20] : (s(bool,'F_0') = s(bool,f(s(bool,X17))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X20) | s(bool,X17) = s(bool,X19)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f24958,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X34)) | s(bool,'F_0') = s(bool,f(s(bool,X33))) | s(bool,X33) = s(bool,X35) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f24928,f294])).
fof(f24928,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X34)) | s(bool,'F_0') = s(bool,f(s(bool,X33))) | s(bool,X33) = s(bool,X35) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f24766,f1281])).
fof(f24766,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'F_0') = s(bool,f(s(bool,X6)))) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f24658,f294])).
fof(f24658,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'F_0') = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f299,f13404])).
fof(f13404,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13099,f283])).
fof(f13099,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f275,f2310,f9608])).
fof(f9608,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) ) | $spl74),
  inference(subsumption_resolution,[],[f9607,f275])).
fof(f9607,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) ) | $spl74),
  inference(forward_demodulation,[],[f9480,f877])).
fof(f877,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f275,f866,f360])).
fof(f360,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.rat.RAT_EQ_RMUL','thm.bool.IMP_ANTISYM_AX')).
fof(f866,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))))),
  inference(unit_resulting_resolution,[],[f272,f347])).
fof(f9480,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) ) | $spl74),
  inference(superposition,[],[f360,f8040])).
fof(f24956,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') = s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,X19) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f24924,f294])).
fof(f24924,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') = s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f24766,f1281])).
fof(f19882,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737,f19738,f19739,f19820,f19873,f19875,f19877,f19879,f19881])).
fof(f19881,plain,(
  ( ! [X2,X3,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1)))))) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19798,f275])).
fof(f19798,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1)))))) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f2274])).
fof(f19879,plain,(
  ( ! [X2,X3,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))),s('type.rat.rat',X3)))) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19796,f275])).
fof(f19796,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))),s('type.rat.rat',X3)))) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f2270])).
fof(f19877,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19794,f275])).
fof(f19794,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f1018])).
fof(f19875,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19789,f275])).
fof(f19789,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f348])).
fof(f19873,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19781,f275])).
fof(f19781,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f398])).
fof(f19820,plain,(
  ( ! [X318,X320,X317,X319] : (s(fun(X317,bool),X318) = s(fun(X317,bool),X319) | s(bool,X320) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X320) = s(bool,g(s(X317,sK17(X318,X319,bool,X317))))) ) | ($spl47 | $spl74 | $spl98)),
  inference(trivial_inequality_removal,[],[f19767])).
fof(f19767,plain,(
  ( ! [X318,X320,X317,X319] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X317,bool),X318) = s(fun(X317,bool),X319) | s(bool,X320) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X320) = s(bool,g(s(X317,sK17(X318,X319,bool,X317))))) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f7873])).
fof(f19739,plain,(
  sP3('T_0',f(s(bool,'T_0')),'F_0') | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f14056])).
fof(f19738,plain,(
  ~sP4('T_0','T_0',f(s(bool,'T_0'))) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f14024])).
fof(f19737,plain,(
  sP3(f(s(bool,'T_0')),'T_0','F_0') | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f14011])).
fof(f19688,plain,(
  ~sP5('T_0','T_0',f(s(bool,'T_0'))) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13090])).
fof(f19687,plain,(
  sP5('F_0','T_0',f(s(bool,'T_0'))) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13089])).
fof(f19686,plain,(
  ~sP5('T_0',f(s(bool,'T_0')),'T_0') | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f12980])).
fof(f19685,plain,(
  sP5('F_0',f(s(bool,'T_0')),'T_0') | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f12979])).
fof(f19831,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | ($spl47 | $spl74 | $spl98)),
  inference(condensation,[],[f19830])).
fof(f19830,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X6) = s(bool,X7)) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19635,f19612])).
fof(f19635,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X6) = s(bool,X7)) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f8301])).
fof(f19828,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19631,f19612])).
fof(f19631,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f7001])).
fof(f19825,plain,(
  ( ! [X8] : (s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19620,f19612])).
fof(f19620,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X8))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f2172])).
fof(f19822,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19613,f19612])).
fof(f19613,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f298])).
fof(f19612,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19603,f13099])).
fof(f19603,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl98),
  inference(duplicate_literal_removal,[],[f19528])).
fof(f19528,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl98),
  inference(superposition,[],[f12776,f283])).
fof(f19472,plain,(
  ( ! [X173,X171,X174,X172,X170] : (s(X172,f(s(bool,happ(s(fun(X170,bool),X171),s(X170,sK15(X171,X170,bool)))))) != s(X172,g(s(bool,happ(s(fun(X170,bool),X171),s(X170,sK15(X171,X170,bool)))))) | s(fun(bool,X172),X173) = s(fun(bool,X172),X174) | s(bool,'T_0') = s(bool,happ(s(fun(X170,bool),X171),s(X170,sK16(X171,X170,bool)))) | s(bool,'T_0') = s(bool,sK17(X173,X174,X172,bool)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X170,bool),X171)))) )),
  inference(superposition,[],[f1340,f1067])).
fof(f19413,plain,(
  ( ! [X66,X65] : (sP4('T_0',f(s(bool,happ(s(fun(X65,bool),X66),s(X65,sK15(X66,X65,bool))))),'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK16(X66,X65,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,bool),X66)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13722,f1067])).
fof(f19350,plain,(
  ( ! [X66,X65] : (~sP2(f(s(bool,happ(s(fun(X65,bool),X66),s(X65,sK15(X66,X65,bool))))),'F_0','F_0') | s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK16(X66,X65,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,bool),X66)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13702,f1067])).
fof(f19310,plain,(
  ( ! [X85,X83,X84] : (sP0(happ(s(fun(X83,bool),X84),s(X83,sK16(X84,X83,bool))),X85) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X83,bool),X84))) | s(bool,'F_0') = s(bool,happ(s(fun(X83,bool),X84),s(X83,sK15(X84,X83,bool))))) ) | $spl74),
  inference(subsumption_resolution,[],[f19309,f275])).
fof(f19309,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X83,bool),X84),s(X83,sK16(X84,X83,bool))),X85) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X83,bool),X84))) | s(bool,'F_0') = s(bool,happ(s(fun(X83,bool),X84),s(X83,sK15(X84,X83,bool))))) ) | $spl74),
  inference(forward_demodulation,[],[f19292,f877])).
fof(f19292,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP0(happ(s(fun(X83,bool),X84),s(X83,sK16(X84,X83,bool))),X85) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X83,bool),X84))) | s(bool,'F_0') = s(bool,happ(s(fun(X83,bool),X84),s(X83,sK15(X84,X83,bool))))) ) | $spl74),
  inference(superposition,[],[f2378,f8040])).
fof(f19305,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (sP0(happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool))),X45) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,happ(s(fun(X43,bool),X44),s(X43,sK15(X44,X43,bool)))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f19284,f275])).
fof(f19284,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool))),X45) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,happ(s(fun(X43,bool),X44),s(X43,sK15(X44,X43,bool)))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f2378,f1281])).
fof(f19298,plain,(
  ( ! [X6,X4,X5,X3] : (sP0(happ(s(fun(X3,bool),X4),s(X3,sK16(X4,X3,bool))),X5) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X3,bool),X4))) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK15(X4,X3,bool)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f19276,f275])).
fof(f19276,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X3,bool),X4),s(X3,sK16(X4,X3,bool))),X5) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X3,bool),X4))) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK15(X4,X3,bool)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2378,f608])).
fof(f19297,plain,(
  ( ! [X66,X67] : (~p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'T_0')))) | sP0(happ(s(fun(bool,bool),X66),s(bool,sK16(X66,bool,bool))),X67) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X66))) | s(bool,'F_0') = s(bool,sK15(X66,bool,bool))) ) | $spl74),
  inference(forward_demodulation,[],[f19274,f877])).
fof(f19274,plain,(
  ( ! [X66,X67] : (~p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))))) | sP0(happ(s(fun(bool,bool),X66),s(bool,sK16(X66,bool,bool))),X67) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X66))) | s(bool,'F_0') = s(bool,sK15(X66,bool,bool))) ) | $spl74),
  inference(superposition,[],[f2378,f8040])).
fof(f19266,plain,(
  ( ! [X37,X35,X38,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0')))) | sP0(happ(s(fun(bool,bool),X34),s(bool,sK16(X34,bool,bool))),X35) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X34))) | s(bool,X36) = s(bool,X37) | s(bool,X37) = s(bool,sK15(X34,bool,bool)) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f2378,f1281])).
fof(f19265,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) | sP0(happ(s(fun(bool,bool),X29),s(bool,sK16(X29,bool,bool))),X31) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X29))) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,sK15(X29,bool,bool)) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f2378,f1281])).
fof(f19264,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X25)))) | sP0(happ(s(fun(bool,bool),X24),s(bool,sK16(X24,bool,bool))),X26) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X24))) | s(bool,X27) = s(bool,sK15(X24,bool,bool)) | s(bool,'T_0') = s(bool,X25) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f2378,f1281])).
fof(f19262,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | sP0(happ(s(fun(bool,bool),X14),s(bool,sK16(X14,bool,bool))),X16) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X14))) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,sK15(X14,bool,bool)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f2378,f1281])).
fof(f19259,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6)))) | sP0(happ(s(fun(bool,bool),X5),s(bool,sK16(X5,bool,bool))),X7) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,sK15(X5,bool,bool))) )),
  inference(superposition,[],[f2378,f608])).
fof(f19258,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) | sP0(happ(s(fun(bool,bool),X2),s(bool,sK16(X2,bool,bool))),X3) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,sK15(X2,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2378,f608])).
fof(f19257,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | sP0(happ(s(fun(bool,bool),X0),s(bool,sK16(X0,bool,bool))),X1) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,sK15(X0,bool,bool))) )),
  inference(superposition,[],[f2378,f283])).
fof(f19256,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK15(X2,X1,fun(X0,bool))))),s(X0,sK15(happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK16(X2,X1,fun(X0,bool)))),X0,bool))))) | sP0(happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK16(X2,X1,fun(X0,bool))))),s(X0,sK16(happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK16(X2,X1,fun(X0,bool)))),X0,bool))),X3) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK16(X2,X1,fun(X0,bool))))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,fun(X0,bool)),X2)))) )),
  inference(superposition,[],[f2378,f1067])).
fof(f19232,plain,(
  ( ! [X66,X65] : (~sP2('F_0',f(s(bool,happ(s(fun(X65,bool),X66),s(X65,sK15(X66,X65,bool))))),'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK16(X66,X65,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,bool),X66)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13671,f1067])).
fof(f19169,plain,(
  ( ! [X66,X65] : (sP5('T_0','F_0',f(s(bool,happ(s(fun(X65,bool),X66),s(X65,sK15(X66,X65,bool)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK16(X66,X65,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,bool),X66)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13656,f1067])).
fof(f19129,plain,(
  ( ! [X85,X83,X84] : (sP0(X85,happ(s(fun(X83,bool),X84),s(X83,sK16(X84,X83,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X83,bool),X84))) | s(bool,'F_0') = s(bool,happ(s(fun(X83,bool),X84),s(X83,sK15(X84,X83,bool))))) ) | $spl74),
  inference(subsumption_resolution,[],[f19128,f275])).
fof(f19128,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,'T_0')) | sP0(X85,happ(s(fun(X83,bool),X84),s(X83,sK16(X84,X83,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X83,bool),X84))) | s(bool,'F_0') = s(bool,happ(s(fun(X83,bool),X84),s(X83,sK15(X84,X83,bool))))) ) | $spl74),
  inference(forward_demodulation,[],[f19111,f877])).
fof(f19111,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP0(X85,happ(s(fun(X83,bool),X84),s(X83,sK16(X84,X83,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X83,bool),X84))) | s(bool,'F_0') = s(bool,happ(s(fun(X83,bool),X84),s(X83,sK15(X84,X83,bool))))) ) | $spl74),
  inference(superposition,[],[f2377,f8040])).
fof(f19124,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (sP0(X45,happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,happ(s(fun(X43,bool),X44),s(X43,sK15(X44,X43,bool)))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f19103,f275])).
fof(f19103,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~p(s(bool,'T_0')) | sP0(X45,happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,happ(s(fun(X43,bool),X44),s(X43,sK15(X44,X43,bool)))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f2377,f1281])).
fof(f19117,plain,(
  ( ! [X6,X4,X5,X3] : (sP0(X5,happ(s(fun(X3,bool),X4),s(X3,sK16(X4,X3,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X3,bool),X4))) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK15(X4,X3,bool)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f19095,f275])).
fof(f19095,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | sP0(X5,happ(s(fun(X3,bool),X4),s(X3,sK16(X4,X3,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X3,bool),X4))) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK15(X4,X3,bool)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2377,f608])).
fof(f19116,plain,(
  ( ! [X66,X67] : (~p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'T_0')))) | sP0(X67,happ(s(fun(bool,bool),X66),s(bool,sK16(X66,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X66))) | s(bool,'F_0') = s(bool,sK15(X66,bool,bool))) ) | $spl74),
  inference(forward_demodulation,[],[f19093,f877])).
fof(f19093,plain,(
  ( ! [X66,X67] : (~p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))))) | sP0(X67,happ(s(fun(bool,bool),X66),s(bool,sK16(X66,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X66))) | s(bool,'F_0') = s(bool,sK15(X66,bool,bool))) ) | $spl74),
  inference(superposition,[],[f2377,f8040])).
fof(f19085,plain,(
  ( ! [X37,X35,X38,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0')))) | sP0(X35,happ(s(fun(bool,bool),X34),s(bool,sK16(X34,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X34))) | s(bool,X36) = s(bool,X37) | s(bool,X37) = s(bool,sK15(X34,bool,bool)) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f2377,f1281])).
fof(f19084,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) | sP0(X31,happ(s(fun(bool,bool),X29),s(bool,sK16(X29,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X29))) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,sK15(X29,bool,bool)) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f2377,f1281])).
fof(f19083,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X25)))) | sP0(X26,happ(s(fun(bool,bool),X24),s(bool,sK16(X24,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X24))) | s(bool,X27) = s(bool,sK15(X24,bool,bool)) | s(bool,'T_0') = s(bool,X25) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f2377,f1281])).
fof(f19081,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | sP0(X16,happ(s(fun(bool,bool),X14),s(bool,sK16(X14,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X14))) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,sK15(X14,bool,bool)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f2377,f1281])).
fof(f19078,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6)))) | sP0(X7,happ(s(fun(bool,bool),X5),s(bool,sK16(X5,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,sK15(X5,bool,bool))) )),
  inference(superposition,[],[f2377,f608])).
fof(f19077,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) | sP0(X3,happ(s(fun(bool,bool),X2),s(bool,sK16(X2,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,sK15(X2,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2377,f608])).
fof(f19076,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | sP0(X1,happ(s(fun(bool,bool),X0),s(bool,sK16(X0,bool,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,sK15(X0,bool,bool))) )),
  inference(superposition,[],[f2377,f283])).
fof(f19075,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK15(X2,X1,fun(X0,bool))))),s(X0,sK15(happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK16(X2,X1,fun(X0,bool)))),X0,bool))))) | sP0(X3,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK16(X2,X1,fun(X0,bool))))),s(X0,sK16(happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK16(X2,X1,fun(X0,bool)))),X0,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK16(X2,X1,fun(X0,bool))))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,fun(X0,bool)),X2)))) )),
  inference(superposition,[],[f2377,f1067])).
fof(f19051,plain,(
  ( ! [X66,X65] : (~sP5('F_0','F_0',f(s(bool,happ(s(fun(X65,bool),X66),s(X65,sK15(X66,X65,bool)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK16(X66,X65,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,bool),X66)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13625,f1067])).
fof(f18988,plain,(
  ( ! [X66,X65] : (sP5('T_0',f(s(bool,happ(s(fun(X65,bool),X66),s(X65,sK15(X66,X65,bool))))),'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK16(X66,X65,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,bool),X66)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13609,f1067])).
fof(f18948,plain,(
  ( ! [X241,X239,X240,X238] : (s(bool,X241) = s(bool,g(s(X238,sK17(X239,X240,bool,X238)))) | s(bool,'F_0') = s(bool,X241) | s(fun(X238,bool),X239) = s(fun(X238,bool),X240) | s(bool,'T_0') = s(bool,f(s(X238,sK17(X239,X240,bool,X238))))) ) | $spl74),
  inference(subsumption_resolution,[],[f18863,f275])).
fof(f18863,plain,(
  ( ! [X241,X239,X240,X238] : (~p(s(bool,'T_0')) | s(bool,X241) = s(bool,g(s(X238,sK17(X239,X240,bool,X238)))) | s(bool,'F_0') = s(bool,X241) | s(fun(X238,bool),X239) = s(fun(X238,bool),X240) | s(bool,'T_0') = s(bool,f(s(X238,sK17(X239,X240,bool,X238))))) ) | $spl74),
  inference(superposition,[],[f9608,f2341])).
fof(f18943,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,g(s(X213,sK17(X214,X215,bool,X213)))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X216))))) | s('type.rat.rat',X216) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(X213,bool),X214) = s(fun(X213,bool),X215) | s(bool,'T_0') = s(bool,f(s(X213,sK17(X214,X215,bool,X213))))) )),
  inference(subsumption_resolution,[],[f18857,f275])).
fof(f18857,plain,(
  ( ! [X216,X215,X213,X214] : (~p(s(bool,'T_0')) | s(bool,g(s(X213,sK17(X214,X215,bool,X213)))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X216))))) | s('type.rat.rat',X216) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(X213,bool),X214) = s(fun(X213,bool),X215) | s(bool,'T_0') = s(bool,f(s(X213,sK17(X214,X215,bool,X213))))) )),
  inference(superposition,[],[f2173,f2341])).
fof(f18942,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~sP8(g(s(X118,sK17(X119,X120,bool,X118))),X121,X122) | sP9(g(s(X118,sK17(X119,X120,bool,X118))),X121,X122) | s(fun(X118,bool),X119) = s(fun(X118,bool),X120) | s(bool,'T_0') = s(bool,f(s(X118,sK17(X119,X120,bool,X118))))) )),
  inference(subsumption_resolution,[],[f18840,f275])).
fof(f18840,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,'T_0')) | ~sP8(g(s(X118,sK17(X119,X120,bool,X118))),X121,X122) | sP9(g(s(X118,sK17(X119,X120,bool,X118))),X121,X122) | s(fun(X118,bool),X119) = s(fun(X118,bool),X120) | s(bool,'T_0') = s(bool,f(s(X118,sK17(X119,X120,bool,X118))))) )),
  inference(superposition,[],[f483,f2341])).
fof(f18941,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~sP10(g(s(X113,sK17(X114,X115,bool,X113))),X116,X117) | sP11(g(s(X113,sK17(X114,X115,bool,X113))),X116,X117) | s(fun(X113,bool),X114) = s(fun(X113,bool),X115) | s(bool,'T_0') = s(bool,f(s(X113,sK17(X114,X115,bool,X113))))) )),
  inference(subsumption_resolution,[],[f18839,f275])).
fof(f18839,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,'T_0')) | ~sP10(g(s(X113,sK17(X114,X115,bool,X113))),X116,X117) | sP11(g(s(X113,sK17(X114,X115,bool,X113))),X116,X117) | s(fun(X113,bool),X114) = s(fun(X113,bool),X115) | s(bool,'T_0') = s(bool,f(s(X113,sK17(X114,X115,bool,X113))))) )),
  inference(superposition,[],[f482,f2341])).
fof(f18940,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,X111)) | p(s(bool,X112)) | ~sP11(g(s(X108,sK17(X109,X110,bool,X108))),X112,X111) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,'T_0') = s(bool,f(s(X108,sK17(X109,X110,bool,X108))))) )),
  inference(subsumption_resolution,[],[f18838,f275])).
fof(f18838,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,'T_0')) | ~p(s(bool,X111)) | p(s(bool,X112)) | ~sP11(g(s(X108,sK17(X109,X110,bool,X108))),X112,X111) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,'T_0') = s(bool,f(s(X108,sK17(X109,X110,bool,X108))))) )),
  inference(superposition,[],[f479,f2341])).
fof(f18939,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,X106)) | sP11(g(s(X103,sK17(X104,X105,bool,X103))),X106,X107) | s(fun(X103,bool),X104) = s(fun(X103,bool),X105) | s(bool,'T_0') = s(bool,f(s(X103,sK17(X104,X105,bool,X103))))) )),
  inference(subsumption_resolution,[],[f18837,f275])).
fof(f18837,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X106)) | sP11(g(s(X103,sK17(X104,X105,bool,X103))),X106,X107) | s(fun(X103,bool),X104) = s(fun(X103,bool),X105) | s(bool,'T_0') = s(bool,f(s(X103,sK17(X104,X105,bool,X103))))) )),
  inference(superposition,[],[f478,f2341])).
fof(f18938,plain,(
  ( ! [X101,X99,X102,X100,X98] : (p(s(bool,X101)) | sP11(g(s(X98,sK17(X99,X100,bool,X98))),X102,X101) | s(fun(X98,bool),X99) = s(fun(X98,bool),X100) | s(bool,'T_0') = s(bool,f(s(X98,sK17(X99,X100,bool,X98))))) )),
  inference(subsumption_resolution,[],[f18836,f275])).
fof(f18836,plain,(
  ( ! [X101,X99,X102,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | sP11(g(s(X98,sK17(X99,X100,bool,X98))),X102,X101) | s(fun(X98,bool),X99) = s(fun(X98,bool),X100) | s(bool,'T_0') = s(bool,f(s(X98,sK17(X99,X100,bool,X98))))) )),
  inference(superposition,[],[f477,f2341])).
fof(f18937,plain,(
  ( ! [X94,X97,X95,X93,X96] : (p(s(bool,X96)) | sP11(X97,X96,g(s(X93,sK17(X94,X95,bool,X93)))) | p(s(bool,X97)) | s(fun(X93,bool),X94) = s(fun(X93,bool),X95) | s(bool,'T_0') = s(bool,f(s(X93,sK17(X94,X95,bool,X93))))) )),
  inference(subsumption_resolution,[],[f18835,f275])).
fof(f18835,plain,(
  ( ! [X94,X97,X95,X93,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | sP11(X97,X96,g(s(X93,sK17(X94,X95,bool,X93)))) | p(s(bool,X97)) | s(fun(X93,bool),X94) = s(fun(X93,bool),X95) | s(bool,'T_0') = s(bool,f(s(X93,sK17(X94,X95,bool,X93))))) )),
  inference(superposition,[],[f476,f2341])).
fof(f18936,plain,(
  ( ! [X92,X90,X88,X91,X89] : (p(s(bool,X91)) | ~p(s(bool,X92)) | ~sP10(X92,X91,g(s(X88,sK17(X89,X90,bool,X88)))) | s(fun(X88,bool),X89) = s(fun(X88,bool),X90) | s(bool,'T_0') = s(bool,f(s(X88,sK17(X89,X90,bool,X88))))) )),
  inference(subsumption_resolution,[],[f18834,f275])).
fof(f18834,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | ~p(s(bool,X92)) | ~sP10(X92,X91,g(s(X88,sK17(X89,X90,bool,X88)))) | s(fun(X88,bool),X89) = s(fun(X88,bool),X90) | s(bool,'T_0') = s(bool,f(s(X88,sK17(X89,X90,bool,X88))))) )),
  inference(superposition,[],[f472,f2341])).
fof(f18935,plain,(
  ( ! [X87,X85,X83,X86,X84] : (p(s(bool,X86)) | p(s(bool,X87)) | ~sP9(g(s(X83,sK17(X84,X85,bool,X83))),X87,X86) | s(fun(X83,bool),X84) = s(fun(X83,bool),X85) | s(bool,'T_0') = s(bool,f(s(X83,sK17(X84,X85,bool,X83))))) )),
  inference(subsumption_resolution,[],[f18833,f275])).
fof(f18833,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X86)) | p(s(bool,X87)) | ~sP9(g(s(X83,sK17(X84,X85,bool,X83))),X87,X86) | s(fun(X83,bool),X84) = s(fun(X83,bool),X85) | s(bool,'T_0') = s(bool,f(s(X83,sK17(X84,X85,bool,X83))))) )),
  inference(superposition,[],[f462,f2341])).
fof(f18934,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,X81)) | sP9(g(s(X78,sK17(X79,X80,bool,X78))),X81,X82) | s(fun(X78,bool),X79) = s(fun(X78,bool),X80) | s(bool,'T_0') = s(bool,f(s(X78,sK17(X79,X80,bool,X78))))) )),
  inference(subsumption_resolution,[],[f18832,f275])).
fof(f18832,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | sP9(g(s(X78,sK17(X79,X80,bool,X78))),X81,X82) | s(fun(X78,bool),X79) = s(fun(X78,bool),X80) | s(bool,'T_0') = s(bool,f(s(X78,sK17(X79,X80,bool,X78))))) )),
  inference(superposition,[],[f461,f2341])).
fof(f18933,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,X76)) | sP9(g(s(X73,sK17(X74,X75,bool,X73))),X77,X76) | s(fun(X73,bool),X74) = s(fun(X73,bool),X75) | s(bool,'T_0') = s(bool,f(s(X73,sK17(X74,X75,bool,X73))))) )),
  inference(subsumption_resolution,[],[f18831,f275])).
fof(f18831,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP9(g(s(X73,sK17(X74,X75,bool,X73))),X77,X76) | s(fun(X73,bool),X74) = s(fun(X73,bool),X75) | s(bool,'T_0') = s(bool,f(s(X73,sK17(X74,X75,bool,X73))))) )),
  inference(superposition,[],[f460,f2341])).
fof(f18932,plain,(
  ( ! [X70,X68,X72,X71,X69] : (p(s(bool,X71)) | p(s(bool,X72)) | ~sP8(g(s(X68,sK17(X69,X70,bool,X68))),X71,X72) | s(fun(X68,bool),X69) = s(fun(X68,bool),X70) | s(bool,'T_0') = s(bool,f(s(X68,sK17(X69,X70,bool,X68))))) )),
  inference(subsumption_resolution,[],[f18830,f275])).
fof(f18830,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP8(g(s(X68,sK17(X69,X70,bool,X68))),X71,X72) | s(fun(X68,bool),X69) = s(fun(X68,bool),X70) | s(bool,'T_0') = s(bool,f(s(X68,sK17(X69,X70,bool,X68))))) )),
  inference(superposition,[],[f455,f2341])).
fof(f18931,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X67,X66,g(s(X63,sK17(X64,X65,bool,X63)))) | s(fun(X63,bool),X64) = s(fun(X63,bool),X65) | s(bool,'T_0') = s(bool,f(s(X63,sK17(X64,X65,bool,X63))))) )),
  inference(subsumption_resolution,[],[f18829,f275])).
fof(f18829,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X67,X66,g(s(X63,sK17(X64,X65,bool,X63)))) | s(fun(X63,bool),X64) = s(fun(X63,bool),X65) | s(bool,'T_0') = s(bool,f(s(X63,sK17(X64,X65,bool,X63))))) )),
  inference(superposition,[],[f447,f2341])).
fof(f18930,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | ~p(s(bool,X62)) | sP7(g(s(X58,sK17(X59,X60,bool,X58))),X61,X62) | s(fun(X58,bool),X59) = s(fun(X58,bool),X60) | s(bool,'T_0') = s(bool,f(s(X58,sK17(X59,X60,bool,X58))))) )),
  inference(subsumption_resolution,[],[f18828,f275])).
fof(f18828,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | sP7(g(s(X58,sK17(X59,X60,bool,X58))),X61,X62) | s(fun(X58,bool),X59) = s(fun(X58,bool),X60) | s(bool,'T_0') = s(bool,f(s(X58,sK17(X59,X60,bool,X58))))) )),
  inference(superposition,[],[f444,f2341])).
fof(f18929,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP6(X56,g(s(X53,sK17(X54,X55,bool,X53))),X57) | s(fun(X53,bool),X54) = s(fun(X53,bool),X55) | s(bool,'T_0') = s(bool,f(s(X53,sK17(X54,X55,bool,X53))))) )),
  inference(subsumption_resolution,[],[f18827,f275])).
fof(f18827,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP6(X56,g(s(X53,sK17(X54,X55,bool,X53))),X57) | s(fun(X53,bool),X54) = s(fun(X53,bool),X55) | s(bool,'T_0') = s(bool,f(s(X53,sK17(X54,X55,bool,X53))))) )),
  inference(superposition,[],[f438,f2341])).
fof(f18928,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP3(X51,g(s(X48,sK17(X49,X50,bool,X48))),X52) | s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,f(s(X48,sK17(X49,X50,bool,X48))))) )),
  inference(subsumption_resolution,[],[f18826,f275])).
fof(f18826,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP3(X51,g(s(X48,sK17(X49,X50,bool,X48))),X52) | s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,f(s(X48,sK17(X49,X50,bool,X48))))) )),
  inference(superposition,[],[f419,f2341])).
fof(f18927,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP4(X46,g(s(X43,sK17(X44,X45,bool,X43))),X47) | s(fun(X43,bool),X44) = s(fun(X43,bool),X45) | s(bool,'T_0') = s(bool,f(s(X43,sK17(X44,X45,bool,X43))))) )),
  inference(subsumption_resolution,[],[f18825,f275])).
fof(f18825,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP4(X46,g(s(X43,sK17(X44,X45,bool,X43))),X47) | s(fun(X43,bool),X44) = s(fun(X43,bool),X45) | s(bool,'T_0') = s(bool,f(s(X43,sK17(X44,X45,bool,X43))))) )),
  inference(superposition,[],[f415,f2341])).
fof(f18926,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X42)) | sP1(X42,g(s(X39,sK17(X40,X41,bool,X39)))) | s(fun(X39,bool),X40) = s(fun(X39,bool),X41) | s(bool,'T_0') = s(bool,f(s(X39,sK17(X40,X41,bool,X39))))) )),
  inference(subsumption_resolution,[],[f18824,f275])).
fof(f18824,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | sP1(X42,g(s(X39,sK17(X40,X41,bool,X39)))) | s(fun(X39,bool),X40) = s(fun(X39,bool),X41) | s(bool,'T_0') = s(bool,f(s(X39,sK17(X40,X41,bool,X39))))) )),
  inference(superposition,[],[f385,f2341])).
fof(f18925,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X38)) | sP1(g(s(X35,sK17(X36,X37,bool,X35))),X38) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'T_0') = s(bool,f(s(X35,sK17(X36,X37,bool,X35))))) )),
  inference(subsumption_resolution,[],[f18823,f275])).
fof(f18823,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | sP1(g(s(X35,sK17(X36,X37,bool,X35))),X38) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'T_0') = s(bool,f(s(X35,sK17(X36,X37,bool,X35))))) )),
  inference(superposition,[],[f382,f2341])).
fof(f18924,plain,(
  ( ! [X33,X31,X34,X32] : (sP0(g(s(X31,sK17(X32,X33,bool,X31))),X34) | s(fun(X31,bool),X32) = s(fun(X31,bool),X33) | s(bool,'T_0') = s(bool,f(s(X31,sK17(X32,X33,bool,X31))))) )),
  inference(subsumption_resolution,[],[f18822,f275])).
fof(f18822,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP0(g(s(X31,sK17(X32,X33,bool,X31))),X34) | s(fun(X31,bool),X32) = s(fun(X31,bool),X33) | s(bool,'T_0') = s(bool,f(s(X31,sK17(X32,X33,bool,X31))))) )),
  inference(superposition,[],[f367,f2341])).
fof(f18923,plain,(
  ( ! [X30,X28,X29,X27] : (sP0(X30,g(s(X27,sK17(X28,X29,bool,X27)))) | s(fun(X27,bool),X28) = s(fun(X27,bool),X29) | s(bool,'T_0') = s(bool,f(s(X27,sK17(X28,X29,bool,X27))))) )),
  inference(subsumption_resolution,[],[f18821,f275])).
fof(f18821,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | sP0(X30,g(s(X27,sK17(X28,X29,bool,X27)))) | s(fun(X27,bool),X28) = s(fun(X27,bool),X29) | s(bool,'T_0') = s(bool,f(s(X27,sK17(X28,X29,bool,X27))))) )),
  inference(superposition,[],[f366,f2341])).
fof(f18922,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X26)) | s(bool,X26) = s(bool,g(s(X23,sK17(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'T_0') = s(bool,f(s(X23,sK17(X24,X25,bool,X23))))) )),
  inference(subsumption_resolution,[],[f18820,f275])).
fof(f18820,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,g(s(X23,sK17(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'T_0') = s(bool,f(s(X23,sK17(X24,X25,bool,X23))))) )),
  inference(superposition,[],[f360,f2341])).
fof(f18921,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,g(s(X9,sK17(X10,X11,bool,X9)))) | s(fun(X9,bool),X10) = s(fun(X9,bool),X11) | s(bool,'T_0') = s(bool,f(s(X9,sK17(X10,X11,bool,X9))))) )),
  inference(subsumption_resolution,[],[f18816,f275])).
fof(f18816,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,g(s(X9,sK17(X10,X11,bool,X9)))) | s(fun(X9,bool),X10) = s(fun(X9,bool),X11) | s(bool,'T_0') = s(bool,f(s(X9,sK17(X10,X11,bool,X9))))) )),
  inference(superposition,[],[f298,f2341])).
fof(f18889,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,g(s(X3,sK17(X4,X5,bool,X3))))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,f(s(X3,sK17(X4,X5,bool,X3))))) )),
  inference(trivial_inequality_removal,[],[f18814])).
fof(f18814,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,g(s(X3,sK17(X4,X5,bool,X3))))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,f(s(X3,sK17(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f294,f2341])).
fof(f18903,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,f(s(X0,sK17(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK17(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f18902])).
fof(f18902,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s(bool,'T_0') = s(bool,X35) | s(fun(X32,bool),X33) = s(fun(X32,bool),X34) | s(bool,'T_0') = s(bool,f(s(X32,sK17(X33,X34,bool,X32)))) | s(bool,X36) = s(bool,g(s(X32,sK17(X33,X34,bool,X32)))) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(duplicate_literal_removal,[],[f18805])).
fof(f18805,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s(bool,'T_0') = s(bool,X35) | s(fun(X32,bool),X33) = s(fun(X32,bool),X34) | s(bool,'T_0') = s(bool,f(s(X32,sK17(X33,X34,bool,X32)))) | s(bool,X36) = s(bool,g(s(X32,sK17(X33,X34,bool,X32)))) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f2341,f1281])).
fof(f18758,plain,(
  ( ! [X62,X63] : (~sP5('F_0',f(s(bool,happ(s(fun(X62,bool),X63),s(X62,sK15(X63,X62,bool))))),'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),X63),s(X62,sK16(X63,X62,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X62,bool),X63)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13490,f1067])).
fof(f18675,plain,(
  ( ! [X198,X199,X197] : (s(bool,X199) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X197),s('type.rat.rat',X198))))))) | s(bool,'F_0') = s(bool,X199) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X198),s('type.rat.rat',X197)))) ) | $spl74),
  inference(subsumption_resolution,[],[f18615,f275])).
fof(f18615,plain,(
  ( ! [X198,X199,X197] : (~p(s(bool,'T_0')) | s(bool,X199) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X197),s('type.rat.rat',X198))))))) | s(bool,'F_0') = s(bool,X199) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X198),s('type.rat.rat',X197)))) ) | $spl74),
  inference(superposition,[],[f9608,f2119])).
fof(f18670,plain,(
  ( ! [X180,X181,X179] : (s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X181))))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X179),s('type.rat.rat',X180))))))) | s('type.rat.rat',X181) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X180),s('type.rat.rat',X179)))) )),
  inference(subsumption_resolution,[],[f18609,f275])).
fof(f18609,plain,(
  ( ! [X180,X181,X179] : (~p(s(bool,'T_0')) | s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X181))))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X179),s('type.rat.rat',X180))))))) | s('type.rat.rat',X181) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X180),s('type.rat.rat',X179)))) )),
  inference(superposition,[],[f2173,f2119])).
fof(f18669,plain,(
  ( ! [X107,X105,X108,X106] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X105),s('type.rat.rat',X106)))))),X107,X108) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X105),s('type.rat.rat',X106)))))),X107,X108) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X106),s('type.rat.rat',X105)))) )),
  inference(subsumption_resolution,[],[f18592,f275])).
fof(f18592,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,'T_0')) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X105),s('type.rat.rat',X106)))))),X107,X108) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X105),s('type.rat.rat',X106)))))),X107,X108) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X106),s('type.rat.rat',X105)))) )),
  inference(superposition,[],[f483,f2119])).
fof(f18668,plain,(
  ( ! [X103,X101,X102,X104] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X101),s('type.rat.rat',X102)))))),X103,X104) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X101),s('type.rat.rat',X102)))))),X103,X104) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X102),s('type.rat.rat',X101)))) )),
  inference(subsumption_resolution,[],[f18591,f275])).
fof(f18591,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'T_0')) | ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X101),s('type.rat.rat',X102)))))),X103,X104) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X101),s('type.rat.rat',X102)))))),X103,X104) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X102),s('type.rat.rat',X101)))) )),
  inference(superposition,[],[f482,f2119])).
fof(f18667,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X32)) | s(bool,X32) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X31))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X30)))) )),
  inference(subsumption_resolution,[],[f18572,f275])).
fof(f18572,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | s(bool,X32) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X31))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X30)))) )),
  inference(superposition,[],[f360,f2119])).
fof(f18666,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X9))))),s('type.rat.rat',X10))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X9))))),s('type.rat.rat',X11))) | s('type.rat.rat',X10) = s('type.rat.rat',X11) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X9),s('type.rat.rat',X8)))) )),
  inference(subsumption_resolution,[],[f18563,f275])).
fof(f18563,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'T_0')) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X9))))),s('type.rat.rat',X10))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X9))))),s('type.rat.rat',X11))) | s('type.rat.rat',X10) = s('type.rat.rat',X11) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X9),s('type.rat.rat',X8)))) )),
  inference(superposition,[],[f398,f2119])).
fof(f18543,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK15(X11,X10,'type.rat.rat')))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK16(X11,X10,'type.rat.rat')))),s('type.rat.rat',X12))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X10,'type.rat.rat'),X11)))) )),
  inference(superposition,[],[f2119,f1067])).
fof(f18540,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2119,f355])).
fof(f18538,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK15(X11,X10,'type.rat.rat')))),s('type.rat.rat',X12))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK16(X11,X10,'type.rat.rat')))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X10,'type.rat.rat'),X11)))) )),
  inference(superposition,[],[f2119,f1067])).
fof(f18535,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))),s('type.rat.rat',X3))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))))) )),
  inference(superposition,[],[f2119,f355])).
fof(f18379,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1137,f408])).
fof(f18339,plain,(
  ( ! [X14,X15,X16] : (sP3('T_0',f(s(bool,X14)),'F_0') | s(bool,X14) = s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl47),
  inference(subsumption_resolution,[],[f18326,f522])).
fof(f18326,plain,(
  ( ! [X14,X15,X16] : (sP3('T_0',f(s(bool,X14)),'F_0') | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl47),
  inference(superposition,[],[f14056,f1281])).
fof(f18317,plain,(
  ( ! [X0] : (sP3('T_0',f(s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f14056,f283])).
fof(f18293,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK15(X1,bool,X0)) | s(bool,X2) = s(bool,sK16(X1,bool,X0)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),X1)))) )),
  inference(subsumption_resolution,[],[f18275,f498])).
fof(f18275,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK15(X1,bool,X0)) | s(bool,X2) = s(bool,sK16(X1,bool,X0)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),X1)))) )),
  inference(superposition,[],[f1449,f283])).
fof(f18274,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK15(X2,X1,fun(bool,X0)))))))) | s(bool,'T_0') != s(bool,sK15(happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK16(X2,X1,fun(bool,X0)))),bool,X0)) | s(bool,X3) = s(bool,sK16(happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK16(X2,X1,fun(bool,X0)))),bool,X0)) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,fun(bool,X0)),X2)))) )),
  inference(superposition,[],[f1449,f1067])).
fof(f18269,plain,(
  ( ! [X14,X15,X16] : (~sP4('T_0','T_0',f(s(bool,X14))) | s(bool,X14) = s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl47),
  inference(subsumption_resolution,[],[f18258,f522])).
fof(f18258,plain,(
  ( ! [X14,X15,X16] : (~sP4('T_0','T_0',f(s(bool,X14))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl47),
  inference(superposition,[],[f14024,f1281])).
fof(f18249,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f14024,f283])).
fof(f18240,plain,(
  ( ! [X14,X15,X16] : (sP3(f(s(bool,X14)),'T_0','F_0') | s(bool,X14) = s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl47),
  inference(subsumption_resolution,[],[f18223,f522])).
fof(f18223,plain,(
  ( ! [X14,X15,X16] : (sP3(f(s(bool,X14)),'T_0','F_0') | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl47),
  inference(superposition,[],[f14011,f1281])).
fof(f18214,plain,(
  ( ! [X0] : (sP3(f(s(bool,X0)),'T_0','F_0') | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f14011,f283])).
fof(f18189,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,sK15(X1,bool,X0)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,sK16(X1,bool,X0)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),X1)))) )),
  inference(subsumption_resolution,[],[f18171,f498])).
fof(f18171,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,X2) != s(bool,sK15(X1,bool,X0)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,sK16(X1,bool,X0)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),X1)))) )),
  inference(superposition,[],[f1339,f283])).
fof(f18170,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK15(X2,X1,fun(bool,X0)))))))) | s(bool,X3) != s(bool,sK15(happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK16(X2,X1,fun(bool,X0)))),bool,X0)) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,sK16(happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK16(X2,X1,fun(bool,X0)))),bool,X0)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,fun(bool,X0)),X2)))) )),
  inference(superposition,[],[f1339,f1067])).
fof(f16014,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2)) )),
  inference(equality_factoring,[],[f1281])).
fof(f18090,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f16007,f608])).
fof(f16007,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(equality_factoring,[],[f1281])).
fof(f15481,plain,(
  ( ! [X580,X576,X578,X581,X577,X579] : (s(X578,f(s(bool,'T_0'))) != s(X578,g(s(bool,'T_0'))) | s(fun(bool,X578),X576) = s(fun(bool,X578),X577) | s(bool,X579) = s(bool,X580) | s(bool,X580) = s(bool,sK17(X576,X577,X578,bool)) | s(bool,X579) = s(bool,X581) | s(bool,'T_0') = s(bool,X581)) )),
  inference(superposition,[],[f401,f1281])).
fof(f15480,plain,(
  ( ! [X573,X575,X571,X572,X574] : (s(bool,'T_0') != s(bool,sK15(X571,bool,X572)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X572),X571)))) | s(bool,X573) = s(bool,X574) | s(bool,X574) = s(bool,sK16(X571,bool,X572)) | s(bool,X573) = s(bool,X575) | s(bool,'T_0') = s(bool,X575)) )),
  inference(superposition,[],[f400,f1281])).
fof(f15479,plain,(
  ( ! [X569,X567,X568,X570,X566] : (s(X567,happ(s(fun(bool,X567),X566),s(bool,'T_0'))) = s(X567,happ(s(fun(bool,X567),X566),s(bool,sK15(X566,bool,X567)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X567),X566))) | s(bool,X568) = s(bool,X569) | s(bool,X569) = s(bool,sK16(X566,bool,X567)) | s(bool,X568) = s(bool,X570) | s(bool,'T_0') = s(bool,X570)) )),
  inference(superposition,[],[f1067,f1281])).
fof(f15476,plain,(
  ( ! [X553,X549,X551,X552,X548,X550] : (s(bool,'T_0') != s(bool,f(s(X548,sK17(X549,X550,bool,X548)))) | s(fun(X548,bool),X549) = s(fun(X548,bool),X550) | s(bool,X551) = s(bool,X552) | s(bool,X552) = s(bool,g(s(X548,sK17(X549,X550,bool,X548)))) | s(bool,X551) = s(bool,X553) | s(bool,'T_0') = s(bool,X553)) )),
  inference(superposition,[],[f401,f1281])).
fof(f17918,plain,(
  ( ! [X528,X524,X526,X522,X525,X527,X521,X523] : (s(X522,happ(s(fun(X521,X522),X523),s(X521,X524))) != s(X522,happ(s(fun(X521,X522),X523),s(X521,X525))) | s(X521,X524) = s(X521,X525) | s(bool,X526) = s(bool,X527) | s(bool,X527) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X521,X522),X523))) | s(bool,X526) = s(bool,X528) | s(bool,'T_0') = s(bool,X528)) )),
  inference(subsumption_resolution,[],[f15472,f275])).
fof(f15472,plain,(
  ( ! [X528,X524,X526,X522,X525,X527,X521,X523] : (~p(s(bool,'T_0')) | s(X522,happ(s(fun(X521,X522),X523),s(X521,X524))) != s(X522,happ(s(fun(X521,X522),X523),s(X521,X525))) | s(X521,X524) = s(X521,X525) | s(bool,X526) = s(bool,X527) | s(bool,X527) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X521,X522),X523))) | s(bool,X526) = s(bool,X528) | s(bool,'T_0') = s(bool,X528)) )),
  inference(superposition,[],[f398,f1281])).
fof(f17914,plain,(
  ( ! [X237,X235,X238,X236,X234] : (s(bool,X234) = s(bool,X235) | s(bool,'F_0') = s(bool,X235) | s(bool,X236) = s(bool,X237) | s(bool,X234) = s(bool,X237) | s(bool,X236) = s(bool,X238) | s(bool,'T_0') = s(bool,X238)) ) | $spl74),
  inference(subsumption_resolution,[],[f15392,f275])).
fof(f15392,plain,(
  ( ! [X237,X235,X238,X236,X234] : (~p(s(bool,'T_0')) | s(bool,X234) = s(bool,X235) | s(bool,'F_0') = s(bool,X235) | s(bool,X236) = s(bool,X237) | s(bool,X234) = s(bool,X237) | s(bool,X236) = s(bool,X238) | s(bool,'T_0') = s(bool,X238)) ) | $spl74),
  inference(superposition,[],[f9608,f1281])).
fof(f17908,plain,(
  ( ! [X208,X206,X209,X207,X205] : (s(bool,X205) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X206))))) | s('type.rat.rat',X206) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X207) = s(bool,X208) | s(bool,X205) = s(bool,X208) | s(bool,X207) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) )),
  inference(subsumption_resolution,[],[f15386,f275])).
fof(f15386,plain,(
  ( ! [X208,X206,X209,X207,X205] : (~p(s(bool,'T_0')) | s(bool,X205) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X206))))) | s('type.rat.rat',X206) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X207) = s(bool,X208) | s(bool,X205) = s(bool,X208) | s(bool,X207) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) )),
  inference(superposition,[],[f2173,f1281])).
fof(f17907,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (~sP8(X145,X146,X147) | sP9(X145,X146,X147) | s(bool,X148) = s(bool,X149) | s(bool,X145) = s(bool,X149) | s(bool,X148) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(subsumption_resolution,[],[f15375,f275])).
fof(f15375,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (~p(s(bool,'T_0')) | ~sP8(X145,X146,X147) | sP9(X145,X146,X147) | s(bool,X148) = s(bool,X149) | s(bool,X145) = s(bool,X149) | s(bool,X148) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f483,f1281])).
fof(f17906,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (~sP10(X139,X140,X141) | sP11(X139,X140,X141) | s(bool,X142) = s(bool,X143) | s(bool,X139) = s(bool,X143) | s(bool,X142) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(subsumption_resolution,[],[f15374,f275])).
fof(f15374,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | ~sP10(X139,X140,X141) | sP11(X139,X140,X141) | s(bool,X142) = s(bool,X143) | s(bool,X139) = s(bool,X143) | s(bool,X142) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f482,f1281])).
fof(f17905,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~p(s(bool,X134)) | p(s(bool,X135)) | ~sP11(X133,X135,X134) | s(bool,X136) = s(bool,X137) | s(bool,X133) = s(bool,X137) | s(bool,X136) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(subsumption_resolution,[],[f15373,f275])).
fof(f15373,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~p(s(bool,'T_0')) | ~p(s(bool,X134)) | p(s(bool,X135)) | ~sP11(X133,X135,X134) | s(bool,X136) = s(bool,X137) | s(bool,X133) = s(bool,X137) | s(bool,X136) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f479,f1281])).
fof(f17904,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (~p(s(bool,X128)) | sP11(X127,X128,X129) | s(bool,X130) = s(bool,X131) | s(bool,X127) = s(bool,X131) | s(bool,X130) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(subsumption_resolution,[],[f15372,f275])).
fof(f15372,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X128)) | sP11(X127,X128,X129) | s(bool,X130) = s(bool,X131) | s(bool,X127) = s(bool,X131) | s(bool,X130) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f478,f1281])).
fof(f17903,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (p(s(bool,X122)) | sP11(X121,X123,X122) | s(bool,X124) = s(bool,X125) | s(bool,X121) = s(bool,X125) | s(bool,X124) = s(bool,X126) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f15371,f275])).
fof(f15371,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,X122)) | sP11(X121,X123,X122) | s(bool,X124) = s(bool,X125) | s(bool,X121) = s(bool,X125) | s(bool,X124) = s(bool,X126) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f477,f1281])).
fof(f17902,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (p(s(bool,X116)) | sP11(X117,X116,X115) | p(s(bool,X117)) | s(bool,X118) = s(bool,X119) | s(bool,X115) = s(bool,X119) | s(bool,X118) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f15370,f275])).
fof(f15370,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (~p(s(bool,'T_0')) | p(s(bool,X116)) | sP11(X117,X116,X115) | p(s(bool,X117)) | s(bool,X118) = s(bool,X119) | s(bool,X115) = s(bool,X119) | s(bool,X118) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f476,f1281])).
fof(f17901,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (p(s(bool,X110)) | ~p(s(bool,X111)) | ~sP10(X111,X110,X109) | s(bool,X112) = s(bool,X113) | s(bool,X109) = s(bool,X113) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f15369,f275])).
fof(f15369,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (~p(s(bool,'T_0')) | p(s(bool,X110)) | ~p(s(bool,X111)) | ~sP10(X111,X110,X109) | s(bool,X112) = s(bool,X113) | s(bool,X109) = s(bool,X113) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f472,f1281])).
fof(f17900,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (p(s(bool,X104)) | p(s(bool,X105)) | ~sP9(X103,X105,X104) | s(bool,X106) = s(bool,X107) | s(bool,X103) = s(bool,X107) | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f15368,f275])).
fof(f15368,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | p(s(bool,X105)) | ~sP9(X103,X105,X104) | s(bool,X106) = s(bool,X107) | s(bool,X103) = s(bool,X107) | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f462,f1281])).
fof(f17899,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,X98)) | sP9(X97,X98,X99) | s(bool,X100) = s(bool,X101) | s(bool,X97) = s(bool,X101) | s(bool,X100) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f15367,f275])).
fof(f15367,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X98)) | sP9(X97,X98,X99) | s(bool,X100) = s(bool,X101) | s(bool,X97) = s(bool,X101) | s(bool,X100) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f461,f1281])).
fof(f17898,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,X92)) | sP9(X91,X93,X92) | s(bool,X94) = s(bool,X95) | s(bool,X91) = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f15366,f275])).
fof(f15366,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | sP9(X91,X93,X92) | s(bool,X94) = s(bool,X95) | s(bool,X91) = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f460,f1281])).
fof(f17897,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | ~sP8(X85,X86,X87) | s(bool,X88) = s(bool,X89) | s(bool,X85) = s(bool,X89) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f15365,f275])).
fof(f15365,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,'T_0')) | p(s(bool,X86)) | p(s(bool,X87)) | ~sP8(X85,X86,X87) | s(bool,X88) = s(bool,X89) | s(bool,X85) = s(bool,X89) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f455,f1281])).
fof(f17896,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP7(X81,X80,X79) | s(bool,X82) = s(bool,X83) | s(bool,X79) = s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f15364,f275])).
fof(f15364,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP7(X81,X80,X79) | s(bool,X82) = s(bool,X83) | s(bool,X79) = s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f447,f1281])).
fof(f17895,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (~p(s(bool,X74)) | ~p(s(bool,X75)) | sP7(X73,X74,X75) | s(bool,X76) = s(bool,X77) | s(bool,X73) = s(bool,X77) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f15363,f275])).
fof(f15363,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X74)) | ~p(s(bool,X75)) | sP7(X73,X74,X75) | s(bool,X76) = s(bool,X77) | s(bool,X73) = s(bool,X77) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f444,f1281])).
fof(f17894,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,X68)) | p(s(bool,X69)) | ~sP6(X68,X67,X69) | s(bool,X70) = s(bool,X71) | s(bool,X67) = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f15362,f275])).
fof(f15362,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X68)) | p(s(bool,X69)) | ~sP6(X68,X67,X69) | s(bool,X70) = s(bool,X71) | s(bool,X67) = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f438,f1281])).
fof(f17893,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | ~sP3(X62,X61,X63) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f15361,f275])).
fof(f15361,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP3(X62,X61,X63) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f419,f1281])).
fof(f17892,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP4(X56,X55,X57) | s(bool,X58) = s(bool,X59) | s(bool,X55) = s(bool,X59) | s(bool,X58) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f15360,f275])).
fof(f15360,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP4(X56,X55,X57) | s(bool,X58) = s(bool,X59) | s(bool,X55) = s(bool,X59) | s(bool,X58) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f415,f1281])).
fof(f17891,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X51)) | sP1(X51,X50) | s(bool,X52) = s(bool,X53) | s(bool,X50) = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f15359,f275])).
fof(f15359,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X51)) | sP1(X51,X50) | s(bool,X52) = s(bool,X53) | s(bool,X50) = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f385,f1281])).
fof(f17890,plain,(
  ( ! [X47,X45,X48,X46,X49] : (p(s(bool,X46)) | sP1(X45,X46) | s(bool,X47) = s(bool,X48) | s(bool,X45) = s(bool,X48) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f15358,f275])).
fof(f15358,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | sP1(X45,X46) | s(bool,X47) = s(bool,X48) | s(bool,X45) = s(bool,X48) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f382,f1281])).
fof(f17889,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP0(X40,X41) | s(bool,X42) = s(bool,X43) | s(bool,X40) = s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f15357,f275])).
fof(f15357,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'T_0')) | sP0(X40,X41) | s(bool,X42) = s(bool,X43) | s(bool,X40) = s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f367,f1281])).
fof(f17888,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP0(X36,X35) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f15356,f275])).
fof(f15356,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP0(X36,X35) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f366,f1281])).
fof(f17887,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | s(bool,X30) = s(bool,X31) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f15355,f275])).
fof(f15355,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | s(bool,X30) = s(bool,X31) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f360,f1281])).
fof(f17886,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X12) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f15351,f275])).
fof(f15351,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X12) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f298,f1281])).
fof(f15273,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f294,f1281])).
fof(f15269,plain,(
  ( ! [X573,X569,X571,X572,X570] : (s(bool,X571) != s(bool,sK15(X569,bool,X570)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X570),X569)))) | s(bool,X571) = s(bool,X572) | s(bool,'T_0') = s(bool,sK16(X569,bool,X570)) | s(bool,X572) = s(bool,X573) | s(bool,'T_0') = s(bool,X573)) )),
  inference(superposition,[],[f400,f1281])).
fof(f15268,plain,(
  ( ! [X565,X567,X568,X564,X566] : (s(X565,happ(s(fun(bool,X565),X564),s(bool,X566))) = s(X565,happ(s(fun(bool,X565),X564),s(bool,sK15(X564,bool,X565)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X565),X564))) | s(bool,X566) = s(bool,X567) | s(bool,'T_0') = s(bool,sK16(X564,bool,X565)) | s(bool,X567) = s(bool,X568) | s(bool,'T_0') = s(bool,X568)) )),
  inference(superposition,[],[f1067,f1281])).
fof(f15265,plain,(
  ( ! [X546,X549,X545,X547,X548,X550] : (s(bool,X548) != s(bool,f(s(X545,sK17(X546,X547,bool,X545)))) | s(fun(X545,bool),X546) = s(fun(X545,bool),X547) | s(bool,X548) = s(bool,X549) | s(bool,'T_0') = s(bool,g(s(X545,sK17(X546,X547,bool,X545)))) | s(bool,X549) = s(bool,X550) | s(bool,'T_0') = s(bool,X550)) )),
  inference(superposition,[],[f401,f1281])).
fof(f17835,plain,(
  ( ! [X451,X452,X450] : (sP4('T_0',f(s(bool,X450)),'F_0') | s(bool,X450) = s(bool,X451) | s(bool,X451) = s(bool,X452) | s(bool,'T_0') = s(bool,X452)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15243,f522])).
fof(f15243,plain,(
  ( ! [X451,X452,X450] : (sP4('T_0',f(s(bool,X450)),'F_0') | s(bool,X450) = s(bool,X451) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X451) = s(bool,X452) | s(bool,'T_0') = s(bool,X452)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13653,f1281])).
fof(f17834,plain,(
  ( ! [X447,X449,X448] : (~sP2(f(s(bool,X447)),'F_0','F_0') | s(bool,X447) = s(bool,X448) | s(bool,X448) = s(bool,X449) | s(bool,'T_0') = s(bool,X449)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15242,f522])).
fof(f15242,plain,(
  ( ! [X447,X449,X448] : (~sP2(f(s(bool,X447)),'F_0','F_0') | s(bool,X447) = s(bool,X448) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X448) = s(bool,X449) | s(bool,'T_0') = s(bool,X449)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13622,f1281])).
fof(f17833,plain,(
  ( ! [X445,X444,X446] : (~sP2('F_0',f(s(bool,X444)),'F_0') | s(bool,X444) = s(bool,X445) | s(bool,X445) = s(bool,X446) | s(bool,'T_0') = s(bool,X446)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15241,f522])).
fof(f15241,plain,(
  ( ! [X445,X444,X446] : (~sP2('F_0',f(s(bool,X444)),'F_0') | s(bool,X444) = s(bool,X445) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X445) = s(bool,X446) | s(bool,'T_0') = s(bool,X446)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13487,f1281])).
fof(f17832,plain,(
  ( ! [X441,X443,X440,X442] : (sP1(f(s(bool,X440)),X441) | s(bool,'F_0') = s(bool,X441) | s(bool,X440) = s(bool,X442) | s(bool,X442) = s(bool,X443) | s(bool,'T_0') = s(bool,X443)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15240,f522])).
fof(f15240,plain,(
  ( ! [X441,X443,X440,X442] : (sP1(f(s(bool,X440)),X441) | s(bool,'F_0') = s(bool,X441) | s(bool,X440) = s(bool,X442) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X442) = s(bool,X443) | s(bool,'T_0') = s(bool,X443)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13486,f1281])).
fof(f17831,plain,(
  ( ! [X436,X438,X437,X439] : (sP1(X437,f(s(bool,X436))) | s(bool,'F_0') = s(bool,X437) | s(bool,X436) = s(bool,X438) | s(bool,X438) = s(bool,X439) | s(bool,'T_0') = s(bool,X439)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15239,f522])).
fof(f15239,plain,(
  ( ! [X436,X438,X437,X439] : (sP1(X437,f(s(bool,X436))) | s(bool,'F_0') = s(bool,X437) | s(bool,X436) = s(bool,X438) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X438) = s(bool,X439) | s(bool,'T_0') = s(bool,X439)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13485,f1281])).
fof(f17830,plain,(
  ( ! [X434,X433,X435] : (sP5('T_0','F_0',f(s(bool,X433))) | s(bool,X433) = s(bool,X434) | s(bool,X434) = s(bool,X435) | s(bool,'T_0') = s(bool,X435)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15238,f522])).
fof(f15238,plain,(
  ( ! [X434,X433,X435] : (sP5('T_0','F_0',f(s(bool,X433))) | s(bool,X433) = s(bool,X434) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X434) = s(bool,X435) | s(bool,'T_0') = s(bool,X435)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13403,f1281])).
fof(f17829,plain,(
  ( ! [X430,X432,X431] : (~sP5('F_0','F_0',f(s(bool,X430))) | s(bool,X430) = s(bool,X431) | s(bool,X431) = s(bool,X432) | s(bool,'T_0') = s(bool,X432)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15237,f522])).
fof(f15237,plain,(
  ( ! [X430,X432,X431] : (~sP5('F_0','F_0',f(s(bool,X430))) | s(bool,X430) = s(bool,X431) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X431) = s(bool,X432) | s(bool,'T_0') = s(bool,X432)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13402,f1281])).
fof(f17828,plain,(
  ( ! [X428,X429,X427] : (sP5('T_0',f(s(bool,X427)),'F_0') | s(bool,X427) = s(bool,X428) | s(bool,X428) = s(bool,X429) | s(bool,'T_0') = s(bool,X429)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15236,f522])).
fof(f15236,plain,(
  ( ! [X428,X429,X427] : (sP5('T_0',f(s(bool,X427)),'F_0') | s(bool,X427) = s(bool,X428) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X428) = s(bool,X429) | s(bool,'T_0') = s(bool,X429)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13396,f1281])).
fof(f17827,plain,(
  ( ! [X424,X426,X425] : (~sP5('F_0',f(s(bool,X424)),'F_0') | s(bool,X424) = s(bool,X425) | s(bool,X425) = s(bool,X426) | s(bool,'T_0') = s(bool,X426)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15235,f522])).
fof(f15235,plain,(
  ( ! [X424,X426,X425] : (~sP5('F_0',f(s(bool,X424)),'F_0') | s(bool,X424) = s(bool,X425) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X425) = s(bool,X426) | s(bool,'T_0') = s(bool,X426)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13395,f1281])).
fof(f17820,plain,(
  ( ! [X401,X403,X402] : (s(bool,X401) = s(bool,f(s(bool,X401))) | s(bool,X401) = s(bool,X402) | s(bool,X402) = s(bool,X403) | s(bool,'T_0') = s(bool,X403)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f15228,f522])).
fof(f15228,plain,(
  ( ! [X401,X403,X402] : (s(bool,X401) = s(bool,f(s(bool,X401))) | s(bool,X401) = s(bool,X402) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X402) = s(bool,X403) | s(bool,'T_0') = s(bool,X403)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13099,f1281])).
fof(f17819,plain,(
  ( ! [X399,X398,X400] : (~sP5('T_0','T_0',f(s(bool,X398))) | s(bool,X398) = s(bool,X399) | s(bool,X399) = s(bool,X400) | s(bool,'T_0') = s(bool,X400)) ) | $spl47),
  inference(subsumption_resolution,[],[f15227,f522])).
fof(f15227,plain,(
  ( ! [X399,X398,X400] : (~sP5('T_0','T_0',f(s(bool,X398))) | s(bool,X398) = s(bool,X399) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X399) = s(bool,X400) | s(bool,'T_0') = s(bool,X400)) ) | $spl47),
  inference(superposition,[],[f13090,f1281])).
fof(f17818,plain,(
  ( ! [X397,X395,X396] : (sP5('F_0','T_0',f(s(bool,X395))) | s(bool,X395) = s(bool,X396) | s(bool,X396) = s(bool,X397) | s(bool,'T_0') = s(bool,X397)) ) | $spl47),
  inference(subsumption_resolution,[],[f15226,f522])).
fof(f15226,plain,(
  ( ! [X397,X395,X396] : (sP5('F_0','T_0',f(s(bool,X395))) | s(bool,X395) = s(bool,X396) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X396) = s(bool,X397) | s(bool,'T_0') = s(bool,X397)) ) | $spl47),
  inference(superposition,[],[f13089,f1281])).
fof(f17817,plain,(
  ( ! [X393,X392,X394] : (~sP5('T_0',f(s(bool,X392)),'T_0') | s(bool,X392) = s(bool,X393) | s(bool,X393) = s(bool,X394) | s(bool,'T_0') = s(bool,X394)) ) | $spl47),
  inference(subsumption_resolution,[],[f15225,f522])).
fof(f15225,plain,(
  ( ! [X393,X392,X394] : (~sP5('T_0',f(s(bool,X392)),'T_0') | s(bool,X392) = s(bool,X393) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X393) = s(bool,X394) | s(bool,'T_0') = s(bool,X394)) ) | $spl47),
  inference(superposition,[],[f12980,f1281])).
fof(f17816,plain,(
  ( ! [X389,X391,X390] : (sP5('F_0',f(s(bool,X389)),'T_0') | s(bool,X389) = s(bool,X390) | s(bool,X390) = s(bool,X391) | s(bool,'T_0') = s(bool,X391)) ) | $spl47),
  inference(subsumption_resolution,[],[f15224,f522])).
fof(f15224,plain,(
  ( ! [X389,X391,X390] : (sP5('F_0',f(s(bool,X389)),'T_0') | s(bool,X389) = s(bool,X390) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X390) = s(bool,X391) | s(bool,'T_0') = s(bool,X391)) ) | $spl47),
  inference(superposition,[],[f12979,f1281])).
fof(f17814,plain,(
  ( ! [X377,X379,X378] : (s(bool,'T_0') != s(bool,f(s(bool,X377))) | s(bool,X377) = s(bool,X378) | s(bool,X378) = s(bool,X379) | s(bool,'T_0') = s(bool,X379)) ) | $spl47),
  inference(subsumption_resolution,[],[f15220,f522])).
fof(f15220,plain,(
  ( ! [X377,X379,X378] : (s(bool,'T_0') != s(bool,f(s(bool,X377))) | s(bool,X377) = s(bool,X378) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X378) = s(bool,X379) | s(bool,'T_0') = s(bool,X379)) ) | $spl47),
  inference(superposition,[],[f2310,f1281])).
fof(f17813,plain,(
  ( ! [X372,X374,X370,X373,X371] : (s(bool,X370) != s(bool,X371) | s(bool,X371) = s(bool,X372) | s(bool,'T_0') = s(bool,X372) | s(bool,X370) = s(bool,X373) | s(bool,X373) = s(bool,X374) | s(bool,'T_0') = s(bool,X374)) )),
  inference(subsumption_resolution,[],[f15219,f522])).
fof(f15219,plain,(
  ( ! [X372,X374,X370,X373,X371] : (s(bool,X370) != s(bool,X371) | s(bool,X371) = s(bool,X372) | s(bool,'T_0') = s(bool,X372) | s(bool,X370) = s(bool,X373) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X373) = s(bool,X374) | s(bool,'T_0') = s(bool,X374)) )),
  inference(superposition,[],[f1700,f1281])).
fof(f17812,plain,(
  ( ! [X364,X360,X362,X365,X361,X363] : (s(X361,f(s(bool,X360))) != s(X361,g(s(bool,X360))) | s(fun(bool,X361),X362) = s(fun(bool,X361),X363) | s(bool,'T_0') = s(bool,sK17(X362,X363,X361,bool)) | s(bool,X360) = s(bool,X364) | s(bool,X364) = s(bool,X365) | s(bool,'T_0') = s(bool,X365)) )),
  inference(subsumption_resolution,[],[f15218,f522])).
fof(f15218,plain,(
  ( ! [X364,X360,X362,X365,X361,X363] : (s(X361,f(s(bool,X360))) != s(X361,g(s(bool,X360))) | s(fun(bool,X361),X362) = s(fun(bool,X361),X363) | s(bool,'T_0') = s(bool,sK17(X362,X363,X361,bool)) | s(bool,X360) = s(bool,X364) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X364) = s(bool,X365) | s(bool,'T_0') = s(bool,X365)) )),
  inference(superposition,[],[f1053,f1281])).
fof(f17811,plain,(
  ( ! [X356,X354,X355] : (s(bool,'T_0') != s(bool,X354) | s(bool,X354) = s(bool,X355) | s(bool,X355) = s(bool,X356) | s(bool,'T_0') = s(bool,X356)) )),
  inference(subsumption_resolution,[],[f15217,f522])).
fof(f15217,plain,(
  ( ! [X356,X354,X355] : (s(bool,'T_0') != s(bool,X354) | s(bool,X354) = s(bool,X355) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X355) = s(bool,X356) | s(bool,'T_0') = s(bool,X356)) )),
  inference(superposition,[],[f522,f1281])).
fof(f16499,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP1(f(s(bool,X1)),'T_0') | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | ($spl47 | $spl74)),
  inference(condensation,[],[f16498])).
fof(f16498,plain,(
  ( ! [X265,X267,X268,X266] : (sP1(f(s(bool,X266)),'T_0') | s(bool,'T_0') = s(bool,X265) | s(bool,X266) = s(bool,X267) | s(bool,X267) = s(bool,X268) | s(bool,'T_0') = s(bool,X268)) ) | ($spl47 | $spl74)),
  inference(duplicate_literal_removal,[],[f15182])).
fof(f15182,plain,(
  ( ! [X265,X267,X268,X266] : (sP1(f(s(bool,X266)),'T_0') | s(bool,'T_0') = s(bool,X265) | s(bool,X266) = s(bool,X267) | s(bool,'T_0') = s(bool,X265) | s(bool,X267) = s(bool,X268) | s(bool,'T_0') = s(bool,X268)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13255,f1281])).
fof(f16501,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP1('T_0',f(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | ($spl47 | $spl74)),
  inference(condensation,[],[f16500])).
fof(f16500,plain,(
  ( ! [X261,X263,X262,X264] : (sP1('T_0',f(s(bool,X262))) | s(bool,'T_0') = s(bool,X261) | s(bool,X262) = s(bool,X263) | s(bool,X263) = s(bool,X264) | s(bool,'T_0') = s(bool,X264)) ) | ($spl47 | $spl74)),
  inference(duplicate_literal_removal,[],[f15181])).
fof(f15181,plain,(
  ( ! [X261,X263,X262,X264] : (sP1('T_0',f(s(bool,X262))) | s(bool,'T_0') = s(bool,X261) | s(bool,X262) = s(bool,X263) | s(bool,'T_0') = s(bool,X261) | s(bool,X263) = s(bool,X264) | s(bool,'T_0') = s(bool,X264)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13242,f1281])).
fof(f16503,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f16502])).
fof(f16502,plain,(
  ( ! [X241,X239,X240,X238] : (s(bool,X239) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X238) | s(bool,X239) = s(bool,X240) | s(bool,X240) = s(bool,X241) | s(bool,'T_0') = s(bool,X241)) )),
  inference(duplicate_literal_removal,[],[f15176])).
fof(f15176,plain,(
  ( ! [X241,X239,X240,X238] : (s(bool,X239) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X238) | s(bool,X239) = s(bool,X240) | s(bool,'T_0') = s(bool,X238) | s(bool,X240) = s(bool,X241) | s(bool,'T_0') = s(bool,X241)) )),
  inference(superposition,[],[f7908,f1281])).
fof(f15129,plain,(
  ( ! [X606,X609,X611,X607,X608,X610] : (s(X608,f(s(bool,X609))) != s(X608,g(s(bool,X609))) | s(fun(bool,X608),X606) = s(fun(bool,X608),X607) | s(bool,X610) = s(bool,sK17(X606,X607,X608,bool)) | s(bool,'T_0') = s(bool,X609) | s(bool,X610) = s(bool,X611) | s(bool,'T_0') = s(bool,X611)) )),
  inference(superposition,[],[f401,f1281])).
fof(f15128,plain,(
  ( ! [X602,X605,X601,X603,X604] : (s(bool,X603) != s(bool,sK15(X601,bool,X602)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X602),X601)))) | s(bool,X604) = s(bool,sK16(X601,bool,X602)) | s(bool,'T_0') = s(bool,X603) | s(bool,X604) = s(bool,X605) | s(bool,'T_0') = s(bool,X605)) )),
  inference(superposition,[],[f400,f1281])).
fof(f15127,plain,(
  ( ! [X600,X596,X598,X597,X599] : (s(X597,happ(s(fun(bool,X597),X596),s(bool,X598))) = s(X597,happ(s(fun(bool,X597),X596),s(bool,sK15(X596,bool,X597)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X597),X596))) | s(bool,X599) = s(bool,sK16(X596,bool,X597)) | s(bool,'T_0') = s(bool,X598) | s(bool,X599) = s(bool,X600) | s(bool,'T_0') = s(bool,X600)) )),
  inference(superposition,[],[f1067,f1281])).
fof(f15124,plain,(
  ( ! [X580,X582,X578,X581,X577,X579] : (s(bool,X580) != s(bool,f(s(X577,sK17(X578,X579,bool,X577)))) | s(fun(X577,bool),X578) = s(fun(X577,bool),X579) | s(bool,X581) = s(bool,g(s(X577,sK17(X578,X579,bool,X577)))) | s(bool,'T_0') = s(bool,X580) | s(bool,X581) = s(bool,X582) | s(bool,'T_0') = s(bool,X582)) )),
  inference(superposition,[],[f401,f1281])).
fof(f17790,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK16(X0,bool,X1)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X1),X0)))) | s(bool,X2) = s(bool,sK15(X0,bool,X1)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f17789])).
fof(f17789,plain,(
  ( ! [X569,X571,X567,X568,X570] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X568),X567)))) | s(bool,'T_0') = s(bool,sK16(X567,bool,X568)) | s(bool,X570) = s(bool,sK15(X567,bool,X568)) | s(bool,'T_0') = s(bool,X569) | s(bool,X570) = s(bool,X571) | s(bool,'T_0') = s(bool,X571)) )),
  inference(subsumption_resolution,[],[f15123,f283])).
fof(f15123,plain,(
  ( ! [X569,X571,X567,X568,X570] : (s(bool,'F_0') != s(bool,X569) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X568),X567)))) | s(bool,'T_0') = s(bool,sK16(X567,bool,X568)) | s(bool,X570) = s(bool,sK15(X567,bool,X568)) | s(bool,'T_0') = s(bool,X569) | s(bool,X570) = s(bool,X571) | s(bool,'T_0') = s(bool,X571)) )),
  inference(superposition,[],[f1026,f1281])).
fof(f17776,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK17(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,f(s(X0,sK17(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f17775])).
fof(f17775,plain,(
  ( ! [X486,X488,X490,X485,X487,X489] : (s(fun(X485,bool),X486) = s(fun(X485,bool),X487) | s(bool,'T_0') = s(bool,g(s(X485,sK17(X486,X487,bool,X485)))) | s(bool,X489) = s(bool,f(s(X485,sK17(X486,X487,bool,X485)))) | s(bool,'T_0') = s(bool,X488) | s(bool,X489) = s(bool,X490) | s(bool,'T_0') = s(bool,X490)) )),
  inference(subsumption_resolution,[],[f15103,f283])).
fof(f15103,plain,(
  ( ! [X486,X488,X490,X485,X487,X489] : (s(bool,'F_0') != s(bool,X488) | s(fun(X485,bool),X486) = s(fun(X485,bool),X487) | s(bool,'T_0') = s(bool,g(s(X485,sK17(X486,X487,bool,X485)))) | s(bool,X489) = s(bool,f(s(X485,sK17(X486,X487,bool,X485)))) | s(bool,'T_0') = s(bool,X488) | s(bool,X489) = s(bool,X490) | s(bool,'T_0') = s(bool,X490)) )),
  inference(superposition,[],[f1054,f1281])).
fof(f17694,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f17693])).
fof(f17693,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) = s(bool,X22) | s(bool,X20) = s(bool,X23) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f14995,f294])).
fof(f14995,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) = s(bool,X22) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f357,f1281])).
fof(f14847,plain,(
  ( ! [X606,X609,X611,X607,X608,X610] : (s(X608,f(s(bool,X609))) != s(X608,g(s(bool,X609))) | s(fun(bool,X608),X606) = s(fun(bool,X608),X607) | s(bool,X609) = s(bool,X610) | s(bool,'T_0') = s(bool,X610) | s(bool,X611) = s(bool,sK17(X606,X607,X608,bool)) | s(bool,'T_0') = s(bool,X611)) )),
  inference(superposition,[],[f401,f1281])).
fof(f14846,plain,(
  ( ! [X602,X605,X601,X603,X604] : (s(bool,X603) != s(bool,sK15(X601,bool,X602)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X602),X601)))) | s(bool,X603) = s(bool,X604) | s(bool,'T_0') = s(bool,X604) | s(bool,X605) = s(bool,sK16(X601,bool,X602)) | s(bool,'T_0') = s(bool,X605)) )),
  inference(superposition,[],[f400,f1281])).
fof(f14845,plain,(
  ( ! [X600,X596,X598,X597,X599] : (s(X597,happ(s(fun(bool,X597),X596),s(bool,X598))) = s(X597,happ(s(fun(bool,X597),X596),s(bool,sK15(X596,bool,X597)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X597),X596))) | s(bool,X598) = s(bool,X599) | s(bool,'T_0') = s(bool,X599) | s(bool,X600) = s(bool,sK16(X596,bool,X597)) | s(bool,'T_0') = s(bool,X600)) )),
  inference(superposition,[],[f1067,f1281])).
fof(f14842,plain,(
  ( ! [X580,X582,X578,X581,X577,X579] : (s(bool,X580) != s(bool,f(s(X577,sK17(X578,X579,bool,X577)))) | s(fun(X577,bool),X578) = s(fun(X577,bool),X579) | s(bool,X580) = s(bool,X581) | s(bool,'T_0') = s(bool,X581) | s(bool,X582) = s(bool,g(s(X577,sK17(X578,X579,bool,X577)))) | s(bool,'T_0') = s(bool,X582)) )),
  inference(superposition,[],[f401,f1281])).
fof(f14598,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun(X111,bool),X112),s(X111,sK15(X112,X111,bool)))) | s(bool,X113) = s(bool,X114) | s(bool,X114) = s(bool,happ(s(fun(X111,bool),X112),s(X111,sK16(X112,X111,bool)))) | s(bool,X113) = s(bool,X115) | s(bool,'T_0') = s(bool,X115) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X111,bool),X112)))) )),
  inference(superposition,[],[f1281,f1067])).
fof(f17454,plain,(
  ( ! [X109,X110,X108] : (s(bool,X108) = s(bool,X109) | s(bool,X109) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,X108) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f14597,f522])).
fof(f14597,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X108) = s(bool,X109) | s(bool,X109) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,X108) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f1281,f1243])).
fof(f17451,plain,(
  ( ! [X90,X91,X89] : (s(bool,X89) = s(bool,X90) | s(bool,X90) = s(bool,f(s(bool,'F_0'))) | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f14592,f522])).
fof(f14592,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X89) = s(bool,X90) | s(bool,X90) = s(bool,f(s(bool,'F_0'))) | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f1281,f13099])).
fof(f14560,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,X16) = s(bool,X17) | s(bool,X18) = s(bool,X19) | s(bool,X17) = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X16) = s(bool,X21)) )),
  inference(superposition,[],[f1281,f6906])).
fof(f14558,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f1281,f1281])).
fof(f14557,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,X156) = s(bool,happ(s(fun(X154,bool),X155),s(X154,sK15(X155,X154,bool)))) | s(bool,X156) = s(bool,X157) | s(bool,'T_0') = s(bool,happ(s(fun(X154,bool),X155),s(X154,sK16(X155,X154,bool)))) | s(bool,X157) = s(bool,X158) | s(bool,'T_0') = s(bool,X158) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X154,bool),X155)))) )),
  inference(superposition,[],[f1281,f1067])).
fof(f14523,plain,(
  ( ! [X171,X169,X172,X170,X168] : (s(bool,X170) = s(bool,happ(s(fun(X168,bool),X169),s(X168,sK15(X169,X168,bool)))) | s(bool,X171) = s(bool,happ(s(fun(X168,bool),X169),s(X168,sK16(X169,X168,bool)))) | s(bool,'T_0') = s(bool,X170) | s(bool,X171) = s(bool,X172) | s(bool,'T_0') = s(bool,X172) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X168,bool),X169)))) )),
  inference(superposition,[],[f1281,f1067])).
fof(f14455,plain,(
  ( ! [X171,X169,X172,X170,X168] : (s(bool,X170) = s(bool,happ(s(fun(X168,bool),X169),s(X168,sK15(X169,X168,bool)))) | s(bool,X170) = s(bool,X171) | s(bool,'T_0') = s(bool,X171) | s(bool,X172) = s(bool,happ(s(fun(X168,bool),X169),s(X168,sK16(X169,X168,bool)))) | s(bool,'T_0') = s(bool,X172) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X168,bool),X169)))) )),
  inference(superposition,[],[f1281,f1067])).
fof(f17310,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X1) = s(bool,X3)) )),
  inference(condensation,[],[f17309])).
fof(f17309,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (s(bool,X104) = s(bool,X105) | s(bool,X105) = s(bool,X106) | s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,X107) | s(bool,X108) = s(bool,X109) | s(bool,X109) = s(bool,X110) | s(bool,'T_0') = s(bool,X110) | s(bool,X104) = s(bool,X108)) )),
  inference(duplicate_literal_removal,[],[f14440])).
fof(f14440,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (s(bool,X104) = s(bool,X105) | s(bool,X105) = s(bool,X106) | s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,X107) | s(bool,X108) = s(bool,X109) | s(bool,X109) = s(bool,X110) | s(bool,'T_0') = s(bool,X110) | s(bool,X104) = s(bool,X108)) )),
  inference(superposition,[],[f1281,f1281])).
fof(f14205,plain,(
  ( ! [X14,X17,X15,X16] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat'))))))),X16,X17) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2193,f1067])).
fof(f14202,plain,(
  ( ! [X4,X2,X5,X3] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))))),X4,X5) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2193,f355])).
fof(f14152,plain,(
  ( ! [X14,X17,X15,X16] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat'))))))),X16,X17) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2192,f1067])).
fof(f14149,plain,(
  ( ! [X4,X2,X5,X3] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))))),X4,X5) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2192,f355])).
fof(f14096,plain,(
  ( ! [X14,X17,X15,X16] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | p(s(bool,X16)) | ~p(s(bool,X17)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2191,f1067])).
fof(f14093,plain,(
  ( ! [X4,X2,X5,X3] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | p(s(bool,X4)) | ~p(s(bool,X5)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2191,f355])).
fof(f14074,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f13090,f283])).
fof(f14058,plain,(
  ( ! [X0] : (sP5('F_0','T_0',f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f13089,f283])).
fof(f14056,plain,(
  sP3('T_0',f(s(bool,'F_0')),'F_0') | $spl47),
  inference(unit_resulting_resolution,[],[f13089,f409])).
fof(f14048,plain,(
  ( ! [X14,X17,X15,X16] : (sP11(X16,X17,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat')))))))) | p(s(bool,X17)) | p(s(bool,X16)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2188,f1067])).
fof(f14045,plain,(
  ( ! [X4,X2,X5,X3] : (sP11(X4,X5,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2))))))) | p(s(bool,X5)) | p(s(bool,X4)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2188,f355])).
fof(f14027,plain,(
  ( ! [X0] : (~sP5('T_0',f(s(bool,X0)),'T_0') | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f12980,f283])).
fof(f14024,plain,(
  ~sP4('T_0','T_0',f(s(bool,'F_0'))) | $spl47),
  inference(unit_resulting_resolution,[],[f275,f512,f507,f12980,f412])).
fof(f14013,plain,(
  ( ! [X0] : (sP5('F_0',f(s(bool,X0)),'T_0') | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f12979,f283])).
fof(f14011,plain,(
  sP3(f(s(bool,'F_0')),'T_0','F_0') | $spl47),
  inference(unit_resulting_resolution,[],[f12979,f409])).
fof(f14003,plain,(
  ( ! [X14,X17,X15,X16] : (~sP10(X16,X17,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat')))))))) | ~p(s(bool,X16)) | p(s(bool,X17)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2187,f1067])).
fof(f14000,plain,(
  ( ! [X4,X2,X5,X3] : (~sP10(X4,X5,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2))))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2187,f355])).
fof(f13981,plain,(
  ( ! [X26,X27] : (sP1(f(s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,X26,bool))))),'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,X26,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X26,bool),X27)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13255,f1067])).
fof(f13942,plain,(
  ( ! [X26,X27] : (sP1('T_0',f(s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,X26,bool)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,X26,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X26,bool),X27)))) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13242,f1067])).
fof(f13910,plain,(
  ( ! [X14,X17,X15,X16] : (~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | p(s(bool,X16)) | p(s(bool,X17)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2186,f1067])).
fof(f13907,plain,(
  ( ! [X4,X2,X5,X3] : (~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | p(s(bool,X4)) | p(s(bool,X5)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2186,f355])).
fof(f13889,plain,(
  ( ! [X0,X1] : (sP1(f(s(bool,X0)),X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13486,f283])).
fof(f13877,plain,(
  ( ! [X0,X1] : (sP1(X1,f(s(bool,X0))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13485,f283])).
fof(f13871,plain,(
  ( ! [X14,X17,X15,X16] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | p(s(bool,X17)) | p(s(bool,X16)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2183,f1067])).
fof(f13868,plain,(
  ( ! [X4,X2,X5,X3] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | p(s(bool,X5)) | p(s(bool,X4)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2183,f355])).
fof(f13828,plain,(
  ( ! [X14,X17,X15,X16] : (~sP7(X16,X17,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat')))))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2182,f1067])).
fof(f13825,plain,(
  ( ! [X4,X2,X5,X3] : (~sP7(X4,X5,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2))))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2182,f355])).
fof(f13785,plain,(
  ( ! [X14,X17,X15,X16] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | ~p(s(bool,X17)) | ~p(s(bool,X16)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2181,f1067])).
fof(f13782,plain,(
  ( ! [X4,X2,X5,X3] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | ~p(s(bool,X5)) | ~p(s(bool,X4)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2181,f355])).
fof(f13746,plain,(
  ( ! [X14,X17,X15,X16] : (~sP6(X16,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X17) | p(s(bool,X17)) | ~p(s(bool,X16)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2180,f1067])).
fof(f13743,plain,(
  ( ! [X4,X2,X5,X3] : (~sP6(X4,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2180,f355])).
fof(f13722,plain,(
  ( ! [X0] : (sP4('T_0',f(s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13653,f283])).
fof(f13702,plain,(
  ( ! [X0] : (~sP2(f(s(bool,X0)),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13622,f283])).
fof(f13693,plain,(
  ( ! [X14,X17,X15,X16] : (~sP3(X16,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X17) | p(s(bool,X17)) | ~p(s(bool,X16)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2179,f1067])).
fof(f13690,plain,(
  ( ! [X4,X2,X5,X3] : (~sP3(X4,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2179,f355])).
fof(f13671,plain,(
  ( ! [X0] : (~sP2('F_0',f(s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13487,f283])).
fof(f13656,plain,(
  ( ! [X0] : (sP5('T_0','F_0',f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13403,f283])).
fof(f13653,plain,(
  sP4('T_0',f(s(bool,'F_0')),'F_0') | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f13403,f410])).
fof(f13646,plain,(
  ( ! [X14,X17,X15,X16] : (~sP4(X16,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X17) | p(s(bool,X17)) | ~p(s(bool,X16)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2178,f1067])).
fof(f13643,plain,(
  ( ! [X4,X2,X5,X3] : (~sP4(X4,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2178,f355])).
fof(f13625,plain,(
  ( ! [X0] : (~sP5('F_0','F_0',f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13402,f283])).
fof(f13622,plain,(
  ~sP2(f(s(bool,'F_0')),'F_0','F_0') | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f498,f509,f506,f13402,f413])).
fof(f13609,plain,(
  ( ! [X0] : (sP5('T_0',f(s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13396,f283])).
fof(f13490,plain,(
  ( ! [X0] : (~sP5('F_0',f(s(bool,X0)),'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13395,f283])).
fof(f13487,plain,(
  ~sP2('F_0',f(s(bool,'F_0')),'F_0') | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f498,f508,f506,f13395,f414])).
fof(f13486,plain,(
  ( ! [X59] : (sP1(f(s(bool,'F_0')),X59) | s(bool,'F_0') = s(bool,X59)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f13466,f498])).
fof(f13466,plain,(
  ( ! [X59] : (p(s(bool,'F_0')) | sP1(f(s(bool,'F_0')),X59) | s(bool,'F_0') = s(bool,X59)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f9616,f13099])).
fof(f13485,plain,(
  ( ! [X58] : (sP1(X58,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X58)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f13465,f498])).
fof(f13465,plain,(
  ( ! [X58] : (p(s(bool,'F_0')) | sP1(X58,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X58)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f9614,f13099])).
fof(f13482,plain,(
  ( ! [X52,X51] : (sP1(f(s(bool,'F_0')),X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f13459,f498])).
fof(f13459,plain,(
  ( ! [X52,X51] : (p(s(bool,'F_0')) | sP1(f(s(bool,'F_0')),X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f1705,f13099])).
fof(f13481,plain,(
  ( ! [X50,X49] : (sP1(X49,f(s(bool,'F_0'))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f13458,f498])).
fof(f13458,plain,(
  ( ! [X50,X49] : (p(s(bool,'F_0')) | sP1(X49,f(s(bool,'F_0'))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f1704,f13099])).
fof(f13480,plain,(
  ( ! [X45,X44] : (s(bool,X44) != s(bool,f(s(bool,'F_0'))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f13453,f498])).
fof(f13453,plain,(
  ( ! [X45,X44] : (p(s(bool,'F_0')) | s(bool,X44) != s(bool,f(s(bool,'F_0'))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f1381,f13099])).
fof(f13473,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | ($spl47 | $spl74)),
  inference(subsumption_resolution,[],[f13412,f522])).
fof(f13412,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,f(s(bool,X7))) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f13099,f6906])).
fof(f13395,plain,(
  ~sP5('F_0',f(s(bool,'F_0')),'F_0') | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f498,f13099,f430])).
fof(f13396,plain,(
  sP5('T_0',f(s(bool,'F_0')),'F_0') | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f275,f13099,f428])).
fof(f13402,plain,(
  ~sP5('F_0','F_0',f(s(bool,'F_0'))) | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f498,f13099,f430])).
fof(f13403,plain,(
  sP5('T_0','F_0',f(s(bool,'F_0'))) | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f275,f13099,f428])).
fof(f13255,plain,(
  ( ! [X0] : (sP1(f(s(bool,X0)),'T_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f12977,f283])).
fof(f13242,plain,(
  ( ! [X0] : (sP1('T_0',f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl47 | $spl74)),
  inference(superposition,[],[f12976,f283])).
fof(f13119,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(trivial_inequality_removal,[],[f13110])).
fof(f13110,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f2310,f608])).
fof(f13239,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,f(s(bool,X7))) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl47),
  inference(subsumption_resolution,[],[f13108,f522])).
fof(f13108,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,f(s(bool,X7))) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl47),
  inference(superposition,[],[f2310,f6906])).
fof(f13100,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f2310,f283])).
fof(f12976,plain,(
  sP1('T_0',f(s(bool,'F_0'))) | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f522,f2310,f9808])).
fof(f12977,plain,(
  sP1(f(s(bool,'F_0')),'T_0') | ($spl47 | $spl74)),
  inference(unit_resulting_resolution,[],[f522,f2310,f9868])).
fof(f12979,plain,(
  sP5('F_0',f(s(bool,'F_0')),'T_0') | $spl47),
  inference(unit_resulting_resolution,[],[f498,f2310,f427])).
fof(f12980,plain,(
  ~sP5('T_0',f(s(bool,'F_0')),'T_0') | $spl47),
  inference(unit_resulting_resolution,[],[f275,f2310,f429])).
fof(f13089,plain,(
  sP5('F_0','T_0',f(s(bool,'F_0'))) | $spl47),
  inference(unit_resulting_resolution,[],[f498,f2310,f427])).
fof(f13090,plain,(
  ~sP5('T_0','T_0',f(s(bool,'F_0'))) | $spl47),
  inference(unit_resulting_resolution,[],[f275,f2310,f429])).
fof(f10254,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X19,X20,bool,bool)))) | s(bool,X21) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X21) = s(bool,sK17(X19,X20,bool,bool))) )),
  inference(superposition,[],[f1054,f6906])).
fof(f12772,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X28,X29,bool,bool)))) | s(bool,'F_0') = s(bool,sK17(X28,X29,bool,bool))) ) | $spl74),
  inference(forward_demodulation,[],[f10257,f877])).
fof(f10257,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,f(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))))) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X28,X29,bool,bool)))) | s(bool,'F_0') = s(bool,sK17(X28,X29,bool,bool))) ) | $spl74),
  inference(superposition,[],[f1054,f8040])).
fof(f12700,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK16(X1,bool,X0)) | s(bool,'T_0') = s(bool,sK15(X1,bool,X0)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),X1)))) )),
  inference(subsumption_resolution,[],[f12689,f498])).
fof(f12689,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK16(X1,bool,X0)) | s(bool,'T_0') = s(bool,sK15(X1,bool,X0)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),X1)))) )),
  inference(superposition,[],[f2213,f283])).
fof(f12688,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X0),happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK15(X2,X1,fun(bool,X0)))))))) | s(bool,'T_0') = s(bool,sK16(happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK16(X2,X1,fun(bool,X0)))),bool,X0)) | s(bool,'T_0') = s(bool,sK15(happ(s(fun(X1,fun(bool,X0)),X2),s(X1,sK16(X2,X1,fun(bool,X0)))),bool,X0)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,fun(bool,X0)),X2)))) )),
  inference(superposition,[],[f2213,f1067])).
fof(f12639,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,happ(s(fun(X57,bool),X58),s(X57,sK15(X58,X57,bool))))) | s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X59))))) = s(bool,happ(s(fun(X57,bool),X58),s(X57,sK16(X58,X57,bool)))) | s('type.rat.rat',X59) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X57,bool),X58)))) )),
  inference(superposition,[],[f2173,f1067])).
fof(f12652,plain,(
  ( ! [X45,X46] : (s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X45))))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X46))))) | s('type.rat.rat',X46) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat',X45) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f12633,f275])).
fof(f12633,plain,(
  ( ! [X45,X46] : (~p(s(bool,'T_0')) | s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X45))))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X46))))) | s('type.rat.rat',X46) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat',X45) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2173,f872])).
fof(f12640,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X18))))) | s('type.rat.rat',X18) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X19) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X17) = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f12621,f275])).
fof(f12621,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'T_0')) | s(bool,X17) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X18))))) | s('type.rat.rat',X18) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X19) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X17) = s(bool,X19)) )),
  inference(superposition,[],[f2173,f6906])).
fof(f12312,plain,(
  ( ! [X14,X17,X15,X16] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | ~p(s(bool,X16)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2190,f1067])).
fof(f12309,plain,(
  ( ! [X4,X2,X5,X3] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | ~p(s(bool,X4)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2190,f355])).
fof(f12251,plain,(
  ( ! [X14,X17,X15,X16] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | p(s(bool,X17)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2189,f1067])).
fof(f12248,plain,(
  ( ! [X4,X2,X5,X3] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | p(s(bool,X5)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2189,f355])).
fof(f12201,plain,(
  ( ! [X14,X17,X15,X16] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | ~p(s(bool,X16)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2185,f1067])).
fof(f12198,plain,(
  ( ! [X4,X2,X5,X3] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | ~p(s(bool,X4)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2185,f355])).
fof(f12150,plain,(
  ( ! [X14,X17,X15,X16] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK15(X15,X14,'type.rat.rat'))))))),X16,X17) | ~p(s(bool,X17)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X14,'type.rat.rat'),X15),s(X14,sK16(X15,X14,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X14,'type.rat.rat'),X15)))) )),
  inference(superposition,[],[f2184,f1067])).
fof(f12147,plain,(
  ( ! [X4,X2,X5,X3] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))))),X4,X5) | ~p(s(bool,X5)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f2184,f355])).
fof(f10922,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f1136,f408])).
fof(f10884,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1135,f408])).
fof(f10735,plain,(
  ( ! [X61] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X61) | s(bool,'F_0') = s(bool,X61)) ) | $spl74),
  inference(subsumption_resolution,[],[f10715,f498])).
fof(f10715,plain,(
  ( ! [X61] : (p(s(bool,'F_0')) | sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X61) | s(bool,'F_0') = s(bool,X61)) ) | $spl74),
  inference(superposition,[],[f9616,f1243])).
fof(f10734,plain,(
  ( ! [X60] : (sP1(X60,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,X60)) ) | $spl74),
  inference(subsumption_resolution,[],[f10714,f498])).
fof(f10714,plain,(
  ( ! [X60] : (p(s(bool,'F_0')) | sP1(X60,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,X60)) ) | $spl74),
  inference(superposition,[],[f9614,f1243])).
fof(f10375,plain,(
  ( ! [X156,X157,X155] : (p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK15(X156,X155,bool))))) | sP1(happ(s(fun(X155,bool),X156),s(X155,sK16(X156,X155,bool))),X157) | s(bool,'F_0') = s(bool,X157) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X155,bool),X156)))) ) | $spl74),
  inference(superposition,[],[f9616,f1067])).
fof(f10374,plain,(
  ( ! [X154,X152,X153] : (p(s(bool,happ(s(fun(X152,bool),X153),s(X152,sK15(X153,X152,bool))))) | sP1(X154,happ(s(fun(X152,bool),X153),s(X152,sK16(X153,X152,bool)))) | s(bool,'F_0') = s(bool,X154) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X152,bool),X153)))) ) | $spl74),
  inference(superposition,[],[f9614,f1067])).
fof(f10373,plain,(
  ( ! [X151,X149,X150] : (~p(s(bool,happ(s(fun(X149,bool),X150),s(X149,sK15(X150,X149,bool))))) | s(bool,X151) = s(bool,happ(s(fun(X149,bool),X150),s(X149,sK16(X150,X149,bool)))) | s(bool,'F_0') = s(bool,X151) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X149,bool),X150)))) ) | $spl74),
  inference(superposition,[],[f9608,f1067])).
fof(f10411,plain,(
  ( ! [X146,X145] : (s(bool,'T_0') != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK15(X146,X145,bool)))) | s(bool,'T_0') = s(bool,happ(s(fun(X145,bool),X146),s(X145,sK16(X146,X145,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X145,bool),X146)))) )),
  inference(forward_demodulation,[],[f10371,f877])).
fof(f10371,plain,(
  ( ! [X146,X145] : (s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK15(X146,X145,bool)))) | s(bool,'T_0') = s(bool,happ(s(fun(X145,bool),X146),s(X145,sK16(X146,X145,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X145,bool),X146)))) )),
  inference(superposition,[],[f7908,f1067])).
fof(f10405,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,X26,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X26,bool),X27))) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,X26,bool))))) ) | $spl74),
  inference(forward_demodulation,[],[f10325,f877])).
fof(f10325,plain,(
  ( ! [X26,X27] : (s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) = s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,X26,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X26,bool),X27))) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,X26,bool))))) ) | $spl74),
  inference(superposition,[],[f1067,f8040])).
fof(f10400,plain,(
  ( ! [X26,X27] : (s(X27,happ(s(fun(bool,X27),X26),s(bool,'T_0'))) = s(X27,happ(s(fun(bool,X27),X26),s(bool,sK15(X26,bool,X27)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X27),X26))) | s(bool,'F_0') = s(bool,sK16(X26,bool,X27))) ) | $spl74),
  inference(forward_demodulation,[],[f10315,f877])).
fof(f10315,plain,(
  ( ! [X26,X27] : (s(X27,happ(s(fun(bool,X27),X26),s(bool,sK15(X26,bool,X27)))) = s(X27,happ(s(fun(bool,X27),X26),s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X27),X26))) | s(bool,'F_0') = s(bool,sK16(X26,bool,X27))) ) | $spl74),
  inference(superposition,[],[f1067,f8040])).
fof(f10161,plain,(
  ( ! [X41,X40] : (s(bool,X41) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X40))))) | s(bool,'F_0') = s(bool,X41) | s('type.rat.rat',X40) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl74),
  inference(subsumption_resolution,[],[f10137,f275])).
fof(f10137,plain,(
  ( ! [X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X40))))) | s(bool,'F_0') = s(bool,X41) | s('type.rat.rat',X40) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl74),
  inference(superposition,[],[f9608,f872])).
fof(f10149,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,'F_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X17) = s(bool,X19)) ) | $spl74),
  inference(subsumption_resolution,[],[f10128,f275])).
fof(f10128,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'T_0')) | s(bool,X17) = s(bool,X18) | s(bool,'F_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X17) = s(bool,X19)) ) | $spl74),
  inference(superposition,[],[f9608,f6906])).
fof(f10012,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(X66,happ(s(fun(X65,X66),X67),s(X65,X68))) != s(X66,happ(s(fun(X65,X66),X67),s(X65,X69))) | s(X65,X68) = s(X65,X69) | s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,X66),X67)))) ) | $spl74),
  inference(subsumption_resolution,[],[f10011,f275])).
fof(f10011,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | s(X66,happ(s(fun(X65,X66),X67),s(X65,X68))) != s(X66,happ(s(fun(X65,X66),X67),s(X65,X69))) | s(X65,X68) = s(X65,X69) | s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,X66),X67)))) ) | $spl74),
  inference(forward_demodulation,[],[f10008,f877])).
fof(f10008,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | s(X66,happ(s(fun(X65,X66),X67),s(X65,X68))) != s(X66,happ(s(fun(X65,X66),X67),s(X65,X69))) | s(X65,X68) = s(X65,X69) | s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X65,X66),X67)))) ) | $spl74),
  inference(superposition,[],[f398,f8040])).
fof(f9893,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,sK15(X26,bool,X27)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X27),X26)))) | s(bool,'F_0') = s(bool,sK16(X26,bool,X27))) ) | $spl74),
  inference(forward_demodulation,[],[f9888,f877])).
fof(f9888,plain,(
  ( ! [X26,X27] : (s(bool,sK15(X26,bool,X27)) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X27),X26)))) | s(bool,'F_0') = s(bool,sK16(X26,bool,X27))) ) | $spl74),
  inference(superposition,[],[f400,f8040])).
fof(f9868,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl74),
  inference(subsumption_resolution,[],[f9847,f498])).
fof(f9847,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | sP1(X0,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl74),
  inference(superposition,[],[f9616,f283])).
fof(f9808,plain,(
  ( ! [X0,X1] : (sP1(X1,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl74),
  inference(subsumption_resolution,[],[f9787,f498])).
fof(f9787,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | sP1(X1,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl74),
  inference(superposition,[],[f9614,f283])).
fof(f9528,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl74),
  inference(superposition,[],[f7908,f8040])).
fof(f9680,plain,(
  ( ! [X123,X124,X122] : (s(X124,f(s(bool,'T_0'))) != s(X124,g(s(bool,'T_0'))) | s(fun(bool,X124),X122) = s(fun(bool,X124),X123) | s(bool,'F_0') = s(bool,sK17(X122,X123,X124,bool))) ) | $spl74),
  inference(forward_demodulation,[],[f9527,f877])).
fof(f9527,plain,(
  ( ! [X123,X124,X122] : (s(X124,f(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))))) != s(X124,g(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))))) | s(fun(bool,X124),X122) = s(fun(bool,X124),X123) | s(bool,'F_0') = s(bool,sK17(X122,X123,X124,bool))) ) | $spl74),
  inference(superposition,[],[f401,f8040])).
fof(f9679,plain,(
  ( ! [X121,X120,X119] : (s(bool,'T_0') != s(bool,f(s(X119,sK17(X120,X121,bool,X119)))) | s(fun(X119,bool),X120) = s(fun(X119,bool),X121) | s(bool,'F_0') = s(bool,g(s(X119,sK17(X120,X121,bool,X119))))) ) | $spl74),
  inference(forward_demodulation,[],[f9526,f877])).
fof(f9526,plain,(
  ( ! [X121,X120,X119] : (s(bool,f(s(X119,sK17(X120,X121,bool,X119)))) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(fun(X119,bool),X120) = s(fun(X119,bool),X121) | s(bool,'F_0') = s(bool,g(s(X119,sK17(X120,X121,bool,X119))))) ) | $spl74),
  inference(superposition,[],[f401,f8040])).
fof(f9648,plain,(
  ( ! [X66,X64,X65] : (~sP8(X64,X65,X66) | sP9(X64,X65,X66) | s(bool,'F_0') = s(bool,X64)) ) | $spl74),
  inference(subsumption_resolution,[],[f9647,f275])).
fof(f9647,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | ~sP8(X64,X65,X66) | sP9(X64,X65,X66) | s(bool,'F_0') = s(bool,X64)) ) | $spl74),
  inference(forward_demodulation,[],[f9500,f877])).
fof(f9500,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~sP8(X64,X65,X66) | sP9(X64,X65,X66) | s(bool,'F_0') = s(bool,X64)) ) | $spl74),
  inference(superposition,[],[f483,f8040])).
fof(f9646,plain,(
  ( ! [X61,X62,X63] : (~sP10(X61,X62,X63) | sP11(X61,X62,X63) | s(bool,'F_0') = s(bool,X61)) ) | $spl74),
  inference(subsumption_resolution,[],[f9645,f275])).
fof(f9645,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | ~sP10(X61,X62,X63) | sP11(X61,X62,X63) | s(bool,'F_0') = s(bool,X61)) ) | $spl74),
  inference(forward_demodulation,[],[f9499,f877])).
fof(f9499,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~sP10(X61,X62,X63) | sP11(X61,X62,X63) | s(bool,'F_0') = s(bool,X61)) ) | $spl74),
  inference(superposition,[],[f482,f8040])).
fof(f9644,plain,(
  ( ! [X59,X60,X58] : (~sP11(X58,X60,X59) | p(s(bool,X60)) | ~p(s(bool,X59)) | s(bool,'F_0') = s(bool,X58)) ) | $spl74),
  inference(subsumption_resolution,[],[f9643,f275])).
fof(f9643,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~sP11(X58,X60,X59) | s(bool,'F_0') = s(bool,X58)) ) | $spl74),
  inference(forward_demodulation,[],[f9498,f877])).
fof(f9498,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~sP11(X58,X60,X59) | s(bool,'F_0') = s(bool,X58)) ) | $spl74),
  inference(superposition,[],[f479,f8040])).
fof(f9642,plain,(
  ( ! [X57,X56,X55] : (sP11(X55,X56,X57) | ~p(s(bool,X56)) | s(bool,'F_0') = s(bool,X55)) ) | $spl74),
  inference(subsumption_resolution,[],[f9641,f275])).
fof(f9641,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP11(X55,X56,X57) | s(bool,'F_0') = s(bool,X55)) ) | $spl74),
  inference(forward_demodulation,[],[f9497,f877])).
fof(f9497,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X56)) | sP11(X55,X56,X57) | s(bool,'F_0') = s(bool,X55)) ) | $spl74),
  inference(superposition,[],[f478,f8040])).
fof(f9640,plain,(
  ( ! [X54,X52,X53] : (sP11(X52,X54,X53) | p(s(bool,X53)) | s(bool,'F_0') = s(bool,X52)) ) | $spl74),
  inference(subsumption_resolution,[],[f9639,f275])).
fof(f9639,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | sP11(X52,X54,X53) | s(bool,'F_0') = s(bool,X52)) ) | $spl74),
  inference(forward_demodulation,[],[f9496,f877])).
fof(f9496,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | p(s(bool,X53)) | sP11(X52,X54,X53) | s(bool,'F_0') = s(bool,X52)) ) | $spl74),
  inference(superposition,[],[f477,f8040])).
fof(f9638,plain,(
  ( ! [X50,X51,X49] : (sP11(X51,X50,X49) | p(s(bool,X50)) | p(s(bool,X51)) | s(bool,'F_0') = s(bool,X49)) ) | $spl74),
  inference(subsumption_resolution,[],[f9637,f275])).
fof(f9637,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,X50)) | sP11(X51,X50,X49) | p(s(bool,X51)) | s(bool,'F_0') = s(bool,X49)) ) | $spl74),
  inference(forward_demodulation,[],[f9495,f877])).
fof(f9495,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | p(s(bool,X50)) | sP11(X51,X50,X49) | p(s(bool,X51)) | s(bool,'F_0') = s(bool,X49)) ) | $spl74),
  inference(superposition,[],[f476,f8040])).
fof(f9636,plain,(
  ( ! [X47,X48,X46] : (~sP10(X48,X47,X46) | ~p(s(bool,X48)) | p(s(bool,X47)) | s(bool,'F_0') = s(bool,X46)) ) | $spl74),
  inference(subsumption_resolution,[],[f9635,f275])).
fof(f9635,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | ~p(s(bool,X48)) | ~sP10(X48,X47,X46) | s(bool,'F_0') = s(bool,X46)) ) | $spl74),
  inference(forward_demodulation,[],[f9494,f877])).
fof(f9494,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | p(s(bool,X47)) | ~p(s(bool,X48)) | ~sP10(X48,X47,X46) | s(bool,'F_0') = s(bool,X46)) ) | $spl74),
  inference(superposition,[],[f472,f8040])).
fof(f9634,plain,(
  ( ! [X45,X43,X44] : (~sP9(X43,X45,X44) | p(s(bool,X45)) | p(s(bool,X44)) | s(bool,'F_0') = s(bool,X43)) ) | $spl74),
  inference(subsumption_resolution,[],[f9633,f275])).
fof(f9633,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | p(s(bool,X45)) | ~sP9(X43,X45,X44) | s(bool,'F_0') = s(bool,X43)) ) | $spl74),
  inference(forward_demodulation,[],[f9493,f877])).
fof(f9493,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | p(s(bool,X44)) | p(s(bool,X45)) | ~sP9(X43,X45,X44) | s(bool,'F_0') = s(bool,X43)) ) | $spl74),
  inference(superposition,[],[f462,f8040])).
fof(f9632,plain,(
  ( ! [X41,X42,X40] : (sP9(X40,X41,X42) | ~p(s(bool,X41)) | s(bool,'F_0') = s(bool,X40)) ) | $spl74),
  inference(subsumption_resolution,[],[f9631,f275])).
fof(f9631,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | sP9(X40,X41,X42) | s(bool,'F_0') = s(bool,X40)) ) | $spl74),
  inference(forward_demodulation,[],[f9492,f877])).
fof(f9492,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X41)) | sP9(X40,X41,X42) | s(bool,'F_0') = s(bool,X40)) ) | $spl74),
  inference(superposition,[],[f461,f8040])).
fof(f9630,plain,(
  ( ! [X39,X37,X38] : (sP9(X37,X39,X38) | ~p(s(bool,X38)) | s(bool,'F_0') = s(bool,X37)) ) | $spl74),
  inference(subsumption_resolution,[],[f9629,f275])).
fof(f9629,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | sP9(X37,X39,X38) | s(bool,'F_0') = s(bool,X37)) ) | $spl74),
  inference(forward_demodulation,[],[f9491,f877])).
fof(f9491,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X38)) | sP9(X37,X39,X38) | s(bool,'F_0') = s(bool,X37)) ) | $spl74),
  inference(superposition,[],[f460,f8040])).
fof(f9628,plain,(
  ( ! [X35,X36,X34] : (~sP8(X34,X35,X36) | p(s(bool,X36)) | p(s(bool,X35)) | s(bool,'F_0') = s(bool,X34)) ) | $spl74),
  inference(subsumption_resolution,[],[f9627,f275])).
fof(f9627,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,X35)) | p(s(bool,X36)) | ~sP8(X34,X35,X36) | s(bool,'F_0') = s(bool,X34)) ) | $spl74),
  inference(forward_demodulation,[],[f9490,f877])).
fof(f9490,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | p(s(bool,X35)) | p(s(bool,X36)) | ~sP8(X34,X35,X36) | s(bool,'F_0') = s(bool,X34)) ) | $spl74),
  inference(superposition,[],[f455,f8040])).
fof(f9626,plain,(
  ( ! [X33,X31,X32] : (~sP7(X33,X32,X31) | p(s(bool,X33)) | ~p(s(bool,X32)) | s(bool,'F_0') = s(bool,X31)) ) | $spl74),
  inference(subsumption_resolution,[],[f9625,f275])).
fof(f9625,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP7(X33,X32,X31) | s(bool,'F_0') = s(bool,X31)) ) | $spl74),
  inference(forward_demodulation,[],[f9489,f877])).
fof(f9489,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP7(X33,X32,X31) | s(bool,'F_0') = s(bool,X31)) ) | $spl74),
  inference(superposition,[],[f447,f8040])).
fof(f9624,plain,(
  ( ! [X30,X28,X29] : (sP7(X28,X29,X30) | ~p(s(bool,X30)) | ~p(s(bool,X29)) | s(bool,'F_0') = s(bool,X28)) ) | $spl74),
  inference(subsumption_resolution,[],[f9623,f275])).
fof(f9623,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | sP7(X28,X29,X30) | s(bool,'F_0') = s(bool,X28)) ) | $spl74),
  inference(forward_demodulation,[],[f9488,f877])).
fof(f9488,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | sP7(X28,X29,X30) | s(bool,'F_0') = s(bool,X28)) ) | $spl74),
  inference(superposition,[],[f444,f8040])).
fof(f9622,plain,(
  ( ! [X26,X27,X25] : (~sP6(X26,X25,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | s(bool,'F_0') = s(bool,X25)) ) | $spl74),
  inference(subsumption_resolution,[],[f9621,f275])).
fof(f9621,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP6(X26,X25,X27) | s(bool,'F_0') = s(bool,X25)) ) | $spl74),
  inference(forward_demodulation,[],[f9487,f877])).
fof(f9487,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP6(X26,X25,X27) | s(bool,'F_0') = s(bool,X25)) ) | $spl74),
  inference(superposition,[],[f438,f8040])).
fof(f9620,plain,(
  ( ! [X24,X23,X22] : (~sP3(X23,X22,X24) | p(s(bool,X24)) | ~p(s(bool,X23)) | s(bool,'F_0') = s(bool,X22)) ) | $spl74),
  inference(subsumption_resolution,[],[f9619,f275])).
fof(f9619,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP3(X23,X22,X24) | s(bool,'F_0') = s(bool,X22)) ) | $spl74),
  inference(forward_demodulation,[],[f9486,f877])).
fof(f9486,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP3(X23,X22,X24) | s(bool,'F_0') = s(bool,X22)) ) | $spl74),
  inference(superposition,[],[f419,f8040])).
fof(f9618,plain,(
  ( ! [X21,X19,X20] : (~sP4(X20,X19,X21) | p(s(bool,X21)) | ~p(s(bool,X20)) | s(bool,'F_0') = s(bool,X19)) ) | $spl74),
  inference(subsumption_resolution,[],[f9617,f275])).
fof(f9617,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP4(X20,X19,X21) | s(bool,'F_0') = s(bool,X19)) ) | $spl74),
  inference(forward_demodulation,[],[f9485,f877])).
fof(f9485,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP4(X20,X19,X21) | s(bool,'F_0') = s(bool,X19)) ) | $spl74),
  inference(superposition,[],[f415,f8040])).
fof(f9616,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | sP1(X18,X17) | s(bool,'F_0') = s(bool,X17)) ) | $spl74),
  inference(subsumption_resolution,[],[f9615,f275])).
fof(f9615,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1(X18,X17) | s(bool,'F_0') = s(bool,X17)) ) | $spl74),
  inference(forward_demodulation,[],[f9484,f877])).
fof(f9484,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | p(s(bool,X18)) | sP1(X18,X17) | s(bool,'F_0') = s(bool,X17)) ) | $spl74),
  inference(superposition,[],[f385,f8040])).
fof(f9614,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | sP1(X15,X16) | s(bool,'F_0') = s(bool,X15)) ) | $spl74),
  inference(subsumption_resolution,[],[f9613,f275])).
fof(f9613,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1(X15,X16) | s(bool,'F_0') = s(bool,X15)) ) | $spl74),
  inference(forward_demodulation,[],[f9483,f877])).
fof(f9483,plain,(
  ( ! [X15,X16] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | p(s(bool,X16)) | sP1(X15,X16) | s(bool,'F_0') = s(bool,X15)) ) | $spl74),
  inference(superposition,[],[f382,f8040])).
fof(f9612,plain,(
  ( ! [X14,X13] : (sP0(X13,X14) | s(bool,'F_0') = s(bool,X13)) ) | $spl74),
  inference(subsumption_resolution,[],[f9611,f275])).
fof(f9611,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | sP0(X13,X14) | s(bool,'F_0') = s(bool,X13)) ) | $spl74),
  inference(forward_demodulation,[],[f9482,f877])).
fof(f9482,plain,(
  ( ! [X14,X13] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP0(X13,X14) | s(bool,'F_0') = s(bool,X13)) ) | $spl74),
  inference(superposition,[],[f367,f8040])).
fof(f9610,plain,(
  ( ! [X12,X11] : (sP0(X12,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl74),
  inference(subsumption_resolution,[],[f9609,f275])).
fof(f9609,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | sP0(X12,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl74),
  inference(forward_demodulation,[],[f9481,f877])).
fof(f9481,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP0(X12,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl74),
  inference(superposition,[],[f366,f8040])).
fof(f9460,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl74),
  inference(superposition,[],[f8040,f6906])).
fof(f9459,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl74),
  inference(superposition,[],[f8040,f8040])).
fof(f2357,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),X3),s(X2,sK15(X3,X2,fun(X0,X1))))),s(X0,sK15(happ(s(fun(X2,fun(X0,X1)),X3),s(X2,sK16(X3,X2,fun(X0,X1)))),X0,X1)))) = s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),X3),s(X2,sK15(X3,X2,fun(X0,X1))))),s(X0,sK16(happ(s(fun(X2,fun(X0,X1)),X3),s(X2,sK16(X3,X2,fun(X0,X1)))),X0,X1)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),X3),s(X2,sK16(X3,X2,fun(X0,X1))))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X2,fun(X0,X1)),X3)))) )),
  inference(superposition,[],[f1067,f1067])).
fof(f2403,plain,(
  ( ! [X118,X116,X119,X117] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X116,X117),happ(s(fun(X118,fun(X116,X117)),X119),s(X118,sK15(X119,X118,fun(X116,X117)))))))) | s(X117,happ(s(fun(X116,X117),happ(s(fun(X118,fun(X116,X117)),X119),s(X118,sK16(X119,X118,fun(X116,X117))))),s(X116,sK15(happ(s(fun(X118,fun(X116,X117)),X119),s(X118,sK16(X119,X118,fun(X116,X117)))),X116,X117)))) = s(X117,happ(s(fun(X116,X117),happ(s(fun(X118,fun(X116,X117)),X119),s(X118,sK16(X119,X118,fun(X116,X117))))),s(X116,sK16(happ(s(fun(X118,fun(X116,X117)),X119),s(X118,sK16(X119,X118,fun(X116,X117)))),X116,X117)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X118,fun(X116,X117)),X119)))) )),
  inference(superposition,[],[f399,f1067])).
fof(f2402,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK15(X113,X112,fun(X110,X111)))))))) | s(X111,happ(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK16(X113,X112,fun(X110,X111))))),s(X110,X114))) != s(X111,happ(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK16(X113,X112,fun(X110,X111))))),s(X110,X115))) | s(X110,X114) = s(X110,X115) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X112,fun(X110,X111)),X113)))) )),
  inference(superposition,[],[f398,f1067])).
fof(f1018,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))))))) )),
  inference(superposition,[],[f348,f355])).
fof(f868,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(superposition,[],[f347,f355])).
fof(f2270,plain,(
  ( ! [X2,X3,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))),s('type.rat.rat',X3))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))))))))) )),
  inference(superposition,[],[f1018,f355])).
fof(f2241,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1)))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f868,f355])).
fof(f2411,plain,(
  ( ! [X138,X139,X137] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X139),s('type.rat.rat',happ(s(fun(X137,'type.rat.rat'),X138),s(X137,sK15(X138,X137,'type.rat.rat'))))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',happ(s(fun(X137,'type.rat.rat'),X138),s(X137,sK16(X138,X137,'type.rat.rat')))),s('type.rat.rat',X139))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X137,'type.rat.rat'),X138)))) )),
  inference(superposition,[],[f868,f1067])).
fof(f2415,plain,(
  ( ! [X149,X150,X148] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',happ(s(fun(X148,'type.rat.rat'),X149),s(X148,sK15(X149,X148,'type.rat.rat')))),s('type.rat.rat',X150))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X150),s('type.rat.rat',happ(s(fun(X148,'type.rat.rat'),X149),s(X148,sK16(X149,X148,'type.rat.rat'))))))))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X148,'type.rat.rat'),X149)))) )),
  inference(superposition,[],[f1018,f1067])).
fof(f2274,plain,(
  ( ! [X2,X3,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))),s('type.rat.rat',X3))))))))) )),
  inference(superposition,[],[f1018,f355])).
fof(f2237,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))),s('type.rat.rat',X3)))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))))) )),
  inference(superposition,[],[f868,f355])).
fof(f2412,plain,(
  ( ! [X142,X140,X141] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',happ(s(fun(X140,'type.rat.rat'),X141),s(X140,sK15(X141,X140,'type.rat.rat')))),s('type.rat.rat',X142)))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X142),s('type.rat.rat',happ(s(fun(X140,'type.rat.rat'),X141),s(X140,sK16(X141,X140,'type.rat.rat')))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X140,'type.rat.rat'),X141)))) )),
  inference(superposition,[],[f868,f1067])).
fof(f2414,plain,(
  ( ! [X146,X147,X145] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X147),s('type.rat.rat',happ(s(fun(X145,'type.rat.rat'),X146),s(X145,sK15(X146,X145,'type.rat.rat')))))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',happ(s(fun(X145,'type.rat.rat'),X146),s(X145,sK16(X146,X145,'type.rat.rat')))),s('type.rat.rat',X147)))))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X145,'type.rat.rat'),X146)))) )),
  inference(superposition,[],[f1018,f1067])).
fof(f2405,plain,(
  ( ! [X121,X122] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X121,'type.rat.rat'),X122),s(X121,sK15(X122,X121,'type.rat.rat'))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X121,'type.rat.rat'),X122),s(X121,sK16(X122,X121,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X121,'type.rat.rat'),X122)))) )),
  inference(superposition,[],[f347,f1067])).
fof(f2406,plain,(
  ( ! [X123,X124] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat',happ(s(fun(X123,'type.rat.rat'),X124),s(X123,sK15(X124,X123,'type.rat.rat')))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X123,'type.rat.rat'),X124),s(X123,sK16(X124,X123,'type.rat.rat'))))))))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X123,'type.rat.rat'),X124)))) )),
  inference(superposition,[],[f348,f1067])).
fof(f2364,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK15(X2,X1,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,bool),X2))) | s(bool,X3) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK16(X2,X1,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1067,f608])).
fof(f2366,plain,(
  ( ! [X8,X7,X9] : (s(bool,X9) = s(bool,happ(s(fun(X7,bool),X8),s(X7,sK15(X8,X7,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X7,bool),X8))) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,X7,bool)))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1067,f608])).
fof(f2368,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,happ(s(fun(X13,bool),X14),s(X13,sK15(X14,X13,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X13,bool),X14))) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X14),s(X13,sK16(X14,X13,bool))))) )),
  inference(superposition,[],[f1067,f283])).
fof(f2371,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,X4,bool))))) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,sK16(X5,X4,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X4,bool),X5)))) )),
  inference(superposition,[],[f295,f1067])).
fof(f2372,plain,(
  ( ! [X6,X7] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK15(X7,X6,bool))))) | s(bool,'F_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK16(X7,X6,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X6,bool),X7)))) )),
  inference(superposition,[],[f298,f1067])).
fof(f2374,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,happ(s(fun(X10,bool),X11),s(X10,sK15(X11,X10,bool))))) | p(s(bool,X12)) | s(bool,X12) = s(bool,happ(s(fun(X10,bool),X11),s(X10,sK16(X11,X10,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X10,bool),X11)))) )),
  inference(superposition,[],[f357,f1067])).
fof(f2376,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK15(X17,X16,bool))))) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,happ(s(fun(X16,bool),X17),s(X16,sK16(X17,X16,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X16,bool),X17)))) )),
  inference(superposition,[],[f360,f1067])).
fof(f2377,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X20,X19,bool))))) | sP0(X21,happ(s(fun(X19,bool),X20),s(X19,sK16(X20,X19,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X19,bool),X20)))) )),
  inference(superposition,[],[f366,f1067])).
fof(f2378,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK15(X23,X22,bool))))) | sP0(happ(s(fun(X22,bool),X23),s(X22,sK16(X23,X22,bool))),X24) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X22,bool),X23)))) )),
  inference(superposition,[],[f367,f1067])).
fof(f2379,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK15(X26,X25,bool))))) | p(s(bool,X27)) | sP1(happ(s(fun(X25,bool),X26),s(X25,sK16(X26,X25,bool))),X27) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X25,bool),X26)))) )),
  inference(superposition,[],[f382,f1067])).
fof(f2380,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,happ(s(fun(X28,bool),X29),s(X28,sK15(X29,X28,bool))))) | p(s(bool,X30)) | sP1(X30,happ(s(fun(X28,bool),X29),s(X28,sK16(X29,X28,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X28,bool),X29)))) )),
  inference(superposition,[],[f385,f1067])).
fof(f2381,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK15(X32,X31,bool))))) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,happ(s(fun(X31,bool),X32),s(X31,sK16(X32,X31,bool))),X34) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X31,bool),X32)))) )),
  inference(superposition,[],[f415,f1067])).
fof(f2382,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X35,bool))))) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,happ(s(fun(X35,bool),X36),s(X35,sK16(X36,X35,bool))),X38) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X35,bool),X36)))) )),
  inference(superposition,[],[f419,f1067])).
fof(f2383,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,sK15(X40,X39,bool))))) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,happ(s(fun(X39,bool),X40),s(X39,sK16(X40,X39,bool))),X42) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X39,bool),X40)))) )),
  inference(superposition,[],[f438,f1067])).
fof(f2384,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK15(X44,X43,bool))))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | sP7(happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool))),X45,X46) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44)))) )),
  inference(superposition,[],[f444,f1067])).
fof(f2385,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun(X47,bool),X48),s(X47,sK15(X48,X47,bool))))) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,happ(s(fun(X47,bool),X48),s(X47,sK16(X48,X47,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X47,bool),X48)))) )),
  inference(superposition,[],[f447,f1067])).
fof(f2386,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(X51,bool),X52),s(X51,sK15(X52,X51,bool))))) | p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(happ(s(fun(X51,bool),X52),s(X51,sK16(X52,X51,bool))),X53,X54) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X51,bool),X52)))) )),
  inference(superposition,[],[f455,f1067])).
fof(f2387,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(X55,bool),X56),s(X55,sK15(X56,X55,bool))))) | ~p(s(bool,X57)) | sP9(happ(s(fun(X55,bool),X56),s(X55,sK16(X56,X55,bool))),X58,X57) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X55,bool),X56)))) )),
  inference(superposition,[],[f460,f1067])).
fof(f2388,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,happ(s(fun(X59,bool),X60),s(X59,sK15(X60,X59,bool))))) | ~p(s(bool,X61)) | sP9(happ(s(fun(X59,bool),X60),s(X59,sK16(X60,X59,bool))),X61,X62) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X59,bool),X60)))) )),
  inference(superposition,[],[f461,f1067])).
fof(f2389,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,happ(s(fun(X63,bool),X64),s(X63,sK15(X64,X63,bool))))) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP9(happ(s(fun(X63,bool),X64),s(X63,sK16(X64,X63,bool))),X66,X65) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X63,bool),X64)))) )),
  inference(superposition,[],[f462,f1067])).
fof(f2390,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,happ(s(fun(X67,bool),X68),s(X67,sK15(X68,X67,bool))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,happ(s(fun(X67,bool),X68),s(X67,sK16(X68,X67,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X67,bool),X68)))) )),
  inference(superposition,[],[f472,f1067])).
fof(f2391,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X71,bool),X72),s(X71,sK15(X72,X71,bool))))) | p(s(bool,X73)) | sP11(X74,X73,happ(s(fun(X71,bool),X72),s(X71,sK16(X72,X71,bool)))) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X71,bool),X72)))) )),
  inference(superposition,[],[f476,f1067])).
fof(f2392,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,happ(s(fun(X75,bool),X76),s(X75,sK15(X76,X75,bool))))) | p(s(bool,X77)) | sP11(happ(s(fun(X75,bool),X76),s(X75,sK16(X76,X75,bool))),X78,X77) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X75,bool),X76)))) )),
  inference(superposition,[],[f477,f1067])).
fof(f2393,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,happ(s(fun(X79,bool),X80),s(X79,sK15(X80,X79,bool))))) | ~p(s(bool,X81)) | sP11(happ(s(fun(X79,bool),X80),s(X79,sK16(X80,X79,bool))),X81,X82) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X79,bool),X80)))) )),
  inference(superposition,[],[f478,f1067])).
fof(f2394,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,happ(s(fun(X83,bool),X84),s(X83,sK15(X84,X83,bool))))) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11(happ(s(fun(X83,bool),X84),s(X83,sK16(X84,X83,bool))),X86,X85) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X83,bool),X84)))) )),
  inference(superposition,[],[f479,f1067])).
fof(f2395,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,sK15(X88,X87,bool))))) | ~sP10(happ(s(fun(X87,bool),X88),s(X87,sK16(X88,X87,bool))),X89,X90) | sP11(happ(s(fun(X87,bool),X88),s(X87,sK16(X88,X87,bool))),X89,X90) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X87,bool),X88)))) )),
  inference(superposition,[],[f482,f1067])).
fof(f2396,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,happ(s(fun(X91,bool),X92),s(X91,sK15(X92,X91,bool))))) | ~sP8(happ(s(fun(X91,bool),X92),s(X91,sK16(X92,X91,bool))),X93,X94) | sP9(happ(s(fun(X91,bool),X92),s(X91,sK16(X92,X91,bool))),X93,X94) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X91,bool),X92)))) )),
  inference(superposition,[],[f483,f1067])).
fof(f2401,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,happ(s(fun(X107,bool),X108),s(X107,sK15(X108,X107,bool))))) | s(bool,X109) = s(bool,happ(s(fun(X107,bool),X108),s(X107,sK16(X108,X107,bool)))) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X107,bool),X108)))) )),
  inference(superposition,[],[f654,f1067])).
fof(f3912,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,happ(s(fun(X23,bool),X24),s(X23,sK15(X24,X23,bool)))) | s(bool,X25) = s(bool,happ(s(fun(X23,bool),X24),s(X23,sK16(X24,X23,bool)))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X23,bool),X24)))) )),
  inference(superposition,[],[f1700,f1067])).
fof(f3984,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,sK15(X40,X39,bool))))) | sP1(X41,happ(s(fun(X39,bool),X40),s(X39,sK16(X40,X39,bool)))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X39,bool),X40)))) )),
  inference(superposition,[],[f1704,f1067])).
fof(f5963,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,sK15(X40,X39,bool))))) | sP1(happ(s(fun(X39,bool),X40),s(X39,sK16(X40,X39,bool))),X41) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X39,bool),X40)))) )),
  inference(superposition,[],[f1705,f1067])).
fof(f6811,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK15(X44,X43,bool))))) | s(bool,X45) != s(bool,happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool)))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44)))) )),
  inference(superposition,[],[f1381,f1067])).
fof(f6884,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK15(X44,X43,bool))))) | s(bool,X45) = s(bool,happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool)))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44)))) )),
  inference(superposition,[],[f1701,f1067])).
fof(f298,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f179])).
fof(f1026,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK15(X0,bool,X1)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X1),X0)))) | s(bool,'T_0') = s(bool,sK16(X0,bool,X1))) )),
  inference(superposition,[],[f400,f283])).
fof(f1053,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,X2,bool))) )),
  inference(superposition,[],[f401,f283])).
fof(f1054,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK17(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK17(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f401,f283])).
fof(f1243,plain,(
  s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f522,f1020,f654])).
fof(f2172,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X8))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2131,f275])).
fof(f2131,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X8))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f298,f872])).
fof(f2362,plain,(
  ( ! [X12,X13] : (s(X13,happ(s(fun(bool,X13),X12),s(bool,'F_0'))) = s(X13,happ(s(fun(bool,X13),X12),s(bool,sK15(X12,bool,X13)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X13),X12))) | s(bool,'T_0') = s(bool,sK16(X12,bool,X13))) )),
  inference(superposition,[],[f1067,f283])).
fof(f7001,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(superposition,[],[f2172,f355])).
fof(f7004,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X6,'type.rat.rat'),X7),s(X6,sK15(X7,X6,'type.rat.rat')))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X6,'type.rat.rat'),X7),s(X6,sK16(X7,X6,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X6,'type.rat.rat'),X7)))) )),
  inference(superposition,[],[f2172,f1067])).
fof(f8301,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X6) = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f7742,f275])).
fof(f7742,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X6) = s(bool,X7)) )),
  inference(superposition,[],[f298,f6906])).
fof(f7873,plain,(
  ( ! [X318,X320,X317,X319] : (s(bool,'T_0') != s(bool,f(s(X317,sK17(X318,X319,bool,X317)))) | s(fun(X317,bool),X318) = s(fun(X317,bool),X319) | s(bool,X320) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X320) = s(bool,g(s(X317,sK17(X318,X319,bool,X317))))) )),
  inference(superposition,[],[f401,f6906])).
fof(f7616,plain,(
  ( ! [X333,X335,X332,X334] : (s(bool,X335) != s(bool,f(s(X332,sK17(X333,X334,bool,X332)))) | s(fun(X332,bool),X333) = s(fun(X332,bool),X334) | s(bool,X335) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,g(s(X332,sK17(X333,X334,bool,X332))))) )),
  inference(superposition,[],[f401,f6906])).
fof(f2341,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK17(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK17(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f2340])).
fof(f2340,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f2333,f283])).
fof(f2333,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1054,f608])).
fof(f1448,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') != s(bool,f(s(X120,sK17(X121,X122,bool,X120)))) | s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,X123) = s(bool,g(s(X120,sK17(X121,X122,bool,X120)))) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f401,f608])).
fof(f1338,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,X125) != s(bool,f(s(X122,sK17(X123,X124,bool,X122)))) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,'T_0') = s(bool,X125) | s(bool,'T_0') = s(bool,g(s(X122,sK17(X123,X124,bool,X122))))) )),
  inference(superposition,[],[f401,f608])).
fof(f1339,plain,(
  ( ! [X130,X131,X129] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X130),X129)))) | s(bool,X131) != s(bool,sK15(X129,bool,X130)) | s(bool,'T_0') = s(bool,X131) | s(bool,'T_0') = s(bool,sK16(X129,bool,X130))) )),
  inference(superposition,[],[f400,f608])).
fof(f1449,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X127),X126)))) | s(bool,'T_0') != s(bool,sK15(X126,bool,X127)) | s(bool,X128) = s(bool,sK16(X126,bool,X127)) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f400,f608])).
fof(f2213,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X1),X0)))) | s(bool,'T_0') = s(bool,sK16(X0,bool,X1)) | s(bool,'T_0') = s(bool,sK15(X0,bool,X1))) )),
  inference(condensation,[],[f2212])).
fof(f2212,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X7),X6)))) | s(bool,'T_0') = s(bool,sK16(X6,bool,X7)) | s(bool,'T_0') = s(bool,sK15(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2208,f283])).
fof(f2208,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X7),X6)))) | s(bool,'T_0') = s(bool,sK16(X6,bool,X7)) | s(bool,'T_0') = s(bool,sK15(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1026,f608])).
fof(f2358,plain,(
  ( ! [X2,X0,X1] : (s(X1,happ(s(fun(bool,X1),X0),s(bool,'T_0'))) = s(X1,happ(s(fun(bool,X1),X0),s(bool,sK15(X0,bool,X1)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X1),X0))) | s(bool,X2) = s(bool,sK16(X0,bool,X1)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1067,f608])).
fof(f2360,plain,(
  ( ! [X6,X8,X7] : (s(X7,happ(s(fun(bool,X7),X6),s(bool,X8))) = s(X7,happ(s(fun(bool,X7),X6),s(bool,sK15(X6,bool,X7)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X7),X6))) | s(bool,'T_0') = s(bool,sK16(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1067,f608])).
fof(f2413,plain,(
  ( ! [X144,X143] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X143,'type.rat.rat'),X144),s(X143,sK15(X144,X143,'type.rat.rat')))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X143,'type.rat.rat'),X144),s(X143,sK16(X144,X143,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X143,'type.rat.rat'),X144)))) )),
  inference(superposition,[],[f872,f1067])).
fof(f6004,plain,(
  ( ! [X12,X10,X11] : (sP0(X12,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK15(X11,X10,'type.rat.rat')))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK16(X11,X10,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X10,'type.rat.rat'),X11)))) )),
  inference(superposition,[],[f2174,f1067])).
fof(f6035,plain,(
  ( ! [X12,X10,X11] : (sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK15(X11,X10,'type.rat.rat'))))))),X12) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK16(X11,X10,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X10,'type.rat.rat'),X11)))) )),
  inference(superposition,[],[f2175,f1067])).
fof(f7057,plain,(
  ( ! [X12,X10,X11] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK15(X11,X10,'type.rat.rat'))))))),X12) | p(s(bool,X12)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK16(X11,X10,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X10,'type.rat.rat'),X11)))) )),
  inference(superposition,[],[f2176,f1067])).
fof(f7095,plain,(
  ( ! [X12,X10,X11] : (sP1(X12,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK15(X11,X10,'type.rat.rat')))))))) | p(s(bool,X12)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',happ(s(fun(X10,'type.rat.rat'),X11),s(X10,sK16(X11,X10,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X10,'type.rat.rat'),X11)))) )),
  inference(superposition,[],[f2177,f1067])).
fof(f8352,plain,(
  ( ! [X300,X302,X304,X301,X303,X299] : (s(X300,happ(s(fun(X299,X300),X301),s(X299,X302))) != s(X300,happ(s(fun(X299,X300),X301),s(X299,X303))) | s(X299,X302) = s(X299,X303) | s(bool,X304) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X304) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X299,X300),X301)))) )),
  inference(subsumption_resolution,[],[f7870,f275])).
fof(f7870,plain,(
  ( ! [X300,X302,X304,X301,X303,X299] : (~p(s(bool,'T_0')) | s(X300,happ(s(fun(X299,X300),X301),s(X299,X302))) != s(X300,happ(s(fun(X299,X300),X301),s(X299,X303))) | s(X299,X302) = s(X299,X303) | s(bool,X304) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X304) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X299,X300),X301)))) )),
  inference(superposition,[],[f398,f6906])).
fof(f1722,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (s(X108,happ(s(fun(X107,X108),X109),s(X107,X110))) != s(X108,happ(s(fun(X107,X108),X109),s(X107,X111))) | s(X107,X110) = s(X107,X111) | s(bool,X112) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X107,X108),X109))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f1446,f275])).
fof(f1446,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (~p(s(bool,'T_0')) | s(X108,happ(s(fun(X107,X108),X109),s(X107,X110))) != s(X108,happ(s(fun(X107,X108),X109),s(X107,X111))) | s(X107,X110) = s(X107,X111) | s(bool,X112) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X107,X108),X109))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f398,f608])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (s(X1,sK15(X2,X1,X0)) != s(X1,sK16(X2,X1,X0)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f348,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))))))) )),
  inference(cnf_transformation,[],[f211])).
fof(f2173,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X14))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2135,f275])).
fof(f2135,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | s(bool,X15) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X14))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f360,f872])).
fof(f2416,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',sK15('const.rat.rat_mul_1'(s('type.rat.rat',X0)),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',sK16('const.rat.rat_mul_1'(s('type.rat.rat',X0)),'type.rat.rat','type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))) )),
  inference(forward_demodulation,[],[f2363,f356])).
fof(f2363,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',sK16('const.rat.rat_mul_1'(s('type.rat.rat',X0)),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))),s('type.rat.rat',sK15('const.rat.rat_mul_1'(s('type.rat.rat',X0)),'type.rat.rat','type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))) )),
  inference(superposition,[],[f1067,f356])).
fof(f2119,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(superposition,[],[f872,f355])).
fof(f6001,plain,(
  ( ! [X2,X3,X1] : (sP0(X3,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))) )),
  inference(superposition,[],[f2174,f355])).
fof(f6032,plain,(
  ( ! [X2,X3,X1] : (sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1)))))),X3) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))) )),
  inference(superposition,[],[f2175,f355])).
fof(f7054,plain,(
  ( ! [X2,X3,X1] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1)))))),X3) | p(s(bool,X3)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))) )),
  inference(superposition,[],[f2176,f355])).
fof(f7092,plain,(
  ( ! [X2,X3,X1] : (sP1(X3,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))))))) | p(s(bool,X3)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))) )),
  inference(superposition,[],[f2177,f355])).
fof(f1249,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK15('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK16('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat'))))),
  inference(forward_demodulation,[],[f1248,f356])).
fof(f1248,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK16('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',sK15('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat'))))),
  inference(forward_demodulation,[],[f1093,f356])).
fof(f1093,plain,(
  s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',sK15('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',sK16('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat'))))),
  inference(unit_resulting_resolution,[],[f1020,f399])).
fof(f2557,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',X0))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) )),
  inference(subsumption_resolution,[],[f2512,f275])).
fof(f2512,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',X0))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) )),
  inference(superposition,[],[f398,f877])).
fof(f1094,plain,(
  s('type.rat.rat',sK15('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat')) != s('type.rat.rat',sK16('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat'))),
  inference(unit_resulting_resolution,[],[f1020,f400])).
fof(f7908,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f7886])).
fof(f7886,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f6906])).
fof(f7878,plain,(
  ( ! [X335,X337,X336,X338] : (s(X337,f(s(bool,'T_0'))) != s(X337,g(s(bool,'T_0'))) | s(fun(bool,X337),X335) = s(fun(bool,X337),X336) | s(bool,X338) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X338) = s(bool,sK17(X335,X336,X337,bool))) )),
  inference(superposition,[],[f401,f6906])).
fof(f8195,plain,(
  ( ! [X288,X287,X289] : (s(bool,X287) != s(bool,X288) | s(bool,X288) = s(bool,X289) | s(bool,'T_0') = s(bool,X289) | s(bool,X287) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(subsumption_resolution,[],[f7600,f522])).
fof(f7600,plain,(
  ( ! [X288,X287,X289] : (s(bool,X287) != s(bool,X288) | s(bool,X288) = s(bool,X289) | s(bool,'T_0') = s(bool,X289) | s(bool,X287) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f1700,f6906])).
fof(f8194,plain,(
  ( ! [X280,X282,X279,X281] : (s(X280,f(s(bool,X279))) != s(X280,g(s(bool,X279))) | s(fun(bool,X280),X281) = s(fun(bool,X280),X282) | s(bool,'T_0') = s(bool,sK17(X281,X282,X280,bool)) | s(bool,X279) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(subsumption_resolution,[],[f7599,f522])).
fof(f7599,plain,(
  ( ! [X280,X282,X279,X281] : (s(X280,f(s(bool,X279))) != s(X280,g(s(bool,X279))) | s(fun(bool,X280),X281) = s(fun(bool,X280),X282) | s(bool,'T_0') = s(bool,sK17(X281,X282,X280,bool)) | s(bool,X279) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f1053,f6906])).
fof(f8193,plain,(
  ( ! [X275] : (s(bool,'T_0') != s(bool,X275) | s(bool,X275) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(subsumption_resolution,[],[f7598,f522])).
fof(f7598,plain,(
  ( ! [X275] : (s(bool,'T_0') != s(bool,X275) | s(bool,X275) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f522,f6906])).
fof(f7199,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f6906,f608])).
fof(f7198,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f6906,f6906])).
fof(f7126,plain,(
  ~sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1129,f408])).
fof(f6906,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f6881,f275])).
fof(f6881,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1701,f877])).
fof(f6905,plain,(
  ( ! [X33,X34,X32] : (s(bool,X33) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X32))))) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s('type.rat.rat',X32) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f6880,f275])).
fof(f6880,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X33) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X32))))) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s('type.rat.rat',X32) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1701,f872])).
fof(f6822,plain,(
  ( ! [X37,X38] : (s(bool,X37) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f6809,f498])).
fof(f6809,plain,(
  ( ! [X37,X38] : (p(s(bool,'F_0')) | s(bool,X37) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f1381,f1243])).
fof(f5972,plain,(
  ( ! [X37,X38] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X37) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f5962,f498])).
fof(f5962,plain,(
  ( ! [X37,X38] : (p(s(bool,'F_0')) | sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X37) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f1705,f1243])).
fof(f5969,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f5955,f498])).
fof(f5955,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | sP1(X16,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1705,f283])).
fof(f3994,plain,(
  ( ! [X37,X38] : (sP1(X37,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f3983,f498])).
fof(f3983,plain,(
  ( ! [X37,X38] : (p(s(bool,'F_0')) | sP1(X37,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f1704,f1243])).
fof(f3991,plain,(
  ( ! [X17,X18,X16] : (sP1(X17,X16) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f3976,f498])).
fof(f3976,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | sP1(X17,X16) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1704,f283])).
fof(f3896,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1700,f283])).
fof(f3808,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(trivial_inequality_removal,[],[f3802])).
fof(f3802,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1095,f608])).
fof(f3680,plain,(
  ( ! [X0] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f1222,f2579])).
fof(f3451,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1134,f408])).
fof(f3412,plain,(
  ~sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f1128,f408])).
fof(f3384,plain,(
  ~sP5('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1122,f408])).
fof(f3355,plain,(
  sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f498,f912,f506,f1121,f414])).
fof(f3356,plain,(
  sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f512,f506,f1121,f412])).
fof(f3357,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f512,f1116,f1121,f412])).
fof(f3327,plain,(
  sP5('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f511,f506,f1120,f414])).
fof(f3328,plain,(
  sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f911,f506,f1120,f414])).
fof(f3311,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1117,f409])).
fof(f3294,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f498,f912,f509,f1116,f413])).
fof(f3264,plain,(
  sP5('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f510,f509,f1115,f413])).
fof(f3265,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f910,f509,f1115,f413])).
fof(f3237,plain,(
  ~sP5('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')),
  inference(unit_resulting_resolution,[],[f1112,f410])).
fof(f2645,plain,(
  ~sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f906,f409])).
fof(f2639,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f897,f410])).
fof(f2580,plain,(
  ( ! [X39,X40] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X39,X40) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X39,X40)) )),
  inference(subsumption_resolution,[],[f2541,f275])).
fof(f2541,plain,(
  ( ! [X39,X40] : (~p(s(bool,'T_0')) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X39,X40) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X39,X40)) )),
  inference(superposition,[],[f483,f877])).
fof(f2579,plain,(
  ( ! [X37,X38] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X37,X38) | ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X37,X38)) )),
  inference(subsumption_resolution,[],[f2540,f275])).
fof(f2540,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X37,X38) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X37,X38)) )),
  inference(superposition,[],[f482,f877])).
fof(f2578,plain,(
  ( ! [X35,X36] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X36,X35) | p(s(bool,X36)) | ~p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f2539,f275])).
fof(f2539,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X36,X35)) )),
  inference(superposition,[],[f479,f877])).
fof(f2577,plain,(
  ( ! [X33,X34] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X33,X34) | ~p(s(bool,X33))) )),
  inference(subsumption_resolution,[],[f2538,f275])).
fof(f2538,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X33,X34)) )),
  inference(superposition,[],[f478,f877])).
fof(f2576,plain,(
  ( ! [X31,X32] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X32,X31) | p(s(bool,X31))) )),
  inference(subsumption_resolution,[],[f2537,f275])).
fof(f2537,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X32,X31)) )),
  inference(superposition,[],[f477,f877])).
fof(f2575,plain,(
  ( ! [X30,X29] : (sP11(X30,X29,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | p(s(bool,X29)) | p(s(bool,X30))) )),
  inference(subsumption_resolution,[],[f2536,f275])).
fof(f2536,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP11(X30,X29,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | p(s(bool,X30))) )),
  inference(superposition,[],[f476,f877])).
fof(f2574,plain,(
  ( ! [X28,X27] : (~sP10(X28,X27,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | ~p(s(bool,X28)) | p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f2535,f275])).
fof(f2535,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~sP10(X28,X27,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(superposition,[],[f472,f877])).
fof(f2573,plain,(
  ( ! [X26,X25] : (~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X26,X25) | p(s(bool,X26)) | p(s(bool,X25))) )),
  inference(subsumption_resolution,[],[f2534,f275])).
fof(f2534,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X26,X25)) )),
  inference(superposition,[],[f462,f877])).
fof(f2572,plain,(
  ( ! [X24,X23] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X23,X24) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f2533,f275])).
fof(f2533,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X23,X24)) )),
  inference(superposition,[],[f461,f877])).
fof(f2571,plain,(
  ( ! [X21,X22] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X22,X21) | ~p(s(bool,X21))) )),
  inference(subsumption_resolution,[],[f2532,f275])).
fof(f2532,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X22,X21)) )),
  inference(superposition,[],[f460,f877])).
fof(f2570,plain,(
  ( ! [X19,X20] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X19,X20) | p(s(bool,X20)) | p(s(bool,X19))) )),
  inference(subsumption_resolution,[],[f2531,f275])).
fof(f2531,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,X19)) | p(s(bool,X20)) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X19,X20)) )),
  inference(superposition,[],[f455,f877])).
fof(f2569,plain,(
  ( ! [X17,X18] : (~sP7(X18,X17,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | p(s(bool,X18)) | ~p(s(bool,X17))) )),
  inference(subsumption_resolution,[],[f2530,f275])).
fof(f2530,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP7(X18,X17,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(superposition,[],[f447,f877])).
fof(f2568,plain,(
  ( ! [X15,X16] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X15,X16) | ~p(s(bool,X16)) | ~p(s(bool,X15))) )),
  inference(subsumption_resolution,[],[f2529,f275])).
fof(f2529,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X15,X16)) )),
  inference(superposition,[],[f444,f877])).
fof(f2567,plain,(
  ( ! [X14,X13] : (~sP6(X13,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X14) | p(s(bool,X14)) | ~p(s(bool,X13))) )),
  inference(subsumption_resolution,[],[f2528,f275])).
fof(f2528,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | p(s(bool,X14)) | ~sP6(X13,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X14)) )),
  inference(superposition,[],[f438,f877])).
fof(f2566,plain,(
  ( ! [X12,X11] : (~sP3(X11,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X12) | p(s(bool,X12)) | ~p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f2527,f275])).
fof(f2527,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | p(s(bool,X12)) | ~sP3(X11,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X12)) )),
  inference(superposition,[],[f419,f877])).
fof(f2565,plain,(
  ( ! [X10,X9] : (~sP4(X9,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X10) | p(s(bool,X10)) | ~p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f2526,f275])).
fof(f2526,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP4(X9,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X10)) )),
  inference(superposition,[],[f415,f877])).
fof(f2564,plain,(
  ( ! [X8] : (sP1(X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f2525,f275])).
fof(f2525,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | p(s(bool,X8)) | sP1(X8,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(superposition,[],[f385,f877])).
fof(f2563,plain,(
  ( ! [X7] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X7) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f2524,f275])).
fof(f2524,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X7)) )),
  inference(superposition,[],[f382,f877])).
fof(f2193,plain,(
  ( ! [X70,X71,X69] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X69)))),X70,X71) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X69)))),X70,X71) | s('type.rat.rat',X69) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2155,f275])).
fof(f2155,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X69)))),X70,X71) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X69)))),X70,X71) | s('type.rat.rat',X69) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f483,f872])).
fof(f2192,plain,(
  ( ! [X68,X66,X67] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X66)))),X67,X68) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X66)))),X67,X68) | s('type.rat.rat',X66) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2154,f275])).
fof(f2154,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X66)))),X67,X68) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X66)))),X67,X68) | s('type.rat.rat',X66) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f482,f872])).
fof(f2191,plain,(
  ( ! [X64,X65,X63] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X63)))),X65,X64) | p(s(bool,X65)) | ~p(s(bool,X64)) | s('type.rat.rat',X63) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2153,f275])).
fof(f2153,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | p(s(bool,X65)) | ~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X63)))),X65,X64) | s('type.rat.rat',X63) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f479,f872])).
fof(f2190,plain,(
  ( ! [X61,X62,X60] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X60)))),X61,X62) | ~p(s(bool,X61)) | s('type.rat.rat',X60) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2152,f275])).
fof(f2152,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X60)))),X61,X62) | s('type.rat.rat',X60) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f478,f872])).
fof(f2189,plain,(
  ( ! [X59,X57,X58] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X57)))),X59,X58) | p(s(bool,X58)) | s('type.rat.rat',X57) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2151,f275])).
fof(f2151,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X57)))),X59,X58) | s('type.rat.rat',X57) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f477,f872])).
fof(f2188,plain,(
  ( ! [X54,X56,X55] : (sP11(X56,X55,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X54))))) | p(s(bool,X55)) | p(s(bool,X56)) | s('type.rat.rat',X54) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2150,f275])).
fof(f2150,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | sP11(X56,X55,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X54))))) | p(s(bool,X56)) | s('type.rat.rat',X54) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f476,f872])).
fof(f2187,plain,(
  ( ! [X52,X53,X51] : (~sP10(X53,X52,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X51))))) | ~p(s(bool,X53)) | p(s(bool,X52)) | s('type.rat.rat',X51) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2149,f275])).
fof(f2149,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | ~p(s(bool,X53)) | ~sP10(X53,X52,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X51))))) | s('type.rat.rat',X51) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f472,f872])).
fof(f2186,plain,(
  ( ! [X50,X48,X49] : (~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X48)))),X50,X49) | p(s(bool,X50)) | p(s(bool,X49)) | s('type.rat.rat',X48) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2148,f275])).
fof(f2148,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X48)))),X50,X49) | s('type.rat.rat',X48) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f462,f872])).
fof(f2185,plain,(
  ( ! [X47,X45,X46] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X45)))),X46,X47) | ~p(s(bool,X46)) | s('type.rat.rat',X45) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2147,f275])).
fof(f2147,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X45)))),X46,X47) | s('type.rat.rat',X45) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f461,f872])).
fof(f2184,plain,(
  ( ! [X43,X44,X42] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X42)))),X44,X43) | ~p(s(bool,X43)) | s('type.rat.rat',X42) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2146,f275])).
fof(f2146,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X42)))),X44,X43) | s('type.rat.rat',X42) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f460,f872])).
fof(f2183,plain,(
  ( ! [X39,X41,X40] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X39)))),X40,X41) | p(s(bool,X41)) | p(s(bool,X40)) | s('type.rat.rat',X39) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2145,f275])).
fof(f2145,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,X40)) | p(s(bool,X41)) | ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X39)))),X40,X41) | s('type.rat.rat',X39) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f455,f872])).
fof(f2182,plain,(
  ( ! [X37,X38,X36] : (~sP7(X38,X37,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X36))))) | p(s(bool,X38)) | ~p(s(bool,X37)) | s('type.rat.rat',X36) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2144,f275])).
fof(f2144,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP7(X38,X37,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X36))))) | s('type.rat.rat',X36) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f447,f872])).
fof(f2181,plain,(
  ( ! [X35,X33,X34] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X33)))),X34,X35) | ~p(s(bool,X35)) | ~p(s(bool,X34)) | s('type.rat.rat',X33) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2143,f275])).
fof(f2143,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X33)))),X34,X35) | s('type.rat.rat',X33) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f444,f872])).
fof(f2180,plain,(
  ( ! [X30,X31,X32] : (~sP6(X31,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X30)))),X32) | p(s(bool,X32)) | ~p(s(bool,X31)) | s('type.rat.rat',X30) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2142,f275])).
fof(f2142,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP6(X31,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X30)))),X32) | s('type.rat.rat',X30) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f438,f872])).
fof(f2179,plain,(
  ( ! [X28,X29,X27] : (~sP3(X28,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X27)))),X29) | p(s(bool,X29)) | ~p(s(bool,X28)) | s('type.rat.rat',X27) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2141,f275])).
fof(f2141,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~sP3(X28,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X27)))),X29) | s('type.rat.rat',X27) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f419,f872])).
fof(f2178,plain,(
  ( ! [X26,X24,X25] : (~sP4(X25,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X24)))),X26) | p(s(bool,X26)) | ~p(s(bool,X25)) | s('type.rat.rat',X24) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2140,f275])).
fof(f2140,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP4(X25,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X24)))),X26) | s('type.rat.rat',X24) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f415,f872])).
fof(f2177,plain,(
  ( ! [X23,X22] : (sP1(X23,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X22))))) | p(s(bool,X23)) | s('type.rat.rat',X22) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2139,f275])).
fof(f2139,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,X23)) | sP1(X23,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X22))))) | s('type.rat.rat',X22) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f385,f872])).
fof(f2176,plain,(
  ( ! [X21,X20] : (sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X20)))),X21) | p(s(bool,X21)) | s('type.rat.rat',X20) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2138,f275])).
fof(f2138,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | p(s(bool,X21)) | sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X20)))),X21) | s('type.rat.rat',X20) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f382,f872])).
fof(f2175,plain,(
  ( ! [X19,X18] : (sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X18)))),X19) | s('type.rat.rat',X18) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2137,f275])).
fof(f2137,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X18)))),X19) | s('type.rat.rat',X18) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f367,f872])).
fof(f2174,plain,(
  ( ! [X17,X16] : (sP0(X17,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X16))))) | s('type.rat.rat',X16) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2136,f275])).
fof(f2136,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | sP0(X17,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X16))))) | s('type.rat.rat',X16) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f366,f872])).
fof(f1995,plain,(
  sP5('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f498,f506,f509,f912,f413])).
fof(f1975,plain,(
  sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f498,f506,f508,f911,f414])).
fof(f1450,plain,(
  ( ! [X132,X130,X131,X129] : (s(X131,f(s(bool,'T_0'))) != s(X131,g(s(bool,'T_0'))) | s(fun(bool,X131),X129) = s(fun(bool,X131),X130) | s(bool,X132) = s(bool,sK17(X129,X130,X131,bool)) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f401,f608])).
fof(f1721,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8(X91,X92,X93) | sP9(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1438,f275])).
fof(f1438,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~sP8(X91,X92,X93) | sP9(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f483,f608])).
fof(f1720,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10(X87,X88,X89) | sP11(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1437,f275])).
fof(f1437,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP10(X87,X88,X89) | sP11(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f482,f608])).
fof(f1719,plain,(
  ( ! [X85,X83,X86,X84] : (~sP11(X83,X85,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1436,f275])).
fof(f1436,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11(X83,X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f479,f608])).
fof(f1718,plain,(
  ( ! [X80,X81,X79,X82] : (sP11(X79,X80,X81) | ~p(s(bool,X80)) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1435,f275])).
fof(f1435,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | sP11(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f478,f608])).
fof(f1717,plain,(
  ( ! [X78,X76,X77,X75] : (sP11(X75,X77,X76) | p(s(bool,X76)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1434,f275])).
fof(f1434,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f477,f608])).
fof(f1716,plain,(
  ( ! [X74,X72,X71,X73] : (sP11(X73,X72,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1433,f275])).
fof(f1433,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11(X73,X72,X71) | p(s(bool,X73)) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f476,f608])).
fof(f1715,plain,(
  ( ! [X70,X68,X69,X67] : (~sP10(X69,X68,X67) | ~p(s(bool,X69)) | p(s(bool,X68)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1432,f275])).
fof(f1432,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f472,f608])).
fof(f1714,plain,(
  ( ! [X66,X64,X65,X63] : (~sP9(X63,X65,X64) | p(s(bool,X65)) | p(s(bool,X64)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1431,f275])).
fof(f1431,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f462,f608])).
fof(f1713,plain,(
  ( ! [X61,X59,X62,X60] : (sP9(X59,X60,X61) | ~p(s(bool,X60)) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1430,f275])).
fof(f1430,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f461,f608])).
fof(f1712,plain,(
  ( ! [X57,X58,X56,X55] : (sP9(X55,X57,X56) | ~p(s(bool,X56)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1429,f275])).
fof(f1429,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f460,f608])).
fof(f1711,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8(X51,X52,X53) | p(s(bool,X53)) | p(s(bool,X52)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1428,f275])).
fof(f1428,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f455,f608])).
fof(f1710,plain,(
  ( ! [X47,X50,X48,X49] : (~sP7(X49,X48,X47) | p(s(bool,X49)) | ~p(s(bool,X48)) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1427,f275])).
fof(f1427,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f447,f608])).
fof(f1709,plain,(
  ( ! [X45,X43,X46,X44] : (sP7(X43,X44,X45) | ~p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1426,f275])).
fof(f1426,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | sP7(X43,X44,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f444,f608])).
fof(f1708,plain,(
  ( ! [X39,X41,X42,X40] : (~sP6(X40,X39,X41) | p(s(bool,X41)) | ~p(s(bool,X40)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1425,f275])).
fof(f1425,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP6(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f438,f608])).
fof(f1707,plain,(
  ( ! [X37,X35,X38,X36] : (~sP3(X36,X35,X37) | p(s(bool,X37)) | ~p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1424,f275])).
fof(f1424,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP3(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f419,f608])).
fof(f1706,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X32,X31,X33) | p(s(bool,X33)) | ~p(s(bool,X32)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1423,f275])).
fof(f1423,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP4(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f415,f608])).
fof(f1705,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1422,f275])).
fof(f1422,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f385,f608])).
fof(f1704,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f1421,f275])).
fof(f1421,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f382,f608])).
fof(f1703,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1420,f275])).
fof(f1420,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f367,f608])).
fof(f1702,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1419,f275])).
fof(f1419,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f366,f608])).
fof(f1701,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1418,f275])).
fof(f1418,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f360,f608])).
fof(f1700,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1414,f275])).
fof(f1414,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f298,f608])).
fof(f1381,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f294,f608])).
fof(f1340,plain,(
  ( ! [X134,X132,X135,X133] : (s(X134,f(s(bool,X135))) != s(X134,g(s(bool,X135))) | s(fun(bool,X134),X132) = s(fun(bool,X134),X133) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,sK17(X132,X133,X134,bool))) )),
  inference(superposition,[],[f401,f608])).
fof(f1257,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f505,f509,f910,f413])).
fof(f1251,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f275,f512,f504,f909,f412])).
fof(f1252,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f512,f900,f909,f412])).
fof(f1095,plain,(
  s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1020,f294])).
fof(f1101,plain,(
  ~sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f498,f1020,f380])).
fof(f1103,plain,(
  ~sP1('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f1020,f380])).
fof(f1104,plain,(
  ~sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1020,f1020,f380])).
fof(f1105,plain,(
  sP1('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f275,f1020,f382])).
fof(f1106,plain,(
  sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f1020,f385])).
fof(f1107,plain,(
  ~sP5('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f275,f275,f1020,f411])).
fof(f1109,plain,(
  sP5('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f510,f507,f505,f1020,f413])).
fof(f1112,plain,(
  ~sP4('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f275,f275,f1020,f415])).
fof(f1115,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X1)) )),
  inference(unit_resulting_resolution,[],[f1020,f417])).
fof(f1116,plain,(
  ( ! [X0,X1] : (sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1020,f418])).
fof(f1117,plain,(
  ~sP3('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f275,f275,f1020,f419])).
fof(f1120,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X1)) )),
  inference(unit_resulting_resolution,[],[f1020,f421])).
fof(f1121,plain,(
  ( ! [X0,X1] : (sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1020,f422])).
fof(f1122,plain,(
  ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f1020,f423])).
fof(f1128,plain,(
  ~sP2('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f1020,f423])).
fof(f1129,plain,(
  ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f1020,f498,f1020,f423])).
fof(f1134,plain,(
  ~sP2('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f498,f1020,f423])).
fof(f1135,plain,(
  ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1020,f498,f1020,f423])).
fof(f1136,plain,(
  ~sP2('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f1020,f1020,f423])).
fof(f1137,plain,(
  ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1020,f1020,f1020,f423])).
fof(f1140,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f522,f1020,f427])).
fof(f1141,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f522,f1020,f427])).
fof(f1142,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f275,f1020,f432])).
fof(f1144,plain,(
  ( ! [X0] : (~sP7('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0)) )),
  inference(unit_resulting_resolution,[],[f275,f1020,f433])).
fof(f1146,plain,(
  ~sP6('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f275,f275,f1020,f438])).
fof(f1149,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X1)) )),
  inference(unit_resulting_resolution,[],[f1020,f440])).
fof(f1150,plain,(
  ( ! [X0,X1] : (sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1020,f441])).
fof(f1151,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f498,f1020,f442])).
fof(f1155,plain,(
  ( ! [X0] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f1020,f1020,f442])).
fof(f1156,plain,(
  ( ! [X0] : (sP7('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0)) )),
  inference(unit_resulting_resolution,[],[f498,f1020,f443])).
fof(f1160,plain,(
  ( ! [X0] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0)) )),
  inference(unit_resulting_resolution,[],[f1020,f1020,f443])).
fof(f1161,plain,(
  ~sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f275,f275,f1020,f447])).
fof(f1162,plain,(
  ( ! [X0] : (~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f1020,f448])).
fof(f1163,plain,(
  ( ! [X0] : (~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f275,f1020,f449])).
fof(f1178,plain,(
  ~sP8('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f1020,f455])).
fof(f1182,plain,(
  ~sP8('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f275,f1020,f455])).
fof(f1183,plain,(
  ~sP8('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1020,f275,f1020,f455])).
fof(f1186,plain,(
  ( ! [X0,X1] : (sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1020,f458])).
fof(f1187,plain,(
  sP9('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f1020,f459])).
fof(f1192,plain,(
  sP9('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f498,f1020,f459])).
fof(f1194,plain,(
  sP9('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f1020,f1020,f459])).
fof(f1197,plain,(
  sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f1020,f459])).
fof(f1198,plain,(
  sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f1020,f498,f1020,f459])).
fof(f1199,plain,(
  sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f1020,f1020,f459])).
fof(f1200,plain,(
  sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1020,f1020,f1020,f459])).
fof(f1201,plain,(
  ~sP9('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f498,f275,f1020,f462])).
fof(f1204,plain,(
  ~sP9('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f275,f498,f1020,f462])).
fof(f1205,plain,(
  ~sP9('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f275,f1020,f1020,f462])).
fof(f1207,plain,(
  ( ! [X0] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f498,f1020,f465])).
fof(f1209,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f498,f1020,f465])).
fof(f1210,plain,(
  ( ! [X0] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f1020,f1020,f465])).
fof(f1212,plain,(
  ( ! [X0] : (~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f275,f1020,f466])).
fof(f1219,plain,(
  ~sP10('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f275,f1020,f472])).
fof(f1222,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f1020,f473])).
fof(f1223,plain,(
  ( ! [X0,X1] : (sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1020,f475])).
fof(f1224,plain,(
  sP11('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')),
  inference(unit_resulting_resolution,[],[f498,f275,f1020,f476])).
fof(f1227,plain,(
  sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f1020,f476])).
fof(f1228,plain,(
  sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f1020,f1020,f476])).
fof(f1229,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f275,f1020,f477])).
fof(f1231,plain,(
  ~sP11('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f275,f1020,f479])).
fof(f1233,plain,(
  ( ! [X0] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f514,f1020,f484])).
fof(f1234,plain,(
  ( ! [X0] : (sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f515,f1020,f484])).
fof(f1235,plain,(
  ~sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0')),
  inference(unit_resulting_resolution,[],[f498,f1020,f499])).
fof(f1239,plain,(
  ~sP0('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f498,f1020,f499])).
fof(f1240,plain,(
  ~sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1020,f1020,f499])).
fof(f1090,plain,(
  ~sP5('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f905,f409])).
fof(f1085,plain,(
  ~sP5('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0')),
  inference(unit_resulting_resolution,[],[f901,f409])).
fof(f1076,plain,(
  sP5('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f512,f507,f900,f412])).
fof(f1071,plain,(
  ~sP5('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f896,f410])).
fof(f1020,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(equality_resolution,[],[f348])).
fof(f879,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f866,f366])).
fof(f880,plain,(
  ( ! [X0] : (sP0('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) )),
  inference(unit_resulting_resolution,[],[f866,f367])).
fof(f881,plain,(
  ~sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f866,f381])).
fof(f883,plain,(
  ~sP1('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f275,f866,f381])).
fof(f884,plain,(
  ~sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f866,f866,f381])).
fof(f885,plain,(
  sP1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f498,f866,f382])).
fof(f886,plain,(
  sP1('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f498,f866,f385])).
fof(f889,plain,(
  ~sP5('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f411])).
fof(f890,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f866,f498,f866,f411])).
fof(f891,plain,(
  sP5('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f511,f507,f504,f866,f412])).
fof(f892,plain,(
  ~sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f415])).
fof(f896,plain,(
  ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f415])).
fof(f897,plain,(
  ~sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f866,f498,f866,f415])).
fof(f900,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f866,f416])).
fof(f901,plain,(
  ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f419])).
fof(f905,plain,(
  ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f419])).
fof(f906,plain,(
  ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f866,f498,f866,f419])).
fof(f909,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f866,f420])).
fof(f910,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f866,f424])).
fof(f911,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X1)) )),
  inference(unit_resulting_resolution,[],[f866,f425])).
fof(f912,plain,(
  ( ! [X0,X1] : (sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f866,f426])).
fof(f913,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f522,f866,f429])).
fof(f914,plain,(
  ~sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f522,f866,f429])).
fof(f915,plain,(
  ( ! [X0] : (~sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f498,f866,f432])).
fof(f916,plain,(
  ( ! [X0] : (~sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f498,f866,f433])).
fof(f923,plain,(
  ~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f438])).
fof(f927,plain,(
  ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f438])).
fof(f928,plain,(
  ~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0')),
  inference(unit_resulting_resolution,[],[f866,f498,f866,f438])).
fof(f931,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f866,f439])).
fof(f932,plain,(
  sP7('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f275,f866,f444])).
fof(f937,plain,(
  sP7('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f275,f275,f866,f444])).
fof(f938,plain,(
  sP7('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f866,f275,f866,f444])).
fof(f942,plain,(
  sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f275,f275,f866,f444])).
fof(f943,plain,(
  sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0')),
  inference(unit_resulting_resolution,[],[f866,f275,f866,f444])).
fof(f944,plain,(
  sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f275,f866,f866,f444])).
fof(f945,plain,(
  sP7('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f866,f866,f866,f444])).
fof(f946,plain,(
  ~sP7('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f447])).
fof(f949,plain,(
  ~sP7('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f447])).
fof(f950,plain,(
  ~sP7('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f866,f498,f866,f447])).
fof(f952,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f498,f866,f448])).
fof(f954,plain,(
  ( ! [X0] : (~sP9('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) )),
  inference(unit_resulting_resolution,[],[f498,f866,f449])).
fof(f958,plain,(
  ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f866,f455])).
fof(f961,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f866,f456])).
fof(f962,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X1)) )),
  inference(unit_resulting_resolution,[],[f866,f457])).
fof(f963,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f275,f866,f460])).
fof(f967,plain,(
  ( ! [X0] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(unit_resulting_resolution,[],[f866,f866,f460])).
fof(f968,plain,(
  ( ! [X0] : (sP9('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) )),
  inference(unit_resulting_resolution,[],[f275,f866,f461])).
fof(f972,plain,(
  ( ! [X0] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) )),
  inference(unit_resulting_resolution,[],[f866,f866,f461])).
fof(f973,plain,(
  ~sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f866,f462])).
fof(f974,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) )),
  inference(unit_resulting_resolution,[],[f498,f866,f466])).
fof(f986,plain,(
  ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f498,f275,f866,f472])).
fof(f990,plain,(
  ~sP10('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f498,f275,f866,f472])).
fof(f991,plain,(
  ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f498,f866,f866,f472])).
fof(f994,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X1)) )),
  inference(unit_resulting_resolution,[],[f866,f474])).
fof(f995,plain,(
  sP11('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f498,f498,f866,f476])).
fof(f998,plain,(
  ( ! [X0] : (sP11('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) )),
  inference(unit_resulting_resolution,[],[f275,f866,f478])).
fof(f1002,plain,(
  ( ! [X0] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) )),
  inference(unit_resulting_resolution,[],[f866,f866,f478])).
fof(f1003,plain,(
  ~sP11('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f479])).
fof(f1006,plain,(
  ~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f866,f479])).
fof(f1007,plain,(
  ~sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))),
  inference(unit_resulting_resolution,[],[f866,f498,f866,f479])).
fof(f1008,plain,(
  ( ! [X0] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f519,f866,f482])).
fof(f1009,plain,(
  ( ! [X0] : (sP11('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f520,f866,f482])).
fof(f1010,plain,(
  ( ! [X0] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f516,f866,f483])).
fof(f1011,plain,(
  ( ! [X0] : (sP9('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f517,f866,f483])).
fof(f847,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f510,f507,f505,f498,f413])).
fof(f833,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f512,f507,f504,f275,f412])).
fof(f799,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f793,f467])).
fof(f793,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f275,f498,f275,f472])).
fof(f786,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f521,f275,f498,f471])).
fof(f763,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f757,f450])).
fof(f757,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f275,f455])).
fof(f750,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f518,f498,f498,f454])).
fof(f727,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f721,f431])).
fof(f721,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f275,f275,f438])).
fof(f713,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f513,f275,f275,f434])).
fof(f709,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f704,f408])).
fof(f704,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f498,f423])).
fof(f701,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f691,f409])).
fof(f691,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f275,f275,f419])).
fof(f685,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f275,f275,f415])).
fof(f671,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f275,f522,f429])).
fof(f673,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f275,f522,f429])).
fof(f662,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f498,f522,f427])).
fof(f663,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f522,f427])).
fof(f654,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f651,f498])).
fof(f651,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f357,f283])).
fof(f584,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f275,f275,f478])).
fof(f577,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f498,f275,f477])).
fof(f575,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f275,f498,f466])).
fof(f574,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f498,f498,f465])).
fof(f567,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f275,f275,f461])).
fof(f564,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f275,f460])).
fof(f561,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f275,f498,f449])).
fof(f557,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f498,f448])).
fof(f556,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f498,f498,f443])).
fof(f550,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f498,f498,f442])).
fof(f547,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f275,f498,f433])).
fof(f546,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f275,f498,f432])).
fof(f543,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f499])).
fof(f536,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f498,f275,f385])).
fof(f533,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f275,f382])).
fof(f530,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f275,f275,f381])).
fof(f529,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f498,f498,f380])).
fof(f522,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f498,f294])).
fof(f521,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f498,f475])).
fof(f520,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f275,f474])).
fof(f519,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f498,f473])).
fof(f518,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f498,f458])).
fof(f517,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f275,f457])).
fof(f516,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f456])).
fof(f515,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f498,f441])).
fof(f514,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f498,f440])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f439])).
fof(f512,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f275,f426])).
fof(f511,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f275,f425])).
fof(f510,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f424])).
fof(f509,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f498,f422])).
fof(f508,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f498,f421])).
fof(f507,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f420])).
fof(f506,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f498,f418])).
fof(f505,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f498,f417])).
fof(f504,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f416])).
fof(f503,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f275,f367])).
fof(f502,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f275,f366])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f153,f170,f169])).
fof(f169,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f170,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X21,X22,X23] : ((p(s(bool,X23)) <=> (p(s(bool,X22)) => p(s(bool,X21)))) <=> ((p(s(bool,X23)) | p(s(bool,X22))) & (p(s(bool,X23)) | ~p(s(bool,X21))) & (~p(s(bool,X22)) | p(s(bool,X21)) | ~p(s(bool,X23)))))),
  file('thm.rat.RAT_EQ_RMUL','thm.sat.sat_27_prove_ltb')).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f268])).
fof(f268,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f169])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f269])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f269])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f265])).
fof(f265,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f170])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f266])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f266])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f266])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f266])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f266])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f266])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f126,f167,f166])).
fof(f166,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f167,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X21,X22,X23] : ((p(s(bool,X23)) <=> (p(s(bool,X22)) | p(s(bool,X21)))) <=> ((p(s(bool,X23)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | p(s(bool,X21)) | ~p(s(bool,X23)))))),
  file('thm.rat.RAT_EQ_RMUL','thm.sat.sat_25_prove_ltb')).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f260])).
fof(f260,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f261])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f257])).
fof(f257,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f167])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f258])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f258])).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f451])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f258])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f258])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f124,f164,f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f164,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f124,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X21,X22,X23] : ((p(s(bool,X23)) <=> (p(s(bool,X22)) & p(s(bool,X21)))) <=> ((p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X21)) | ~p(s(bool,X23)))))),
  file('thm.rat.RAT_EQ_RMUL','thm.sat.sat_23_prove_ltb')).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f255])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f163])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f253])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f253])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f253])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f250])).
fof(f250,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f164])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f251])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f251])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f251])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f251])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f122,f161,f160,f159,f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f159,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f160,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f161,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X21,X22,X23] : ((p(s(bool,X23)) <=> s(bool,X21) = s(bool,X22)) <=> ((p(s(bool,X23)) | p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X23)) | ~p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23))) & (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)))))),
  file('thm.rat.RAT_EQ_RMUL','thm.sat.sat_2|49|_prove_ltb')).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f248])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f158])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f247])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f247])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f159])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f245])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f160])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f242])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f242])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f161])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f239])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK17(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK17(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK17(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK17(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_EQ_RMUL',aHL_EXT)).
fof(f380,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f381,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f382,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f385,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f499,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f365,f498])).
fof(f365,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f366,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f367,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f357,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f144])).
fof(f295,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f177])).
fof(f19880,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737,f19738,f19739,f19820,f19873,f19875,f19877,f19879])).
fof(f19878,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737,f19738,f19739,f19820,f19873,f19875,f19877])).
fof(f19876,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737,f19738,f19739,f19820,f19873,f19875])).
fof(f19874,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737,f19738,f19739,f19820,f19873])).
fof(f19872,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737,f19738,f19739,f19820])).
fof(f19871,plain,(
  $false | ($spl47 | $spl48 | $spl71 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19742,f6731])).
fof(f6731,plain,(
  sP3('T_0',g(s(bool,'T_0')),'F_0') | $spl71),
  inference(unit_resulting_resolution,[],[f6630,f409])).
fof(f6630,plain,(
  sP5('F_0','T_0',g(s(bool,'T_0'))) | $spl71),
  inference(unit_resulting_resolution,[],[f498,f6477,f427])).
fof(f19742,plain,(
  ~sP3('T_0',g(s(bool,'T_0')),'F_0') | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f14157])).
fof(f14157,plain,(
  ~sP3('T_0',g(s(bool,'F_0')),'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f498,f511,f506,f13274,f414])).
fof(f13274,plain,(
  ~sP5('F_0','T_0',g(s(bool,'F_0'))) | $spl48),
  inference(unit_resulting_resolution,[],[f498,f2312,f430])).
fof(f19870,plain,(
  $false | ($spl47 | $spl48 | $spl71 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19741,f6717])).
fof(f6717,plain,(
  ~sP4('T_0','T_0',g(s(bool,'T_0'))) | $spl71),
  inference(unit_resulting_resolution,[],[f275,f512,f507,f6582,f412])).
fof(f6582,plain,(
  ~sP5('T_0',g(s(bool,'T_0')),'T_0') | $spl71),
  inference(unit_resulting_resolution,[],[f275,f6477,f429])).
fof(f19741,plain,(
  sP4('T_0','T_0',g(s(bool,'T_0'))) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f14118])).
fof(f14118,plain,(
  sP4('T_0','T_0',g(s(bool,'F_0'))) | $spl48),
  inference(unit_resulting_resolution,[],[f13273,f410])).
fof(f13273,plain,(
  sP5('T_0',g(s(bool,'F_0')),'T_0') | $spl48),
  inference(unit_resulting_resolution,[],[f275,f2312,f428])).
fof(f19869,plain,(
  $false | ($spl47 | $spl48 | $spl71 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19740,f6707])).
fof(f6707,plain,(
  sP3(g(s(bool,'T_0')),'T_0','F_0') | $spl71),
  inference(unit_resulting_resolution,[],[f6581,f409])).
fof(f6581,plain,(
  sP5('F_0',g(s(bool,'T_0')),'T_0') | $spl71),
  inference(unit_resulting_resolution,[],[f498,f6477,f427])).
fof(f19740,plain,(
  ~sP3(g(s(bool,'T_0')),'T_0','F_0') | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f14102])).
fof(f14102,plain,(
  ~sP3(g(s(bool,'F_0')),'T_0','F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f275,f512,f506,f13272,f412])).
fof(f13272,plain,(
  ~sP5('F_0',g(s(bool,'F_0')),'T_0') | $spl48),
  inference(unit_resulting_resolution,[],[f498,f2312,f430])).
fof(f19868,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737,f19738,f19739])).
fof(f19867,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737,f19738])).
fof(f19866,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688,f19737])).
fof(f19864,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19863,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19863,plain,(
  ( ! [X37,X38] : (~sP8(g(s(bool,'T_0')),X37,X38) | sP9(g(s(bool,'T_0')),X37,X38)) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(forward_demodulation,[],[f19714,f19612])).
fof(f19714,plain,(
  ( ! [X37,X38] : (sP9(g(s(bool,'T_0')),X37,X38) | ~sP8(g(s(bool,'F_0')),X37,X38)) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13380])).
fof(f13380,plain,(
  ( ! [X37,X38] : (~sP8(g(s(bool,'F_0')),X37,X38) | sP9(g(s(bool,'F_0')),X37,X38)) ) | $spl48),
  inference(subsumption_resolution,[],[f13317,f275])).
fof(f13317,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | ~sP8(g(s(bool,'F_0')),X37,X38) | sP9(g(s(bool,'F_0')),X37,X38)) ) | $spl48),
  inference(superposition,[],[f483,f2312])).
fof(f19862,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19861,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19861,plain,(
  ( ! [X35,X36] : (~sP10(g(s(bool,'T_0')),X35,X36) | sP11(g(s(bool,'T_0')),X35,X36)) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(forward_demodulation,[],[f19713,f19612])).
fof(f19713,plain,(
  ( ! [X35,X36] : (sP11(g(s(bool,'T_0')),X35,X36) | ~sP10(g(s(bool,'F_0')),X35,X36)) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13379])).
fof(f13379,plain,(
  ( ! [X35,X36] : (~sP10(g(s(bool,'F_0')),X35,X36) | sP11(g(s(bool,'F_0')),X35,X36)) ) | $spl48),
  inference(subsumption_resolution,[],[f13316,f275])).
fof(f13316,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~sP10(g(s(bool,'F_0')),X35,X36) | sP11(g(s(bool,'F_0')),X35,X36)) ) | $spl48),
  inference(superposition,[],[f482,f2312])).
fof(f19860,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19712,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19712,plain,(
  ( ! [X33,X34] : (~sP11(g(s(bool,'T_0')),X34,X33) | ~p(s(bool,X33)) | p(s(bool,X34))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13378])).
fof(f13378,plain,(
  ( ! [X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP11(g(s(bool,'F_0')),X34,X33)) ) | $spl48),
  inference(subsumption_resolution,[],[f13315,f275])).
fof(f13315,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP11(g(s(bool,'F_0')),X34,X33)) ) | $spl48),
  inference(superposition,[],[f479,f2312])).
fof(f19859,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19711,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19711,plain,(
  ( ! [X31,X32] : (sP11(g(s(bool,'T_0')),X31,X32) | ~p(s(bool,X31))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13377])).
fof(f13377,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | sP11(g(s(bool,'F_0')),X31,X32)) ) | $spl48),
  inference(subsumption_resolution,[],[f13314,f275])).
fof(f13314,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | sP11(g(s(bool,'F_0')),X31,X32)) ) | $spl48),
  inference(superposition,[],[f478,f2312])).
fof(f19858,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19710,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19710,plain,(
  ( ! [X30,X29] : (sP11(g(s(bool,'T_0')),X30,X29) | p(s(bool,X29))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13376])).
fof(f13376,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | sP11(g(s(bool,'F_0')),X30,X29)) ) | $spl48),
  inference(subsumption_resolution,[],[f13313,f275])).
fof(f13313,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP11(g(s(bool,'F_0')),X30,X29)) ) | $spl48),
  inference(superposition,[],[f477,f2312])).
fof(f19857,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19709,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19709,plain,(
  ( ! [X28,X27] : (sP11(X28,X27,g(s(bool,'T_0'))) | p(s(bool,X27)) | p(s(bool,X28))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13375])).
fof(f13375,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | sP11(X28,X27,g(s(bool,'F_0'))) | p(s(bool,X28))) ) | $spl48),
  inference(subsumption_resolution,[],[f13312,f275])).
fof(f13312,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP11(X28,X27,g(s(bool,'F_0'))) | p(s(bool,X28))) ) | $spl48),
  inference(superposition,[],[f476,f2312])).
fof(f19856,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19708,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19708,plain,(
  ( ! [X26,X25] : (~sP10(X26,X25,g(s(bool,'T_0'))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13374])).
fof(f13374,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | ~sP10(X26,X25,g(s(bool,'F_0')))) ) | $spl48),
  inference(subsumption_resolution,[],[f13311,f275])).
fof(f13311,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~sP10(X26,X25,g(s(bool,'F_0')))) ) | $spl48),
  inference(superposition,[],[f472,f2312])).
fof(f19855,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19707,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19707,plain,(
  ( ! [X24,X23] : (~sP9(g(s(bool,'T_0')),X24,X23) | p(s(bool,X23)) | p(s(bool,X24))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13373])).
fof(f13373,plain,(
  ( ! [X24,X23] : (p(s(bool,X23)) | p(s(bool,X24)) | ~sP9(g(s(bool,'F_0')),X24,X23)) ) | $spl48),
  inference(subsumption_resolution,[],[f13310,f275])).
fof(f13310,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | p(s(bool,X23)) | p(s(bool,X24)) | ~sP9(g(s(bool,'F_0')),X24,X23)) ) | $spl48),
  inference(superposition,[],[f462,f2312])).
fof(f19854,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19706,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19706,plain,(
  ( ! [X21,X22] : (sP9(g(s(bool,'T_0')),X21,X22) | ~p(s(bool,X21))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13372])).
fof(f13372,plain,(
  ( ! [X21,X22] : (~p(s(bool,X21)) | sP9(g(s(bool,'F_0')),X21,X22)) ) | $spl48),
  inference(subsumption_resolution,[],[f13309,f275])).
fof(f13309,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | sP9(g(s(bool,'F_0')),X21,X22)) ) | $spl48),
  inference(superposition,[],[f461,f2312])).
fof(f19853,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19705,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19705,plain,(
  ( ! [X19,X20] : (sP9(g(s(bool,'T_0')),X20,X19) | ~p(s(bool,X19))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13371])).
fof(f13371,plain,(
  ( ! [X19,X20] : (~p(s(bool,X19)) | sP9(g(s(bool,'F_0')),X20,X19)) ) | $spl48),
  inference(subsumption_resolution,[],[f13308,f275])).
fof(f13308,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | sP9(g(s(bool,'F_0')),X20,X19)) ) | $spl48),
  inference(superposition,[],[f460,f2312])).
fof(f19852,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19704,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19704,plain,(
  ( ! [X17,X18] : (~sP8(g(s(bool,'T_0')),X17,X18) | p(s(bool,X17)) | p(s(bool,X18))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13370])).
fof(f13370,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | ~sP8(g(s(bool,'F_0')),X17,X18)) ) | $spl48),
  inference(subsumption_resolution,[],[f13307,f275])).
fof(f13307,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | p(s(bool,X18)) | ~sP8(g(s(bool,'F_0')),X17,X18)) ) | $spl48),
  inference(superposition,[],[f455,f2312])).
fof(f19851,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19703,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19703,plain,(
  ( ! [X15,X16] : (~sP7(X16,X15,g(s(bool,'T_0'))) | ~p(s(bool,X15)) | p(s(bool,X16))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13369])).
fof(f13369,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | ~sP7(X16,X15,g(s(bool,'F_0')))) ) | $spl48),
  inference(subsumption_resolution,[],[f13306,f275])).
fof(f13306,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP7(X16,X15,g(s(bool,'F_0')))) ) | $spl48),
  inference(superposition,[],[f447,f2312])).
fof(f19850,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19702,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19702,plain,(
  ( ! [X14,X13] : (sP7(g(s(bool,'T_0')),X13,X14) | ~p(s(bool,X13)) | ~p(s(bool,X14))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13368])).
fof(f13368,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | ~p(s(bool,X14)) | sP7(g(s(bool,'F_0')),X13,X14)) ) | $spl48),
  inference(subsumption_resolution,[],[f13305,f275])).
fof(f13305,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | sP7(g(s(bool,'F_0')),X13,X14)) ) | $spl48),
  inference(superposition,[],[f444,f2312])).
fof(f19849,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19701,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19701,plain,(
  ( ! [X12,X11] : (~sP6(X11,g(s(bool,'T_0')),X12) | ~p(s(bool,X11)) | p(s(bool,X12))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13367])).
fof(f13367,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | p(s(bool,X12)) | ~sP6(X11,g(s(bool,'F_0')),X12)) ) | $spl48),
  inference(subsumption_resolution,[],[f13304,f275])).
fof(f13304,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | p(s(bool,X12)) | ~sP6(X11,g(s(bool,'F_0')),X12)) ) | $spl48),
  inference(superposition,[],[f438,f2312])).
fof(f19848,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19700,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19700,plain,(
  ( ! [X10,X9] : (~sP3(X9,g(s(bool,'T_0')),X10) | ~p(s(bool,X9)) | p(s(bool,X10))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13366])).
fof(f13366,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | ~sP3(X9,g(s(bool,'F_0')),X10)) ) | $spl48),
  inference(subsumption_resolution,[],[f13303,f275])).
fof(f13303,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP3(X9,g(s(bool,'F_0')),X10)) ) | $spl48),
  inference(superposition,[],[f419,f2312])).
fof(f19847,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19699,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19699,plain,(
  ( ! [X8,X7] : (~sP4(X7,g(s(bool,'T_0')),X8) | ~p(s(bool,X7)) | p(s(bool,X8))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13365])).
fof(f13365,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,g(s(bool,'F_0')),X8)) ) | $spl48),
  inference(subsumption_resolution,[],[f13302,f275])).
fof(f13302,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,g(s(bool,'F_0')),X8)) ) | $spl48),
  inference(superposition,[],[f415,f2312])).
fof(f19846,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19698,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19698,plain,(
  ( ! [X6] : (sP1(X6,g(s(bool,'T_0'))) | p(s(bool,X6))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13364])).
fof(f13364,plain,(
  ( ! [X6] : (sP1(X6,g(s(bool,'F_0'))) | p(s(bool,X6))) ) | $spl48),
  inference(subsumption_resolution,[],[f13301,f275])).
fof(f13301,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,g(s(bool,'F_0')))) ) | $spl48),
  inference(superposition,[],[f385,f2312])).
fof(f19845,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19697,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19697,plain,(
  ( ! [X5] : (sP1(g(s(bool,'T_0')),X5) | p(s(bool,X5))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13363])).
fof(f13363,plain,(
  ( ! [X5] : (sP1(g(s(bool,'F_0')),X5) | p(s(bool,X5))) ) | $spl48),
  inference(subsumption_resolution,[],[f13300,f275])).
fof(f13300,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(g(s(bool,'F_0')),X5)) ) | $spl48),
  inference(superposition,[],[f382,f2312])).
fof(f19844,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19696,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19696,plain,(
  ( ! [X4] : (sP0(g(s(bool,'T_0')),X4)) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13362])).
fof(f13362,plain,(
  ( ! [X4] : (sP0(g(s(bool,'F_0')),X4)) ) | $spl48),
  inference(subsumption_resolution,[],[f13299,f275])).
fof(f13299,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | sP0(g(s(bool,'F_0')),X4)) ) | $spl48),
  inference(superposition,[],[f367,f2312])).
fof(f19843,plain,(
  $false | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f19695,f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19695,plain,(
  ( ! [X3] : (sP0(X3,g(s(bool,'T_0')))) ) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13361])).
fof(f13361,plain,(
  ( ! [X3] : (sP0(X3,g(s(bool,'F_0')))) ) | $spl48),
  inference(subsumption_resolution,[],[f13298,f275])).
fof(f13298,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | sP0(X3,g(s(bool,'F_0')))) ) | $spl48),
  inference(superposition,[],[f366,f2312])).
fof(f19842,plain,(
  $false | ($spl47 | $spl48 | $spl71 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19694,f6631])).
fof(f6631,plain,(
  ~sP5('T_0','T_0',g(s(bool,'T_0'))) | $spl71),
  inference(unit_resulting_resolution,[],[f275,f6477,f429])).
fof(f19694,plain,(
  sP5('T_0','T_0',g(s(bool,'T_0'))) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13275])).
fof(f13275,plain,(
  sP5('T_0','T_0',g(s(bool,'F_0'))) | $spl48),
  inference(unit_resulting_resolution,[],[f275,f2312,f428])).
fof(f19841,plain,(
  $false | ($spl47 | $spl48 | $spl71 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19693,f6630])).
fof(f19693,plain,(
  ~sP5('F_0','T_0',g(s(bool,'T_0'))) | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13274])).
fof(f19840,plain,(
  $false | ($spl47 | $spl48 | $spl71 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19692,f6582])).
fof(f19692,plain,(
  sP5('T_0',g(s(bool,'T_0')),'T_0') | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13273])).
fof(f19839,plain,(
  $false | ($spl47 | $spl48 | $spl71 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19691,f6581])).
fof(f19691,plain,(
  ~sP5('F_0',g(s(bool,'T_0')),'T_0') | ($spl47 | $spl48 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f13272])).
fof(f19838,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687,f19688])).
fof(f19837,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686,f19687])).
fof(f19836,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685,f19686])).
fof(f19835,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831,f19685])).
fof(f19832,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828,f19831])).
fof(f19829,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825,f19828])).
fof(f19827,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19621,f19612])).
fof(f19621,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f2310])).
fof(f19826,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822,f19825])).
fof(f19824,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19615,f275])).
fof(f19615,plain,(
  ~p(s(bool,'T_0')) | ($spl47 | $spl74 | $spl98)),
  inference(backward_demodulation,[],[f19612,f498])).
fof(f19823,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6822,f6905,f6906,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7878,f7880,f7908,f8040,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7004,f7001,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f9459,f9460,f9608,f9610,f9612,f9614,f9616,f9618,f9620,f9622,f9624,f9626,f9628,f9630,f9632,f9634,f9636,f9638,f9640,f9642,f9644,f9646,f9648,f9679,f9680,f9528,f7879,f8088,f9808,f9816,f9868,f9876,f9893,f10012,f10149,f10161,f10400,f10405,f10411,f10373,f10374,f10375,f10734,f10735,f10884,f10922,f12147,f12150,f12198,f12201,f12248,f12251,f12309,f12312,f12640,f12652,f12639,f12688,f12700,f12772,f10256,f10270,f10251,f10250,f10248,f10228,f12778,f10226,f10225,f10224,f10223,f10222,f10221,f7874,f7622,f7617,f8157,f12780,f2330,f2329,f2328,f10254,f2310,f13099,f13090,f13089,f12980,f12979,f12977,f12976,f13100,f13239,f13119,f13242,f13255,f13403,f13402,f13396,f13395,f13404,f13473,f13480,f13481,f13482,f13485,f13486,f13487,f13490,f13609,f13622,f13625,f13643,f13646,f13653,f13656,f13671,f13690,f13693,f13702,f13722,f13743,f13746,f13782,f13785,f13825,f13828,f13868,f13871,f13877,f13889,f13907,f13910,f13942,f13981,f14000,f14003,f14011,f14013,f14024,f14027,f14045,f14048,f14056,f14058,f14074,f14093,f14096,f14149,f14152,f14202,f14205,f17310,f14455,f14523,f14557,f14558,f14560,f17451,f17454,f14598,f14842,f14845,f14846,f14847,f17694,f17698,f17704,f17706,f17709,f17711,f17776,f17778,f17785,f17787,f17790,f15124,f15127,f15128,f15129,f15130,f17792,f16503,f16501,f16499,f17811,f17812,f17813,f17814,f17816,f17817,f17818,f17819,f17820,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f15265,f15268,f15269,f15273,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17914,f17918,f15476,f15479,f15480,f15481,f17929,f17930,f17932,f17933,f17997,f18003,f18004,f18090,f16014,f12776,f18170,f18189,f18192,f18197,f18214,f18240,f18249,f18269,f18274,f18293,f18296,f18301,f18317,f18339,f18379,f18535,f18538,f18540,f18543,f18666,f18667,f18668,f18669,f18670,f18675,f18758,f18903,f18889,f18921,f18922,f18923,f18924,f18925,f18926,f18927,f18928,f18929,f18930,f18931,f18932,f18933,f18934,f18935,f18936,f18937,f18938,f18939,f18940,f18941,f18942,f18943,f18948,f18988,f19051,f19075,f19076,f19077,f19078,f19081,f19083,f19084,f19085,f19116,f19117,f19124,f19129,f19169,f19232,f19256,f19257,f19258,f19259,f19262,f19264,f19265,f19266,f19297,f19298,f19305,f19310,f19350,f19413,f19472,f19481,f19482,f19483,f19487,f19612,f19822])).
fof(f19611,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19521,f13099])).
fof(f19521,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl98),
  inference(unit_resulting_resolution,[],[f522,f12776])).
fof(f19610,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19609,f13099])).
fof(f19609,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | ($spl47 | $spl74 | $spl98)),
  inference(forward_demodulation,[],[f19522,f13099])).
fof(f19522,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl47 | $spl98)),
  inference(unit_resulting_resolution,[],[f2310,f12776])).
fof(f19608,plain,(
  $false | ($spl47 | $spl71 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19607,f13099])).
fof(f19607,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | ($spl71 | $spl98)),
  inference(forward_demodulation,[],[f19523,f6656])).
fof(f6656,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl71),
  inference(forward_demodulation,[],[f6567,f1243])).
fof(f6567,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl71),
  inference(unit_resulting_resolution,[],[f1186,f1163,f6477,f1721])).
fof(f19523,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,g(s(bool,'T_0'))))) | ($spl71 | $spl98)),
  inference(unit_resulting_resolution,[],[f6477,f12776])).
fof(f19606,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19605,f13099])).
fof(f19605,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl98),
  inference(forward_demodulation,[],[f19524,f1243])).
fof(f19524,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl98),
  inference(unit_resulting_resolution,[],[f1095,f12776])).
fof(f19604,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(subsumption_resolution,[],[f19525,f522])).
fof(f19525,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl47 | $spl74 | $spl98)),
  inference(unit_resulting_resolution,[],[f13099,f12776])).
fof(f19526,plain,(
  $false | ($spl47 | $spl74 | $spl98)),
  inference(unit_resulting_resolution,[],[f522,f13099,f12776])).
fof(f2308,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl44),
  inference(cnf_transformation,[],[f2308_D])).
fof(f2308_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f19517,plain,(
  ( ! [X39,X43,X44,X40] : (s(bool,X40) != s(bool,f(s(bool,X39))) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X40) | s(bool,X43) = s(bool,g(s(bool,X39))) | s(bool,'T_0') = s(bool,X39)) ) | $spl134),
  inference(cnf_transformation,[],[f19517_D])).
fof(f19517_D,plain,(
  ( ! [X39,X43,X44,X40] : (s(bool,X40) != s(bool,f(s(bool,X39))) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X40) | s(bool,X43) = s(bool,g(s(bool,X39))) | s(bool,'T_0') = s(bool,X39)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f19487,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X64) != s(bool,f(s(bool,X63))) | s(fun(bool,bool),X65) = s(fun(bool,bool),X66) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,sK17(X65,X66,bool,bool)) | s(bool,X67) = s(bool,g(s(bool,X63))) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f1340,f1281])).
fof(f19483,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s(bool,X40) != s(bool,f(s(bool,X39))) | s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,sK17(X41,X42,bool,bool)) | s(bool,X43) = s(bool,g(s(bool,X39))) | s(bool,'T_0') = s(bool,X40) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f1340,f1281])).
fof(f19516,plain,(
  ( ! [X28,X31,X27,X32] : (s(bool,X28) != s(bool,f(s(bool,X27))) | s(bool,'T_0') = s(bool,X32) | s(bool,X32) = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X31) | s(bool,X28) = s(bool,X31) | s(bool,'T_0') = s(bool,X27)) ) | $spl132),
  inference(cnf_transformation,[],[f19516_D])).
fof(f19516_D,plain,(
  ( ! [X28,X31,X27,X32] : (s(bool,X28) != s(bool,f(s(bool,X27))) | s(bool,'T_0') = s(bool,X32) | s(bool,X32) = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X31) | s(bool,X28) = s(bool,X31) | s(bool,'T_0') = s(bool,X27)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f19482,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,X34) != s(bool,f(s(bool,X33))) | s(fun(bool,bool),X35) = s(fun(bool,bool),X36) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,sK17(X35,X36,bool,bool)) | s(bool,X37) = s(bool,g(s(bool,X33))) | s(bool,'T_0') = s(bool,X37) | s(bool,X34) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f1340,f1281])).
fof(f19481,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (s(bool,X28) != s(bool,f(s(bool,X27))) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,sK17(X29,X30,bool,bool)) | s(bool,X28) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f1340,f1281])).
fof(f18297,plain,(
  ( ! [X35,X31,X36,X34,X32] : (s(bool,'T_0') != s(bool,sK15(X32,bool,X31)) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,sK16(X32,bool,X31))) ) | $spl130),
  inference(cnf_transformation,[],[f18297_D])).
fof(f18297_D,plain,(
  ( ! [X35,X31,X36,X34,X32] : (s(bool,'T_0') != s(bool,sK15(X32,bool,X31)) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,sK16(X32,bool,X31))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f18301,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (p(s(bool,X57)) | s(bool,'T_0') != s(bool,sK15(X56,bool,X55)) | s(bool,X58) = s(bool,sK16(X56,bool,X55)) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X55),X56))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f18286,f294])).
fof(f18286,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (p(s(bool,X57)) | s(bool,'T_0') != s(bool,sK15(X56,bool,X55)) | s(bool,X58) = s(bool,sK16(X56,bool,X55)) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X55),X56))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f1449,f1281])).
fof(f18296,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,X33)) | s(bool,'T_0') != s(bool,sK15(X32,bool,X31)) | s(bool,X34) = s(bool,sK16(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f18282,f294])).
fof(f18282,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,X33)) | s(bool,'T_0') != s(bool,sK15(X32,bool,X31)) | s(bool,X34) = s(bool,sK16(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1449,f1281])).
fof(f18193,plain,(
  ( ! [X35,X31,X36,X34,X32] : (s(bool,X34) != s(bool,sK15(X32,bool,X31)) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,sK16(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) | $spl128),
  inference(cnf_transformation,[],[f18193_D])).
fof(f18193_D,plain,(
  ( ! [X35,X31,X36,X34,X32] : (s(bool,X34) != s(bool,sK15(X32,bool,X31)) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,sK16(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f18197,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (p(s(bool,X57)) | s(bool,X58) != s(bool,sK15(X56,bool,X55)) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,sK16(X56,bool,X55)) | s(bool,X59) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X55),X56))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f18182,f294])).
fof(f18182,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (p(s(bool,X57)) | s(bool,X58) != s(bool,sK15(X56,bool,X55)) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,sK16(X56,bool,X55)) | s(bool,X59) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X55),X56))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f1339,f1281])).
fof(f18192,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,X33)) | s(bool,X34) != s(bool,sK15(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,sK16(X32,bool,X31)) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f18178,f294])).
fof(f18178,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,X33)) | s(bool,X34) != s(bool,sK15(X32,bool,X31)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,sK16(X32,bool,X31)) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X31),X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1339,f1281])).
fof(f17788,plain,(
  ( ! [X565,X563,X566,X562] : (s(bool,'T_0') = s(bool,sK16(X563,bool,X562)) | s(bool,'T_0') = s(bool,X566) | s(bool,X565) = s(bool,X566) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,sK15(X563,bool,X562))) ) | $spl122),
  inference(cnf_transformation,[],[f17788_D])).
fof(f17788_D,plain,(
  ( ! [X565,X563,X566,X562] : (s(bool,'T_0') = s(bool,sK16(X563,bool,X562)) | s(bool,'T_0') = s(bool,X566) | s(bool,X565) = s(bool,X566) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,sK15(X563,bool,X562))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f18004,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK16(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK15(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(subsumption_resolution,[],[f15615,f294])).
fof(f15615,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK16(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK15(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566) | s(bool,'T_0') = s(bool,X564)) )),
  inference(superposition,[],[f2213,f1281])).
fof(f17786,plain,(
  ( ! [X561,X557,X560,X556,X558] : (s(X557,happ(s(fun(X556,X557),X558),s(X556,sK15(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK16(X558,X556,X557)))) | s(bool,'T_0') = s(bool,X561) | s(bool,X560) = s(bool,X561) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558)))) ) | $spl120),
  inference(cnf_transformation,[],[f17786_D])).
fof(f17786_D,plain,(
  ( ! [X561,X557,X560,X556,X558] : (s(X557,happ(s(fun(X556,X557),X558),s(X556,sK15(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK16(X558,X556,X557)))) | s(bool,'T_0') = s(bool,X561) | s(bool,X560) = s(bool,X561) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f18003,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK15(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK16(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(subsumption_resolution,[],[f15614,f294])).
fof(f15614,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK15(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK16(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561) | s(bool,'T_0') = s(bool,X559)) )),
  inference(superposition,[],[f399,f1281])).
fof(f17779,plain,(
  ( ! [X528,X530,X531,X527] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X528),s('type.rat.rat',X527))) | s(bool,'T_0') = s(bool,X531) | s(bool,X530) = s(bool,X531) | s(bool,X530) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X527),s('type.rat.rat',X528)))))))) ) | $spl118),
  inference(cnf_transformation,[],[f17779_D])).
fof(f17779_D,plain,(
  ( ! [X528,X530,X531,X527] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X528),s('type.rat.rat',X527))) | s(bool,'T_0') = s(bool,X531) | s(bool,X530) = s(bool,X531) | s(bool,X530) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X527),s('type.rat.rat',X528)))))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f17997,plain,(
  ( ! [X528,X530,X529,X531,X527] : (p(s(bool,X529)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X528),s('type.rat.rat',X527))) | s(bool,X530) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X527),s('type.rat.rat',X528))))))) | s(bool,X530) = s(bool,X531) | s(bool,'T_0') = s(bool,X531)) )),
  inference(subsumption_resolution,[],[f15608,f294])).
fof(f15608,plain,(
  ( ! [X528,X530,X529,X531,X527] : (p(s(bool,X529)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X528),s('type.rat.rat',X527))) | s(bool,X530) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X527),s('type.rat.rat',X528))))))) | s(bool,X530) = s(bool,X531) | s(bool,'T_0') = s(bool,X531) | s(bool,'T_0') = s(bool,X529)) )),
  inference(superposition,[],[f868,f1281])).
fof(f17712,plain,(
  ( ! [X259,X260,X256,X258] : (sP1(X256,X258) | s(bool,'T_0') = s(bool,X260) | s(bool,X259) = s(bool,X260) | s(bool,X256) = s(bool,X259) | s(bool,'F_0') = s(bool,X258)) ) | $spl116),
  inference(cnf_transformation,[],[f17712_D])).
fof(f17712_D,plain,(
  ( ! [X259,X260,X256,X258] : (sP1(X256,X258) | s(bool,'T_0') = s(bool,X260) | s(bool,X259) = s(bool,X260) | s(bool,X256) = s(bool,X259) | s(bool,'F_0') = s(bool,X258)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f17933,plain,(
  ( ! [X257,X259,X260,X256,X258] : (p(s(bool,X257)) | sP1(X256,X258) | s(bool,'F_0') = s(bool,X258) | s(bool,X256) = s(bool,X259) | s(bool,X259) = s(bool,X260) | s(bool,'T_0') = s(bool,X260)) ) | $spl74),
  inference(subsumption_resolution,[],[f15532,f294])).
fof(f15532,plain,(
  ( ! [X257,X259,X260,X256,X258] : (p(s(bool,X257)) | sP1(X256,X258) | s(bool,'F_0') = s(bool,X258) | s(bool,X256) = s(bool,X259) | s(bool,X259) = s(bool,X260) | s(bool,'T_0') = s(bool,X260) | s(bool,'T_0') = s(bool,X257)) ) | $spl74),
  inference(superposition,[],[f9616,f1281])).
fof(f17710,plain,(
  ( ! [X255,X253,X251,X254] : (sP1(X253,X251) | s(bool,'T_0') = s(bool,X255) | s(bool,X254) = s(bool,X255) | s(bool,X251) = s(bool,X254) | s(bool,'F_0') = s(bool,X253)) ) | $spl114),
  inference(cnf_transformation,[],[f17710_D])).
fof(f17710_D,plain,(
  ( ! [X255,X253,X251,X254] : (sP1(X253,X251) | s(bool,'T_0') = s(bool,X255) | s(bool,X254) = s(bool,X255) | s(bool,X251) = s(bool,X254) | s(bool,'F_0') = s(bool,X253)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f17932,plain,(
  ( ! [X255,X253,X251,X254,X252] : (p(s(bool,X252)) | sP1(X253,X251) | s(bool,'F_0') = s(bool,X253) | s(bool,X251) = s(bool,X254) | s(bool,X254) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | $spl74),
  inference(subsumption_resolution,[],[f15531,f294])).
fof(f15531,plain,(
  ( ! [X255,X253,X251,X254,X252] : (p(s(bool,X252)) | sP1(X253,X251) | s(bool,'F_0') = s(bool,X253) | s(bool,X251) = s(bool,X254) | s(bool,X254) = s(bool,X255) | s(bool,'T_0') = s(bool,X255) | s(bool,'T_0') = s(bool,X252)) ) | $spl74),
  inference(superposition,[],[f9614,f1281])).
fof(f17707,plain,(
  ( ! [X212,X210,X211,X209,X207] : (sP1(X207,X209) | s(bool,'T_0') = s(bool,X212) | s(bool,X211) = s(bool,X212) | s(bool,X207) = s(bool,X211) | s(bool,'T_0') = s(bool,X210) | s(bool,X209) = s(bool,X210)) ) | $spl112),
  inference(cnf_transformation,[],[f17707_D])).
fof(f17707_D,plain,(
  ( ! [X212,X210,X211,X209,X207] : (sP1(X207,X209) | s(bool,'T_0') = s(bool,X212) | s(bool,X211) = s(bool,X212) | s(bool,X207) = s(bool,X211) | s(bool,'T_0') = s(bool,X210) | s(bool,X209) = s(bool,X210)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f17930,plain,(
  ( ! [X212,X210,X208,X211,X209,X207] : (p(s(bool,X208)) | sP1(X207,X209) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X207) = s(bool,X211) | s(bool,X211) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(subsumption_resolution,[],[f15522,f294])).
fof(f15522,plain,(
  ( ! [X212,X210,X208,X211,X209,X207] : (p(s(bool,X208)) | sP1(X207,X209) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X207) = s(bool,X211) | s(bool,X211) = s(bool,X212) | s(bool,'T_0') = s(bool,X212) | s(bool,'T_0') = s(bool,X208)) )),
  inference(superposition,[],[f1705,f1281])).
fof(f17705,plain,(
  ( ! [X206,X204,X205,X203,X201] : (sP1(X203,X201) | s(bool,'T_0') = s(bool,X206) | s(bool,X205) = s(bool,X206) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X204) | s(bool,X203) = s(bool,X204)) ) | $spl110),
  inference(cnf_transformation,[],[f17705_D])).
fof(f17705_D,plain,(
  ( ! [X206,X204,X205,X203,X201] : (sP1(X203,X201) | s(bool,'T_0') = s(bool,X206) | s(bool,X205) = s(bool,X206) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X204) | s(bool,X203) = s(bool,X204)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f17929,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,X202)) | sP1(X203,X201) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) )),
  inference(subsumption_resolution,[],[f15521,f294])).
fof(f15521,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,X202)) | sP1(X203,X201) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206) | s(bool,'T_0') = s(bool,X202)) )),
  inference(superposition,[],[f1704,f1281])).
fof(f17703,plain,(
  ( ! [X189,X188,X186] : (sP1(X186,'T_0') | s(bool,'T_0') = s(bool,X189) | s(bool,X188) = s(bool,X189) | s(bool,X186) = s(bool,X188)) ) | $spl108),
  inference(cnf_transformation,[],[f17703_D])).
fof(f17703_D,plain,(
  ( ! [X189,X188,X186] : (sP1(X186,'T_0') | s(bool,'T_0') = s(bool,X189) | s(bool,X188) = s(bool,X189) | s(bool,X186) = s(bool,X188)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f17928,plain,(
  ( ! [X189,X187,X188,X186] : (p(s(bool,X187)) | sP1(X186,'T_0') | s(bool,X186) = s(bool,X188) | s(bool,X188) = s(bool,X189) | s(bool,'T_0') = s(bool,X189)) ) | $spl8),
  inference(subsumption_resolution,[],[f15518,f294])).
fof(f15518,plain,(
  ( ! [X189,X187,X188,X186] : (p(s(bool,X187)) | sP1(X186,'T_0') | s(bool,X186) = s(bool,X188) | s(bool,X188) = s(bool,X189) | s(bool,'T_0') = s(bool,X189) | s(bool,'T_0') = s(bool,X187)) ) | $spl8),
  inference(superposition,[],[f1651,f1281])).
fof(f17701,plain,(
  ( ! [X185,X182,X184] : (sP1('T_0',X182) | s(bool,'T_0') = s(bool,X185) | s(bool,X184) = s(bool,X185) | s(bool,X182) = s(bool,X184)) ) | $spl106),
  inference(cnf_transformation,[],[f17701_D])).
fof(f17701_D,plain,(
  ( ! [X185,X182,X184] : (sP1('T_0',X182) | s(bool,'T_0') = s(bool,X185) | s(bool,X184) = s(bool,X185) | s(bool,X182) = s(bool,X184)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f17927,plain,(
  ( ! [X185,X182,X184,X183] : (p(s(bool,X183)) | sP1('T_0',X182) | s(bool,X182) = s(bool,X184) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) ) | $spl4),
  inference(subsumption_resolution,[],[f15517,f294])).
fof(f15517,plain,(
  ( ! [X185,X182,X184,X183] : (p(s(bool,X183)) | sP1('T_0',X182) | s(bool,X182) = s(bool,X184) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185) | s(bool,'T_0') = s(bool,X183)) ) | $spl4),
  inference(superposition,[],[f1647,f1281])).
fof(f17793,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) | $spl126),
  inference(cnf_transformation,[],[f17793_D])).
fof(f17793_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f17792,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f15132,f294])).
fof(f15132,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f294,f1281])).
fof(f17791,plain,(
  ( ! [X613,X615,X616,X612] : (s(fun(bool,bool),X612) = s(fun(bool,bool),X613) | s(bool,'T_0') = s(bool,X616) | s(bool,X615) = s(bool,X616) | s(bool,X615) = s(bool,sK17(X612,X613,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X612,X613,bool,bool))))) ) | $spl124),
  inference(cnf_transformation,[],[f17791_D])).
fof(f17791_D,plain,(
  ( ! [X613,X615,X616,X612] : (s(fun(bool,bool),X612) = s(fun(bool,bool),X613) | s(bool,'T_0') = s(bool,X616) | s(bool,X615) = s(bool,X616) | s(bool,X615) = s(bool,sK17(X612,X613,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X612,X613,bool,bool))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f12776,plain,(
  ( ! [X24] : (s(bool,'F_0') != s(bool,f(s(bool,X24))) | s(bool,'T_0') = s(bool,X24)) ) | $spl98),
  inference(cnf_transformation,[],[f12776_D])).
fof(f12776_D,plain,(
  ( ! [X24] : (s(bool,'F_0') != s(bool,f(s(bool,X24))) | s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f15130,plain,(
  ( ! [X613,X615,X616,X612,X614] : (s(bool,'F_0') != s(bool,f(s(bool,X614))) | s(fun(bool,bool),X612) = s(fun(bool,bool),X613) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X612,X613,bool,bool)))) | s(bool,X615) = s(bool,sK17(X612,X613,bool,bool)) | s(bool,'T_0') = s(bool,X614) | s(bool,X615) = s(bool,X616) | s(bool,'T_0') = s(bool,X616)) )),
  inference(superposition,[],[f1054,f1281])).
fof(f17787,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK16(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK15(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(subsumption_resolution,[],[f15122,f294])).
fof(f15122,plain,(
  ( ! [X565,X563,X564,X566,X562] : (p(s(bool,X564)) | s(bool,'T_0') = s(bool,sK16(X563,bool,X562)) | s(bool,'T_0') = s(bool,sK15(X563,bool,X562)) | s(bool,X565) = s(bool,'const.bool.ONE_ONE_1'(s(fun(bool,X562),X563))) | s(bool,'T_0') = s(bool,X564) | s(bool,X565) = s(bool,X566) | s(bool,'T_0') = s(bool,X566)) )),
  inference(superposition,[],[f2213,f1281])).
fof(f17785,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK15(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK16(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(subsumption_resolution,[],[f15121,f294])).
fof(f15121,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,X559)) | s(X557,happ(s(fun(X556,X557),X558),s(X556,sK15(X558,X556,X557)))) = s(X557,happ(s(fun(X556,X557),X558),s(X556,sK16(X558,X556,X557)))) | s(bool,X560) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X556,X557),X558))) | s(bool,'T_0') = s(bool,X559) | s(bool,X560) = s(bool,X561) | s(bool,'T_0') = s(bool,X561)) )),
  inference(superposition,[],[f399,f1281])).
fof(f17778,plain,(
  ( ! [X528,X530,X529,X531,X527] : (p(s(bool,X529)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X528),s('type.rat.rat',X527))) | s(bool,X530) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X527),s('type.rat.rat',X528))))))) | s(bool,X530) = s(bool,X531) | s(bool,'T_0') = s(bool,X531)) )),
  inference(subsumption_resolution,[],[f15115,f294])).
fof(f15115,plain,(
  ( ! [X528,X530,X529,X531,X527] : (p(s(bool,X529)) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X528),s('type.rat.rat',X527))) | s(bool,X530) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X527),s('type.rat.rat',X528))))))) | s(bool,'T_0') = s(bool,X529) | s(bool,X530) = s(bool,X531) | s(bool,'T_0') = s(bool,X531)) )),
  inference(superposition,[],[f868,f1281])).
fof(f17711,plain,(
  ( ! [X257,X259,X260,X256,X258] : (p(s(bool,X257)) | sP1(X256,X258) | s(bool,'F_0') = s(bool,X258) | s(bool,X256) = s(bool,X259) | s(bool,X259) = s(bool,X260) | s(bool,'T_0') = s(bool,X260)) ) | $spl74),
  inference(subsumption_resolution,[],[f15039,f294])).
fof(f15039,plain,(
  ( ! [X257,X259,X260,X256,X258] : (p(s(bool,X257)) | sP1(X256,X258) | s(bool,'F_0') = s(bool,X258) | s(bool,X256) = s(bool,X259) | s(bool,'T_0') = s(bool,X257) | s(bool,X259) = s(bool,X260) | s(bool,'T_0') = s(bool,X260)) ) | $spl74),
  inference(superposition,[],[f9616,f1281])).
fof(f17709,plain,(
  ( ! [X255,X253,X251,X254,X252] : (p(s(bool,X252)) | sP1(X253,X251) | s(bool,'F_0') = s(bool,X253) | s(bool,X251) = s(bool,X254) | s(bool,X254) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | $spl74),
  inference(subsumption_resolution,[],[f15038,f294])).
fof(f15038,plain,(
  ( ! [X255,X253,X251,X254,X252] : (p(s(bool,X252)) | sP1(X253,X251) | s(bool,'F_0') = s(bool,X253) | s(bool,X251) = s(bool,X254) | s(bool,'T_0') = s(bool,X252) | s(bool,X254) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | $spl74),
  inference(superposition,[],[f9614,f1281])).
fof(f17706,plain,(
  ( ! [X212,X210,X208,X211,X209,X207] : (p(s(bool,X208)) | sP1(X207,X209) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X207) = s(bool,X211) | s(bool,X211) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(subsumption_resolution,[],[f15029,f294])).
fof(f15029,plain,(
  ( ! [X212,X210,X208,X211,X209,X207] : (p(s(bool,X208)) | sP1(X207,X209) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X207) = s(bool,X211) | s(bool,'T_0') = s(bool,X208) | s(bool,X211) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f1705,f1281])).
fof(f17704,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,X202)) | sP1(X203,X201) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) )),
  inference(subsumption_resolution,[],[f15028,f294])).
fof(f15028,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,X202)) | sP1(X203,X201) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X202) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) )),
  inference(superposition,[],[f1704,f1281])).
fof(f17702,plain,(
  ( ! [X189,X187,X188,X186] : (p(s(bool,X187)) | sP1(X186,'T_0') | s(bool,X186) = s(bool,X188) | s(bool,X188) = s(bool,X189) | s(bool,'T_0') = s(bool,X189)) ) | $spl8),
  inference(subsumption_resolution,[],[f15025,f294])).
fof(f15025,plain,(
  ( ! [X189,X187,X188,X186] : (p(s(bool,X187)) | sP1(X186,'T_0') | s(bool,X186) = s(bool,X188) | s(bool,'T_0') = s(bool,X187) | s(bool,X188) = s(bool,X189) | s(bool,'T_0') = s(bool,X189)) ) | $spl8),
  inference(superposition,[],[f1651,f1281])).
fof(f17700,plain,(
  ( ! [X185,X182,X184,X183] : (p(s(bool,X183)) | sP1('T_0',X182) | s(bool,X182) = s(bool,X184) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) ) | $spl4),
  inference(subsumption_resolution,[],[f15024,f294])).
fof(f15024,plain,(
  ( ! [X185,X182,X184,X183] : (p(s(bool,X183)) | sP1('T_0',X182) | s(bool,X182) = s(bool,X184) | s(bool,'T_0') = s(bool,X183) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) ) | $spl4),
  inference(superposition,[],[f1647,f1281])).
fof(f17699,plain,(
  ( ! [X180,X178,X176,X181,X179] : (s(bool,X176) != s(bool,X178) | s(bool,'T_0') = s(bool,X181) | s(bool,X180) = s(bool,X181) | s(bool,X176) = s(bool,X180) | s(bool,'T_0') = s(bool,X179) | s(bool,X178) = s(bool,X179)) ) | $spl104),
  inference(cnf_transformation,[],[f17699_D])).
fof(f17699_D,plain,(
  ( ! [X180,X178,X176,X181,X179] : (s(bool,X176) != s(bool,X178) | s(bool,'T_0') = s(bool,X181) | s(bool,X180) = s(bool,X181) | s(bool,X176) = s(bool,X180) | s(bool,'T_0') = s(bool,X179) | s(bool,X178) = s(bool,X179)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f17698,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (p(s(bool,X177)) | s(bool,X176) != s(bool,X178) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179) | s(bool,X176) = s(bool,X180) | s(bool,X180) = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) )),
  inference(subsumption_resolution,[],[f15023,f294])).
fof(f15023,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (p(s(bool,X177)) | s(bool,X176) != s(bool,X178) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179) | s(bool,X176) = s(bool,X180) | s(bool,'T_0') = s(bool,X177) | s(bool,X180) = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f1381,f1281])).
fof(f12774,plain,(
  ( ! [X26,X25] : (s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,'T_0') = s(bool,sK17(X25,X26,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X25,X26,bool,bool))))) ) | $spl94),
  inference(cnf_transformation,[],[f12774_D])).
fof(f12774_D,plain,(
  ( ! [X26,X25] : (s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,'T_0') = s(bool,sK17(X25,X26,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X25,X26,bool,bool))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2328,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X6,X7,bool,bool)))) | s(bool,'T_0') = s(bool,sK17(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1054,f608])).
fof(f2329,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X11))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X9,X10,bool,bool)))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,sK17(X9,X10,bool,bool))) )),
  inference(superposition,[],[f1054,f608])).
fof(f2315,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl53),
  inference(cnf_transformation,[],[f2315_D])).
fof(f2315_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f2330,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X12,X13,bool,bool)))) | s(bool,'T_0') = s(bool,sK17(X12,X13,bool,bool))) )),
  inference(superposition,[],[f1054,f283])).
fof(f2313,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl50),
  inference(cnf_transformation,[],[f2313_D])).
fof(f2313_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2312,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl48),
  inference(cnf_transformation,[],[f2312_D])).
fof(f2312_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12780,plain,(
  ( ! [X337,X336,X338] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X336) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X337) = s(fun(bool,bool),X338) | s(bool,'T_0') = s(bool,sK17(X337,X338,bool,bool)) | s(bool,'T_0') = s(bool,X336)) )),
  inference(forward_demodulation,[],[f7470,f877])).
fof(f7470,plain,(
  ( ! [X337,X336,X338] : (s(bool,X336) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X337) = s(fun(bool,bool),X338) | s(bool,'T_0') = s(bool,sK17(X337,X338,bool,bool)) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X336)) )),
  inference(superposition,[],[f1053,f6906])).
fof(f8157,plain,(
  ( ! [X356,X357,X355] : (s(bool,'T_0') = s(bool,sK17(X355,X356,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X357))) | s(fun(bool,bool),X355) = s(fun(bool,bool),X356) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X355,X356,bool,bool)))) | s(bool,'T_0') = s(bool,X357)) )),
  inference(forward_demodulation,[],[f7475,f877])).
fof(f7475,plain,(
  ( ! [X356,X357,X355] : (s(bool,'F_0') != s(bool,f(s(bool,X357))) | s(fun(bool,bool),X355) = s(fun(bool,bool),X356) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X355,X356,bool,bool)))) | s(bool,sK17(X355,X356,bool,bool)) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X357)) )),
  inference(superposition,[],[f1054,f6906])).
fof(f12777,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(bool,X20) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl100),
  inference(cnf_transformation,[],[f12777_D])).
fof(f12777_D,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(bool,X20) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f7617,plain,(
  ( ! [X337,X336,X338] : (s(bool,X336) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X337) = s(fun(bool,bool),X338) | s(bool,'T_0') = s(bool,sK17(X337,X338,bool,bool)) | s(bool,X336) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1053,f6906])).
fof(f12775,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,f(s(bool,X27))) | s(bool,X27) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl96),
  inference(cnf_transformation,[],[f12775_D])).
fof(f12775_D,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,f(s(bool,X27))) | s(bool,X27) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f7622,plain,(
  ( ! [X356,X357,X355] : (s(bool,'F_0') != s(bool,f(s(bool,X357))) | s(fun(bool,bool),X355) = s(fun(bool,bool),X356) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X355,X356,bool,bool)))) | s(bool,X357) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,sK17(X355,X356,bool,bool))) )),
  inference(superposition,[],[f1054,f6906])).
fof(f12779,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X16) = s(bool,g(s(bool,'F_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f12779_D])).
fof(f12779_D,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X16) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f2310,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl47),
  inference(cnf_transformation,[],[f2310_D])).
fof(f2310_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f7874,plain,(
  ( ! [X321,X323,X322] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X321) = s(fun(bool,bool),X322) | s(bool,'T_0') = s(bool,sK17(X321,X322,bool,bool)) | s(bool,X323) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X323) = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1053,f6906])).
fof(f10221,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1053,f283])).
fof(f2307,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f2307_D])).
fof(f2307_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f10222,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK17(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1053,f608])).
fof(f10223,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK17(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1053,f608])).
fof(f10224,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK17(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1053,f608])).
fof(f10225,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK17(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1053,f608])).
fof(f10226,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,sK17(X14,X15,bool,bool)) | s(bool,X16) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X16) = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1053,f6906])).
fof(f12778,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK17(X18,X19,bool,bool)) | s(bool,'T_0') = s(bool,X17)) )),
  inference(forward_demodulation,[],[f10227,f877])).
fof(f10227,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK17(X18,X19,bool,bool)) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f1053,f6906])).
fof(f10228,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK17(X21,X22,bool,bool)) | s(bool,X20) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1053,f6906])).
fof(f10248,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,sK17(X2,X3,bool,bool))) )),
  inference(superposition,[],[f1054,f283])).
fof(f10250,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X7,X8,bool,bool)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK17(X7,X8,bool,bool))) )),
  inference(superposition,[],[f1054,f608])).
fof(f10251,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') != s(bool,f(s(bool,X12))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X10,X11,bool,bool)))) | s(bool,'T_0') = s(bool,sK17(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1054,f608])).
fof(f10270,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,sK17(X22,X23,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X24))) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X22,X23,bool,bool)))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(forward_demodulation,[],[f10255,f877])).
fof(f10255,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,f(s(bool,X24))) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X22,X23,bool,bool)))) | s(bool,sK17(X22,X23,bool,bool)) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f1054,f6906])).
fof(f10256,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,X27))) | s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X25,X26,bool,bool)))) | s(bool,X27) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,sK17(X25,X26,bool,bool))) )),
  inference(superposition,[],[f1054,f6906])).
fof(f1648,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f1648_D])).
fof(f1648_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f12770,plain,(
  ( ! [X57,X56] : (sP0(X57,f(s(bool,X56))) | s(bool,'F_0') = s(bool,X56)) ) | $spl92),
  inference(cnf_transformation,[],[f12770_D])).
fof(f12770_D,plain,(
  ( ! [X57,X56] : (sP0(X57,f(s(bool,X56))) | s(bool,'F_0') = s(bool,X56)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f12771,plain,(
  ( ! [X68,X66,X67] : (sP0(X67,f(s(bool,X66))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X66)) ) | ($spl57 | $spl74)),
  inference(subsumption_resolution,[],[f12739,f7908])).
fof(f12739,plain,(
  ( ! [X68,X66,X67] : (sP0(X67,f(s(bool,X66))) | s(bool,'T_0') = s(bool,X68) | s(bool,X68) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X66)) ) | ($spl57 | $spl74)),
  inference(superposition,[],[f11411,f8040])).
fof(f11411,plain,(
  ( ! [X2,X0,X1] : (sP0(X1,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl57),
  inference(forward_demodulation,[],[f11398,f877])).
fof(f11398,plain,(
  ( ! [X2,X0,X1] : (sP0(X1,f(s(bool,X0))) | s(bool,X2) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X0) = s(bool,X2)) ) | $spl57),
  inference(superposition,[],[f10530,f6906])).
fof(f10530,plain,(
  ( ! [X0] : (sP0(X0,f(s(bool,'T_0')))) ) | $spl57),
  inference(unit_resulting_resolution,[],[f522,f2339,f1702])).
fof(f12769,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X58) | sP0(X57,f(s(bool,X56))) | s(bool,'F_0') = s(bool,X56)) ) | ($spl57 | $spl74)),
  inference(duplicate_literal_removal,[],[f12768])).
fof(f12768,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X58) | sP0(X57,f(s(bool,X56))) | s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X56)) ) | ($spl57 | $spl74)),
  inference(forward_demodulation,[],[f12735,f10569])).
fof(f10569,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl57),
  inference(forward_demodulation,[],[f10546,f877])).
fof(f10546,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl57),
  inference(unit_resulting_resolution,[],[f522,f2339,f6906])).
fof(f12735,plain,(
  ( ! [X57,X58,X56] : (sP0(X57,f(s(bool,X56))) | s(bool,'T_0') = s(bool,X58) | s(bool,X58) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X56)) ) | ($spl57 | $spl74)),
  inference(superposition,[],[f11411,f11135])).
fof(f11135,plain,(
  ( ! [X53] : (s(bool,X53) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X53)) ) | ($spl57 | $spl74)),
  inference(subsumption_resolution,[],[f11079,f275])).
fof(f11079,plain,(
  ( ! [X53] : (~p(s(bool,'T_0')) | s(bool,X53) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X53)) ) | ($spl57 | $spl74)),
  inference(superposition,[],[f9608,f10569])).
fof(f9877,plain,(
  ( ! [X38] : (sP1('F_0',X38) | s(bool,'F_0') = s(bool,X38)) ) | $spl90),
  inference(cnf_transformation,[],[f9877_D])).
fof(f9877_D,plain,(
  ( ! [X38] : (sP1('F_0',X38) | s(bool,'F_0') = s(bool,X38)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f9876,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP1('F_0',X38) | s(bool,'F_0') = s(bool,X38)) ) | $spl74),
  inference(subsumption_resolution,[],[f9860,f294])).
fof(f9860,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP1('F_0',X38) | s(bool,'F_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) ) | $spl74),
  inference(superposition,[],[f9616,f283])).
fof(f9817,plain,(
  ( ! [X38] : (sP1(X38,'F_0') | s(bool,'F_0') = s(bool,X38)) ) | $spl88),
  inference(cnf_transformation,[],[f9817_D])).
fof(f9817_D,plain,(
  ( ! [X38] : (sP1(X38,'F_0') | s(bool,'F_0') = s(bool,X38)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f9816,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP1(X38,'F_0') | s(bool,'F_0') = s(bool,X38)) ) | $spl74),
  inference(subsumption_resolution,[],[f9800,f294])).
fof(f9800,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP1(X38,'F_0') | s(bool,'F_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) ) | $spl74),
  inference(superposition,[],[f9614,f283])).
fof(f2339,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl57),
  inference(cnf_transformation,[],[f2339_D])).
fof(f2339_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f6475,plain,(
  ( ! [X138,X137] : (s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl68),
  inference(cnf_transformation,[],[f6475_D])).
fof(f6475_D,plain,(
  ( ! [X138,X137] : (s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9727,plain,(
  ( ! [X136,X135] : (s(fun(bool,bool),X135) = s(fun(bool,bool),X136) | s(bool,'F_0') != s(bool,f(s(bool,'T_0')))) ) | ($spl44 | $spl71)),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1975,f1995,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2302,f2303,f2304,f2305,f2306,f2308,f2326,f2327,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3987,f3990,f3991,f3992,f3994,f5966,f5968,f5969,f5970,f5972,f6453,f6454,f6455,f6456,f6457,f6458,f6459,f6460,f6461,f6462,f6463,f6464,f6465,f6466,f6467,f6468,f6469,f6470,f6471,f6472,f6477,f6631,f6630,f6629,f6582,f6581,f6634,f6646,f6707,f6709,f6717,f6720,f6731,f6733,f6762,f6775,f6822,f6827,f6838,f6905,f6906,f6907,f6920,f6919,f6914,f6913,f6984,f6985,f6986,f6987,f6990,f7028,f7038,f7041,f7063,f7066,f7078,f7103,f7114,f7126,f8007,f7996,f7990,f8041,f7198,f7199,f7981,f8193,f8194,f8195,f7682,f7878,f7880,f7908,f1094,f2557,f2171,f1249,f7092,f7054,f6032,f6001,f2119,f356,f355,f6446,f6452,f2416,f2173,f348,f347,f400,f399,f398,f1067,f1722,f8352,f7095,f7057,f6035,f6004,f2413,f2360,f2358,f2213,f1449,f1339,f1338,f1448,f2341,f7616,f7873,f8301,f7689,f7004,f7001,f6921,f6656,f6451,f2362,f2172,f1243,f1054,f1053,f1026,f498,f299,f298,f6884,f6811,f5963,f3984,f3912,f2401,f2396,f2395,f2394,f2393,f2392,f2391,f2390,f2389,f2388,f2387,f2386,f2385,f2384,f2383,f2382,f2381,f2380,f2379,f2378,f2377,f2376,f2374,f2373,f2372,f2371,f2370,f2368,f2366,f2364,f2406,f2405,f2414,f2412,f2237,f2274,f2415,f2411,f2241,f2270,f868,f1018,f2402,f2403,f2357,f7879,f8088,f9726,f6441])).
fof(f6441,plain,(
  ( ! [X136,X135] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X135) = s(fun(bool,bool),X136) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X135,X136,bool,bool))))) ) | $spl44),
  inference(duplicate_literal_removal,[],[f6438])).
fof(f6438,plain,(
  ( ! [X136,X135] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X135) = s(fun(bool,bool),X136) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X135,X136,bool,bool)))) | s(fun(bool,bool),X135) = s(fun(bool,bool),X136)) ) | $spl44),
  inference(superposition,[],[f1054,f2308])).
fof(f6451,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,sK17(X6,X7,bool,bool)) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | $spl44),
  inference(subsumption_resolution,[],[f6402,f275])).
fof(f6402,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,sK17(X6,X7,bool,bool)) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | $spl44),
  inference(superposition,[],[f298,f2308])).
fof(f6921,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6656,f608])).
fof(f7689,plain,(
  ( ! [X182,X183] : (s(bool,'F_0') = s(bool,g(s(bool,X182))) | s(bool,X183) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X182) = s(bool,X183)) ) | $spl71),
  inference(superposition,[],[f6656,f6906])).
fof(f6452,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,sK17(X16,X17,bool,bool)) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17)) ) | $spl44),
  inference(subsumption_resolution,[],[f6406,f275])).
fof(f6406,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,sK17(X16,X17,bool,bool)) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17)) ) | $spl44),
  inference(superposition,[],[f360,f2308])).
fof(f6446,plain,(
  ( ! [X2,X3] : (p(s(bool,sK17(X2,X3,bool,bool))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3)) ) | $spl44),
  inference(trivial_inequality_removal,[],[f6400])).
fof(f6400,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,sK17(X2,X3,bool,bool))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3)) ) | $spl44),
  inference(superposition,[],[f294,f2308])).
fof(f7682,plain,(
  ( ! [X169,X168] : (s(bool,X168) != s(bool,g(s(bool,X168))) | s(bool,X169) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X168) = s(bool,X169)) ) | $spl71),
  inference(superposition,[],[f6477,f6906])).
fof(f8041,plain,(
  ( ! [X35] : (s(bool,'F_0') = s(bool,X35) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl71),
  inference(subsumption_resolution,[],[f7191,f6477])).
fof(f7191,plain,(
  ( ! [X35] : (s(bool,'F_0') = s(bool,X35) | s(bool,X35) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl71),
  inference(superposition,[],[f6906,f6656])).
fof(f7114,plain,(
  ( ! [X0,X1] : (sP4('T_0',g(s(bool,X0)),'F_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f7063,f608])).
fof(f7103,plain,(
  ( ! [X0,X1] : (~sP2(g(s(bool,X0)),'F_0','F_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f7038,f608])).
fof(f7078,plain,(
  ( ! [X0,X1] : (~sP2('F_0',g(s(bool,X0)),'F_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6987,f608])).
fof(f7066,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0',g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6920,f608])).
fof(f7063,plain,(
  sP4('T_0',g(s(bool,'T_0')),'F_0') | $spl71),
  inference(unit_resulting_resolution,[],[f6920,f410])).
fof(f7041,plain,(
  ( ! [X0,X1] : (~sP5('F_0','F_0',g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6919,f608])).
fof(f7038,plain,(
  ~sP2(g(s(bool,'T_0')),'F_0','F_0') | $spl71),
  inference(unit_resulting_resolution,[],[f498,f509,f506,f6919,f413])).
fof(f7028,plain,(
  ( ! [X0,X1] : (sP5('T_0',g(s(bool,X0)),'F_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6914,f608])).
fof(f6990,plain,(
  ( ! [X0,X1] : (~sP5('F_0',g(s(bool,X0)),'F_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6913,f608])).
fof(f6987,plain,(
  ~sP2('F_0',g(s(bool,'T_0')),'F_0') | $spl71),
  inference(unit_resulting_resolution,[],[f498,f508,f506,f6913,f414])).
fof(f6986,plain,(
  ( ! [X52,X51] : (sP1(g(s(bool,'T_0')),X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl71),
  inference(subsumption_resolution,[],[f6970,f498])).
fof(f6970,plain,(
  ( ! [X52,X51] : (p(s(bool,'F_0')) | sP1(g(s(bool,'T_0')),X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl71),
  inference(superposition,[],[f1705,f6656])).
fof(f6985,plain,(
  ( ! [X50,X49] : (sP1(X49,g(s(bool,'T_0'))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | $spl71),
  inference(subsumption_resolution,[],[f6969,f498])).
fof(f6969,plain,(
  ( ! [X50,X49] : (p(s(bool,'F_0')) | sP1(X49,g(s(bool,'T_0'))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | $spl71),
  inference(superposition,[],[f1704,f6656])).
fof(f6984,plain,(
  ( ! [X45,X44] : (s(bool,X44) != s(bool,g(s(bool,'T_0'))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl71),
  inference(subsumption_resolution,[],[f6964,f498])).
fof(f6964,plain,(
  ( ! [X45,X44] : (p(s(bool,'F_0')) | s(bool,X44) != s(bool,g(s(bool,'T_0'))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl71),
  inference(superposition,[],[f1381,f6656])).
fof(f6913,plain,(
  ~sP5('F_0',g(s(bool,'T_0')),'F_0') | $spl71),
  inference(unit_resulting_resolution,[],[f498,f6656,f430])).
fof(f6914,plain,(
  sP5('T_0',g(s(bool,'T_0')),'F_0') | $spl71),
  inference(unit_resulting_resolution,[],[f275,f6656,f428])).
fof(f6919,plain,(
  ~sP5('F_0','F_0',g(s(bool,'T_0'))) | $spl71),
  inference(unit_resulting_resolution,[],[f498,f6656,f430])).
fof(f6920,plain,(
  sP5('T_0','F_0',g(s(bool,'T_0'))) | $spl71),
  inference(unit_resulting_resolution,[],[f275,f6656,f428])).
fof(f6907,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) = s(bool,sK17(X39,X40,bool,bool)) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40)) ) | $spl44),
  inference(subsumption_resolution,[],[f6883,f275])).
fof(f6883,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | s(bool,X41) = s(bool,sK17(X39,X40,bool,bool)) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40)) ) | $spl44),
  inference(superposition,[],[f1701,f2308])).
fof(f6838,plain,(
  ( ! [X0,X1] : (sP3('T_0',g(s(bool,X0)),'F_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6731,f608])).
fof(f6827,plain,(
  ( ! [X0,X1] : (~sP4('T_0','T_0',g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6717,f608])).
fof(f6775,plain,(
  ( ! [X0,X1] : (sP3(g(s(bool,X0)),'T_0','F_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6707,f608])).
fof(f6762,plain,(
  ( ! [X0,X1] : (~sP5('T_0','T_0',g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6631,f608])).
fof(f6733,plain,(
  ( ! [X0,X1] : (sP5('F_0','T_0',g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6630,f608])).
fof(f6720,plain,(
  ( ! [X0,X1] : (~sP5('T_0',g(s(bool,X0)),'T_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6582,f608])).
fof(f6709,plain,(
  ( ! [X0,X1] : (sP5('F_0',g(s(bool,X0)),'T_0') | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6581,f608])).
fof(f6646,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl71),
  inference(trivial_inequality_removal,[],[f6640])).
fof(f6640,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl71),
  inference(superposition,[],[f6477,f608])).
fof(f6634,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl71),
  inference(superposition,[],[f6477,f608])).
fof(f6629,plain,(
  ~p(s(bool,g(s(bool,'T_0')))) | $spl71),
  inference(unit_resulting_resolution,[],[f275,f6477,f360])).
fof(f6472,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8(sK17(X91,X92,bool,bool),X93,X94) | sP9(sK17(X91,X92,bool,bool),X93,X94) | s(fun(bool,bool),X91) = s(fun(bool,bool),X92)) ) | $spl44),
  inference(subsumption_resolution,[],[f6426,f275])).
fof(f6426,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~sP8(sK17(X91,X92,bool,bool),X93,X94) | sP9(sK17(X91,X92,bool,bool),X93,X94) | s(fun(bool,bool),X91) = s(fun(bool,bool),X92)) ) | $spl44),
  inference(superposition,[],[f483,f2308])).
fof(f6471,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10(sK17(X87,X88,bool,bool),X89,X90) | sP11(sK17(X87,X88,bool,bool),X89,X90) | s(fun(bool,bool),X87) = s(fun(bool,bool),X88)) ) | $spl44),
  inference(subsumption_resolution,[],[f6425,f275])).
fof(f6425,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP10(sK17(X87,X88,bool,bool),X89,X90) | sP11(sK17(X87,X88,bool,bool),X89,X90) | s(fun(bool,bool),X87) = s(fun(bool,bool),X88)) ) | $spl44),
  inference(superposition,[],[f482,f2308])).
fof(f6470,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11(sK17(X83,X84,bool,bool),X86,X85) | s(fun(bool,bool),X83) = s(fun(bool,bool),X84)) ) | $spl44),
  inference(subsumption_resolution,[],[f6424,f275])).
fof(f6424,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11(sK17(X83,X84,bool,bool),X86,X85) | s(fun(bool,bool),X83) = s(fun(bool,bool),X84)) ) | $spl44),
  inference(superposition,[],[f479,f2308])).
fof(f6469,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X81)) | sP11(sK17(X79,X80,bool,bool),X81,X82) | s(fun(bool,bool),X79) = s(fun(bool,bool),X80)) ) | $spl44),
  inference(subsumption_resolution,[],[f6423,f275])).
fof(f6423,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | sP11(sK17(X79,X80,bool,bool),X81,X82) | s(fun(bool,bool),X79) = s(fun(bool,bool),X80)) ) | $spl44),
  inference(superposition,[],[f478,f2308])).
fof(f6468,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X77)) | sP11(sK17(X75,X76,bool,bool),X78,X77) | s(fun(bool,bool),X75) = s(fun(bool,bool),X76)) ) | $spl44),
  inference(subsumption_resolution,[],[f6422,f275])).
fof(f6422,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X77)) | sP11(sK17(X75,X76,bool,bool),X78,X77) | s(fun(bool,bool),X75) = s(fun(bool,bool),X76)) ) | $spl44),
  inference(superposition,[],[f477,f2308])).
fof(f6467,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X73)) | sP11(X74,X73,sK17(X71,X72,bool,bool)) | p(s(bool,X74)) | s(fun(bool,bool),X71) = s(fun(bool,bool),X72)) ) | $spl44),
  inference(subsumption_resolution,[],[f6421,f275])).
fof(f6421,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X73)) | sP11(X74,X73,sK17(X71,X72,bool,bool)) | p(s(bool,X74)) | s(fun(bool,bool),X71) = s(fun(bool,bool),X72)) ) | $spl44),
  inference(superposition,[],[f476,f2308])).
fof(f6466,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,sK17(X67,X68,bool,bool)) | s(fun(bool,bool),X67) = s(fun(bool,bool),X68)) ) | $spl44),
  inference(subsumption_resolution,[],[f6420,f275])).
fof(f6420,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,sK17(X67,X68,bool,bool)) | s(fun(bool,bool),X67) = s(fun(bool,bool),X68)) ) | $spl44),
  inference(superposition,[],[f472,f2308])).
fof(f6465,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | p(s(bool,X66)) | ~sP9(sK17(X63,X64,bool,bool),X66,X65) | s(fun(bool,bool),X63) = s(fun(bool,bool),X64)) ) | $spl44),
  inference(subsumption_resolution,[],[f6419,f275])).
fof(f6419,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP9(sK17(X63,X64,bool,bool),X66,X65) | s(fun(bool,bool),X63) = s(fun(bool,bool),X64)) ) | $spl44),
  inference(superposition,[],[f462,f2308])).
fof(f6464,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | sP9(sK17(X59,X60,bool,bool),X61,X62) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60)) ) | $spl44),
  inference(subsumption_resolution,[],[f6418,f275])).
fof(f6418,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | sP9(sK17(X59,X60,bool,bool),X61,X62) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60)) ) | $spl44),
  inference(superposition,[],[f461,f2308])).
fof(f6463,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP9(sK17(X55,X56,bool,bool),X58,X57) | s(fun(bool,bool),X55) = s(fun(bool,bool),X56)) ) | $spl44),
  inference(subsumption_resolution,[],[f6417,f275])).
fof(f6417,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9(sK17(X55,X56,bool,bool),X58,X57) | s(fun(bool,bool),X55) = s(fun(bool,bool),X56)) ) | $spl44),
  inference(superposition,[],[f460,f2308])).
fof(f6462,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(sK17(X51,X52,bool,bool),X53,X54) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52)) ) | $spl44),
  inference(subsumption_resolution,[],[f6416,f275])).
fof(f6416,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(sK17(X51,X52,bool,bool),X53,X54) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52)) ) | $spl44),
  inference(superposition,[],[f455,f2308])).
fof(f6461,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,sK17(X47,X48,bool,bool)) | s(fun(bool,bool),X47) = s(fun(bool,bool),X48)) ) | $spl44),
  inference(subsumption_resolution,[],[f6415,f275])).
fof(f6415,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,sK17(X47,X48,bool,bool)) | s(fun(bool,bool),X47) = s(fun(bool,bool),X48)) ) | $spl44),
  inference(superposition,[],[f447,f2308])).
fof(f6460,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | sP7(sK17(X43,X44,bool,bool),X45,X46) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44)) ) | $spl44),
  inference(subsumption_resolution,[],[f6414,f275])).
fof(f6414,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | sP7(sK17(X43,X44,bool,bool),X45,X46) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44)) ) | $spl44),
  inference(superposition,[],[f444,f2308])).
fof(f6459,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,sK17(X39,X40,bool,bool),X42) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40)) ) | $spl44),
  inference(subsumption_resolution,[],[f6413,f275])).
fof(f6413,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,sK17(X39,X40,bool,bool),X42) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40)) ) | $spl44),
  inference(superposition,[],[f438,f2308])).
fof(f6458,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,sK17(X35,X36,bool,bool),X38) | s(fun(bool,bool),X35) = s(fun(bool,bool),X36)) ) | $spl44),
  inference(subsumption_resolution,[],[f6412,f275])).
fof(f6412,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,sK17(X35,X36,bool,bool),X38) | s(fun(bool,bool),X35) = s(fun(bool,bool),X36)) ) | $spl44),
  inference(superposition,[],[f419,f2308])).
fof(f6457,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,sK17(X31,X32,bool,bool),X34) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32)) ) | $spl44),
  inference(subsumption_resolution,[],[f6411,f275])).
fof(f6411,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,sK17(X31,X32,bool,bool),X34) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32)) ) | $spl44),
  inference(superposition,[],[f415,f2308])).
fof(f6456,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,sK17(X28,X29,bool,bool)) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29)) ) | $spl44),
  inference(subsumption_resolution,[],[f6410,f275])).
fof(f6410,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,sK17(X28,X29,bool,bool)) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29)) ) | $spl44),
  inference(superposition,[],[f385,f2308])).
fof(f6455,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1(sK17(X25,X26,bool,bool),X27) | s(fun(bool,bool),X25) = s(fun(bool,bool),X26)) ) | $spl44),
  inference(subsumption_resolution,[],[f6409,f275])).
fof(f6409,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1(sK17(X25,X26,bool,bool),X27) | s(fun(bool,bool),X25) = s(fun(bool,bool),X26)) ) | $spl44),
  inference(superposition,[],[f382,f2308])).
fof(f6454,plain,(
  ( ! [X24,X23,X22] : (sP0(sK17(X22,X23,bool,bool),X24) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23)) ) | $spl44),
  inference(subsumption_resolution,[],[f6408,f275])).
fof(f6408,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(sK17(X22,X23,bool,bool),X24) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23)) ) | $spl44),
  inference(superposition,[],[f367,f2308])).
fof(f6453,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,sK17(X19,X20,bool,bool)) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20)) ) | $spl44),
  inference(subsumption_resolution,[],[f6407,f275])).
fof(f6407,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,sK17(X19,X20,bool,bool)) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20)) ) | $spl44),
  inference(superposition,[],[f366,f2308])).
fof(f9726,plain,(
  ( ! [X138,X137] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | ($spl44 | $spl71)),
  inference(forward_demodulation,[],[f6440,f6656])).
fof(f6440,plain,(
  ( ! [X138,X137] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl44),
  inference(duplicate_literal_removal,[],[f6439])).
fof(f6439,plain,(
  ( ! [X138,X137] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | $spl44),
  inference(superposition,[],[f401,f2308])).
fof(f2337,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool))))) ) | $spl54),
  inference(cnf_transformation,[],[f2337_D])).
fof(f2337_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8088,plain,(
  ( ! [X200,X199,X201] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X199) = s(fun(bool,bool),X200) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X199,X200,bool,bool)))) | s(bool,X201) = s(bool,sK17(X199,X200,bool,bool)) | s(bool,'T_0') = s(bool,X201)) )),
  inference(forward_demodulation,[],[f7279,f877])).
fof(f7279,plain,(
  ( ! [X200,X199,X201] : (s(bool,'F_0') != s(bool,f(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))))) | s(fun(bool,bool),X199) = s(fun(bool,bool),X200) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X199,X200,bool,bool)))) | s(bool,X201) = s(bool,sK17(X199,X200,bool,bool)) | s(bool,'T_0') = s(bool,X201)) )),
  inference(superposition,[],[f1054,f6906])).
fof(f9725,plain,(
  ( ! [X341,X339,X340] : (s(fun(bool,bool),X339) = s(fun(bool,bool),X340) | s(bool,X341) = s(bool,sK17(X339,X340,bool,bool)) | s(bool,X341) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X339,X340,bool,bool))))) ) | $spl86),
  inference(cnf_transformation,[],[f9725_D])).
fof(f9725_D,plain,(
  ( ! [X341,X339,X340] : (s(fun(bool,bool),X339) = s(fun(bool,bool),X340) | s(bool,X341) = s(bool,sK17(X339,X340,bool,bool)) | s(bool,X341) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X339,X340,bool,bool))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f7879,plain,(
  ( ! [X341,X339,X340] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X339) = s(fun(bool,bool),X340) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X339,X340,bool,bool)))) | s(bool,X341) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X341) = s(bool,sK17(X339,X340,bool,bool))) )),
  inference(superposition,[],[f1054,f6906])).
fof(f9723,plain,(
  ( ! [X107,X106] : (sP9(f(s(bool,X106)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X107) | s(bool,'F_0') = s(bool,X106)) ) | $spl84),
  inference(cnf_transformation,[],[f9723_D])).
fof(f9723_D,plain,(
  ( ! [X107,X106] : (sP9(f(s(bool,X106)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X107) | s(bool,'F_0') = s(bool,X106)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9722,plain,(
  ( ! [X107,X108,X106] : (s(bool,'T_0') = s(bool,X108) | sP9(f(s(bool,X106)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X107) | s(bool,'F_0') = s(bool,X106)) ) | ($spl57 | $spl74)),
  inference(forward_demodulation,[],[f9721,f877])).
fof(f9721,plain,(
  ( ! [X107,X108,X106] : (sP9(f(s(bool,X106)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X107) | s(bool,X108) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X106)) ) | ($spl57 | $spl74)),
  inference(subsumption_resolution,[],[f9574,f8193])).
fof(f9574,plain,(
  ( ! [X107,X108,X106] : (sP9(f(s(bool,X106)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X107) | s(bool,X108) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X108) | s(bool,'F_0') = s(bool,X106)) ) | ($spl57 | $spl74)),
  inference(superposition,[],[f3416,f8040])).
fof(f3416,plain,(
  ( ! [X2,X0,X1] : (sP9(f(s(bool,X0)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl57),
  inference(superposition,[],[f3389,f608])).
fof(f3389,plain,(
  ( ! [X0] : (sP9(f(s(bool,'T_0')),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) ) | $spl57),
  inference(unit_resulting_resolution,[],[f962,f3016])).
fof(f3016,plain,(
  ( ! [X37,X38] : (sP9(f(s(bool,'T_0')),X37,X38) | ~sP8(f(s(bool,'T_0')),X37,X38)) ) | $spl57),
  inference(subsumption_resolution,[],[f2972,f275])).
fof(f2972,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | ~sP8(f(s(bool,'T_0')),X37,X38) | sP9(f(s(bool,'T_0')),X37,X38)) ) | $spl57),
  inference(superposition,[],[f483,f2692])).
fof(f2692,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl57),
  inference(unit_resulting_resolution,[],[f498,f2339,f654])).
fof(f9720,plain,(
  ( ! [X103,X104] : (sP9(f(s(bool,X103)),X104,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X103)) ) | $spl82),
  inference(cnf_transformation,[],[f9720_D])).
fof(f9720_D,plain,(
  ( ! [X103,X104] : (sP9(f(s(bool,X103)),X104,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X103)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f9719,plain,(
  ( ! [X103,X105,X104] : (s(bool,'T_0') = s(bool,X105) | sP9(f(s(bool,X103)),X104,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X103)) ) | ($spl57 | $spl74)),
  inference(forward_demodulation,[],[f9718,f877])).
fof(f9718,plain,(
  ( ! [X103,X105,X104] : (sP9(f(s(bool,X103)),X104,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X105) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X103)) ) | ($spl57 | $spl74)),
  inference(subsumption_resolution,[],[f9573,f8193])).
fof(f9573,plain,(
  ( ! [X103,X105,X104] : (sP9(f(s(bool,X103)),X104,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X105) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X105) | s(bool,'F_0') = s(bool,X103)) ) | ($spl57 | $spl74)),
  inference(superposition,[],[f3400,f8040])).
fof(f3400,plain,(
  ( ! [X2,X0,X1] : (sP9(f(s(bool,X0)),X1,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl57),
  inference(superposition,[],[f3387,f608])).
fof(f3387,plain,(
  ( ! [X0] : (sP9(f(s(bool,'T_0')),X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl57),
  inference(unit_resulting_resolution,[],[f961,f3016])).
fof(f9717,plain,(
  ( ! [X101,X100] : (sP11(f(s(bool,X100)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X101) | s(bool,'F_0') = s(bool,X100)) ) | $spl80),
  inference(cnf_transformation,[],[f9717_D])).
fof(f9717_D,plain,(
  ( ! [X101,X100] : (sP11(f(s(bool,X100)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X101) | s(bool,'F_0') = s(bool,X100)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f9716,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,X102) | sP11(f(s(bool,X100)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X101) | s(bool,'F_0') = s(bool,X100)) ) | ($spl57 | $spl74)),
  inference(forward_demodulation,[],[f9715,f877])).
fof(f9715,plain,(
  ( ! [X101,X102,X100] : (sP11(f(s(bool,X100)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X101) | s(bool,X102) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X100)) ) | ($spl57 | $spl74)),
  inference(subsumption_resolution,[],[f9572,f8193])).
fof(f9572,plain,(
  ( ! [X101,X102,X100] : (sP11(f(s(bool,X100)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X101) | s(bool,X102) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X100)) ) | ($spl57 | $spl74)),
  inference(superposition,[],[f3366,f8040])).
fof(f3366,plain,(
  ( ! [X2,X0,X1] : (sP11(f(s(bool,X0)),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl57),
  inference(superposition,[],[f3345,f608])).
fof(f3345,plain,(
  ( ! [X0] : (sP11(f(s(bool,'T_0')),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) ) | $spl57),
  inference(unit_resulting_resolution,[],[f994,f3015])).
fof(f3015,plain,(
  ( ! [X35,X36] : (sP11(f(s(bool,'T_0')),X35,X36) | ~sP10(f(s(bool,'T_0')),X35,X36)) ) | $spl57),
  inference(subsumption_resolution,[],[f2971,f275])).
fof(f2971,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~sP10(f(s(bool,'T_0')),X35,X36) | sP11(f(s(bool,'T_0')),X35,X36)) ) | $spl57),
  inference(superposition,[],[f482,f2692])).
fof(f8526,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f8525,f8508])).
fof(f8508,plain,(
  ( ! [X8] : (s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl76),
  inference(subsumption_resolution,[],[f8473,f8465])).
fof(f8465,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,X25)) ) | $spl76),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f355,f356,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1094,f1249,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1339,f1449,f400,f399,f398,f1067,f1722,f868,f1018,f348,f347,f1338,f1448,f1243,f1054,f1053,f1026,f498,f299,f298,f1975,f1995,f2119,f2171,f2172,f2173,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2213,f2237,f2241,f2270,f2274,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2341,f2357,f2358,f2360,f2362,f2416,f2364,f2366,f2368,f2370,f2371,f2372,f2373,f2374,f2376,f2377,f2378,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2386,f2387,f2388,f2389,f2390,f2391,f2392,f2393,f2394,f2395,f2396,f2401,f2402,f2403,f2405,f2406,f2411,f2412,f2413,f2414,f2415,f2557,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3912,f3987,f3990,f3991,f3992,f3994,f3984,f5966,f5968,f5969,f5970,f5972,f5963,f6001,f6004,f6032,f6035,f6822,f6811,f6905,f6906,f6884,f7001,f7004,f7054,f7057,f7092,f7095,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7616,f8301,f8352,f7873,f7878,f7880,f7908,f8357,f8418,f8420,f8433,f8438,f8440,f8448,f8450,f8452,f8453,f8456,f8458,f8460,f8464])).
fof(f8464,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,happ(s(fun(X23,bool),X24),s(X23,sK16(X24,X23,bool)))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X23,bool),X24)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8411,f8357])).
fof(f8411,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X25) = s(bool,happ(s(fun(X23,bool),X24),s(X23,sK16(X24,X23,bool)))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X23,bool),X24)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f3912])).
fof(f8460,plain,(
  ( ! [X80,X81,X79,X82] : (sP11(happ(s(fun(X79,bool),X80),s(X79,sK16(X80,X79,bool))),X81,X82) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X79,bool),X80)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8459,f8453])).
fof(f8459,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP11(happ(s(fun(X79,bool),X80),s(X79,sK16(X80,X79,bool))),X81,X82) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X79,bool),X80)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8401,f8453])).
fof(f8401,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X81)) | sP11(happ(s(fun(X79,bool),X80),s(X79,sK16(X80,X79,bool))),X81,X82) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X79,bool),X80)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2393])).
fof(f8458,plain,(
  ( ! [X57,X58,X56,X55] : (sP9(happ(s(fun(X55,bool),X56),s(X55,sK16(X56,X55,bool))),X58,X57) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X55,bool),X56)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8457,f8453])).
fof(f8457,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP9(happ(s(fun(X55,bool),X56),s(X55,sK16(X56,X55,bool))),X58,X57) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X55,bool),X56)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8395,f8453])).
fof(f8395,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X57)) | sP9(happ(s(fun(X55,bool),X56),s(X55,sK16(X56,X55,bool))),X58,X57) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X55,bool),X56)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2387])).
fof(f8456,plain,(
  ( ! [X45,X43,X46,X44] : (sP7(happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool))),X45,X46) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8455,f8453])).
fof(f8455,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP7(happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool))),X45,X46) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8454,f8453])).
fof(f8454,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X45)) | sP7(happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool))),X45,X46) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8392,f8453])).
fof(f8392,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | sP7(happ(s(fun(X43,bool),X44),s(X43,sK16(X44,X43,bool))),X45,X46) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X43,bool),X44)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2384])).
fof(f8453,plain,(
  ( ! [X27] : (p(s(bool,X27))) ) | $spl76),
  inference(global_subsumption,[],[f274,f273,f272,f275,f283,f295,f294,f355,f356,f360,f357,f367,f366,f499,f385,f382,f381,f380,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1020,f1071,f1076,f1085,f1090,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1094,f1249,f1252,f1251,f1257,f1281,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1339,f1449,f400,f399,f398,f1067,f1722,f868,f1018,f348,f347,f1338,f1448,f1243,f1054,f1053,f1026,f498,f299,f298,f1975,f1995,f2119,f2171,f2172,f2173,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2213,f2237,f2241,f2270,f2274,f2302,f2303,f2304,f2305,f2306,f2326,f2327,f2341,f2357,f2358,f2360,f2362,f2416,f2364,f2366,f2368,f2370,f2371,f2372,f2373,f2374,f2376,f2377,f2378,f2379,f2380,f2381,f2382,f2383,f2384,f2385,f2386,f2387,f2388,f2389,f2390,f2391,f2392,f2393,f2394,f2395,f2396,f2401,f2402,f2403,f2405,f2406,f2411,f2412,f2413,f2414,f2415,f2557,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2639,f2645,f3237,f3265,f3264,f3294,f3311,f3328,f3327,f3357,f3356,f3355,f3384,f3412,f3451,f3680,f3808,f3896,f3912,f3987,f3990,f3991,f3992,f3994,f3984,f5966,f5968,f5969,f5970,f5972,f5963,f6001,f6004,f6032,f6035,f6822,f6811,f6905,f6906,f6884,f7001,f7004,f7054,f7057,f7092,f7095,f7126,f8007,f7996,f7990,f7198,f7199,f7981,f8193,f8194,f8195,f7616,f8301,f8352,f7873,f7878,f7880,f7908,f8357,f8418,f8420,f8433,f8438,f8440,f8448,f8450,f8452])).
fof(f8452,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1(happ(s(fun(X25,bool),X26),s(X25,sK16(X26,X25,bool))),X27) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X25,bool),X26)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8451,f275])).
fof(f8451,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1(happ(s(fun(X25,bool),X26),s(X25,sK16(X26,X25,bool))),X27) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X25,bool),X26)))) ) | $spl76),
  inference(forward_demodulation,[],[f8387,f877])).
fof(f8387,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | p(s(bool,X27)) | sP1(happ(s(fun(X25,bool),X26),s(X25,sK16(X26,X25,bool))),X27) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X25,bool),X26)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2379])).
fof(f8450,plain,(
  ( ! [X24,X23,X22] : (sP0(happ(s(fun(X22,bool),X23),s(X22,sK16(X23,X22,bool))),X24) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X22,bool),X23)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8449,f275])).
fof(f8449,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X22,bool),X23),s(X22,sK16(X23,X22,bool))),X24) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X22,bool),X23)))) ) | $spl76),
  inference(forward_demodulation,[],[f8386,f877])).
fof(f8386,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP0(happ(s(fun(X22,bool),X23),s(X22,sK16(X23,X22,bool))),X24) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X22,bool),X23)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2378])).
fof(f8448,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,happ(s(fun(X19,bool),X20),s(X19,sK16(X20,X19,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X19,bool),X20)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8447,f275])).
fof(f8447,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,happ(s(fun(X19,bool),X20),s(X19,sK16(X20,X19,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X19,bool),X20)))) ) | $spl76),
  inference(forward_demodulation,[],[f8385,f877])).
fof(f8385,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | sP0(X21,happ(s(fun(X19,bool),X20),s(X19,sK16(X20,X19,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X19,bool),X20)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2377])).
fof(f8440,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK16(X7,X6,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X6,bool),X7)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8439,f275])).
fof(f8439,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK16(X7,X6,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X6,bool),X7)))) ) | $spl76),
  inference(forward_demodulation,[],[f8381,f877])).
fof(f8381,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK16(X7,X6,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X6,bool),X7)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2372])).
fof(f8438,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,sK16(X5,X4,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X4,bool),X5)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8437,f275])).
fof(f8437,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,sK16(X5,X4,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X4,bool),X5)))) ) | $spl76),
  inference(forward_demodulation,[],[f8380,f877])).
fof(f8380,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,sK16(X5,X4,bool)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X4,bool),X5)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2371])).
fof(f8433,plain,(
  ( ! [X123,X124] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat',happ(s(fun(X123,'type.rat.rat'),X124),s(X123,sK15(X124,X123,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X123,'type.rat.rat'),X124)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8432,f275])).
fof(f8432,plain,(
  ( ! [X123,X124] : (~p(s(bool,'T_0')) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat',happ(s(fun(X123,'type.rat.rat'),X124),s(X123,sK15(X124,X123,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X123,'type.rat.rat'),X124)))) ) | $spl76),
  inference(forward_demodulation,[],[f8375,f877])).
fof(f8375,plain,(
  ( ! [X123,X124] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat',happ(s(fun(X123,'type.rat.rat'),X124),s(X123,sK15(X124,X123,'type.rat.rat')))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X123,'type.rat.rat'),X124)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2406])).
fof(f8420,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8419,f275])).
fof(f8419,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) ) | $spl76),
  inference(forward_demodulation,[],[f8364,f877])).
fof(f8364,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f1018])).
fof(f8418,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (s(X111,happ(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK16(X113,X112,fun(X110,X111))))),s(X110,X114))) != s(X111,happ(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK16(X113,X112,fun(X110,X111))))),s(X110,X115))) | s(X110,X114) = s(X110,X115) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X112,fun(X110,X111)),X113)))) ) | $spl76),
  inference(subsumption_resolution,[],[f8417,f275])).
fof(f8417,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (~p(s(bool,'T_0')) | s(X111,happ(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK16(X113,X112,fun(X110,X111))))),s(X110,X114))) != s(X111,happ(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK16(X113,X112,fun(X110,X111))))),s(X110,X115))) | s(X110,X114) = s(X110,X115) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X112,fun(X110,X111)),X113)))) ) | $spl76),
  inference(forward_demodulation,[],[f8363,f877])).
fof(f8363,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | s(X111,happ(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK16(X113,X112,fun(X110,X111))))),s(X110,X114))) != s(X111,happ(s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),X113),s(X112,sK16(X113,X112,fun(X110,X111))))),s(X110,X115))) | s(X110,X114) = s(X110,X115) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X112,fun(X110,X111)),X113)))) ) | $spl76),
  inference(backward_demodulation,[],[f8357,f2402])).
fof(f8473,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X8))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl76),
  inference(backward_demodulation,[],[f8465,f2172])).
fof(f8525,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl76),
  inference(subsumption_resolution,[],[f8498,f275])).
fof(f8498,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl76),
  inference(backward_demodulation,[],[f8465,f348])).
fof(f8522,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f8520,f8521])).
fof(f8521,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | $spl76),
  inference(subsumption_resolution,[],[f8519,f8508])).
fof(f8519,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | $spl76),
  inference(backward_demodulation,[],[f8508,f2557])).
fof(f8520,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat',sK16('const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'type.rat.rat','type.rat.rat')) | $spl76),
  inference(backward_demodulation,[],[f8508,f1094])).
fof(f8507,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f8468,f275])).
fof(f8468,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(backward_demodulation,[],[f8465,f498])).
fof(f8506,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f8505,f8453])).
fof(f8505,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(subsumption_resolution,[],[f8466,f8465])).
fof(f8466,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl76),
  inference(backward_demodulation,[],[f8465,f298])).
fof(f8357,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl76),
  inference(cnf_transformation,[],[f8357_D])).
fof(f8357_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f8358,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl78),
  inference(cnf_transformation,[],[f8358_D])).
fof(f8358_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7880,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X0) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f6906])).
fof(f8032,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X17) = s(bool,X19) | s(bool,X19) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl72),
  inference(cnf_transformation,[],[f8032_D])).
fof(f8032_D,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X17) = s(bool,X19) | s(bool,X19) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7981,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(duplicate_literal_removal,[],[f7212])).
fof(f7212,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f6906,f6906])).
fof(f7990,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X25) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) )),
  inference(duplicate_literal_removal,[],[f7185])).
fof(f7185,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X25) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f6906,f6906])).
fof(f8040,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X1)) ) | $spl74),
  inference(cnf_transformation,[],[f8040_D])).
fof(f8040_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f7996,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f7177])).
fof(f7177,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6906,f283])).
fof(f8007,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X19) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X17) = s(bool,X19)) )),
  inference(duplicate_literal_removal,[],[f7162])).
fof(f7162,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | s(bool,X17) = s(bool,X19)) )),
  inference(superposition,[],[f6906,f6906])).
fof(f6477,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl71),
  inference(cnf_transformation,[],[f6477_D])).
fof(f6477_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f6474,plain,(
  ( ! [X138,X137] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X137) = s(fun(bool,bool),X138)) ) | ($spl44 | $spl57)),
  inference(forward_demodulation,[],[f6440,f2692])).
fof(f5971,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl66),
  inference(cnf_transformation,[],[f5971_D])).
fof(f5971_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5970,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f5958,f294])).
fof(f5958,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1705,f283])).
fof(f5967,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl64),
  inference(cnf_transformation,[],[f5967_D])).
fof(f5967_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5968,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5954,f294])).
fof(f5954,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1705,f608])).
fof(f5966,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f5953,f294])).
fof(f5953,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1705,f608])).
fof(f5921,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f3998,f4011])).
fof(f4011,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f3989,f360])).
fof(f3998,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f298])).
fof(f5920,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4002,f275])).
fof(f4002,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f522,f3989,f360])).
fof(f5919,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4004,f275])).
fof(f4004,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1095,f3989,f360])).
fof(f5918,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4008,f3989])).
fof(f4008,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1095,f3989,f360])).
fof(f5917,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4012,f275])).
fof(f4012,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f522,f3989,f360])).
fof(f5916,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4014,f275])).
fof(f4014,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1095,f3989,f360])).
fof(f5915,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4018,f3989])).
fof(f4018,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1095,f3989,f360])).
fof(f5914,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4025,f3989])).
fof(f4025,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1105,f3989,f381])).
fof(f5913,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4026,f275])).
fof(f4026,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f536,f3989,f381])).
fof(f5912,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4027,f3989])).
fof(f4027,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f886,f3989,f381])).
fof(f5911,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f4028,f3989])).
fof(f4028,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f3160,f3989,f381])).
fof(f3160,plain,(
  sP1('F_0',f(s(bool,'T_0'))) | $spl57),
  inference(unit_resulting_resolution,[],[f498,f3000])).
fof(f3000,plain,(
  ( ! [X6] : (sP1(X6,f(s(bool,'T_0'))) | p(s(bool,X6))) ) | $spl57),
  inference(subsumption_resolution,[],[f2956,f275])).
fof(f2956,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,f(s(bool,'T_0')))) ) | $spl57),
  inference(superposition,[],[f385,f2692])).
fof(f5910,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4031,f275])).
fof(f4031,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1106,f3989,f381])).
fof(f5909,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4036,f275])).
fof(f4036,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f533,f3989,f381])).
fof(f5908,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4037,f275])).
fof(f4037,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1105,f3989,f381])).
fof(f5907,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f4041,f3989])).
fof(f4041,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f3138,f3989,f381])).
fof(f3138,plain,(
  sP1(f(s(bool,'T_0')),'F_0') | $spl57),
  inference(unit_resulting_resolution,[],[f498,f2999])).
fof(f2999,plain,(
  ( ! [X5] : (sP1(f(s(bool,'T_0')),X5) | p(s(bool,X5))) ) | $spl57),
  inference(subsumption_resolution,[],[f2955,f275])).
fof(f2955,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(f(s(bool,'T_0')),X5)) ) | $spl57),
  inference(superposition,[],[f382,f2692])).
fof(f5906,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4042,f3989])).
fof(f4042,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f885,f3989,f381])).
fof(f5905,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4043,f3989])).
fof(f4043,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1106,f3989,f381])).
fof(f5904,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4066,f275])).
fof(f4066,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f663,f498,f3989,f411])).
fof(f5903,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f4067,f3989])).
fof(f4067,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f2689,f498,f3989,f411])).
fof(f2689,plain,(
  sP5('F_0',f(s(bool,'T_0')),'F_0') | $spl57),
  inference(unit_resulting_resolution,[],[f498,f2339,f427])).
fof(f5902,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4068,f3989])).
fof(f4068,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1975,f498,f3989,f411])).
fof(f5901,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4070,f3989])).
fof(f4070,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1109,f498,f3989,f411])).
fof(f5900,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4072,f275])).
fof(f4072,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1141,f498,f3989,f411])).
fof(f5899,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4145,f275])).
fof(f4145,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f847,f498,f3989,f411])).
fof(f5898,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4146,f275])).
fof(f4146,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1109,f498,f3989,f411])).
fof(f5897,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4147,f3989])).
fof(f4147,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1257,f498,f3989,f411])).
fof(f5896,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4148,f3989])).
fof(f4148,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1141,f498,f3989,f411])).
fof(f5895,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4153,f4335])).
fof(f4335,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f416])).
fof(f4153,plain,(
  ~sP4('T_0','F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f508,f673,f3989,f412])).
fof(f5894,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4154,f4335])).
fof(f4154,plain,(
  ~sP4('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f1120,f1107,f3989,f412])).
fof(f5893,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f4158,f4335])).
fof(f4158,plain,(
  ~sP4('T_0','F_0',f(s(bool,'T_0'))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f510,f508,f2688,f3989,f412])).
fof(f2688,plain,(
  ~sP5('T_0',f(s(bool,'T_0')),'F_0') | $spl57),
  inference(unit_resulting_resolution,[],[f275,f2339,f429])).
fof(f5892,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4159,f4335])).
fof(f4159,plain,(
  ~sP4('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f911,f508,f889,f3989,f412])).
fof(f5891,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4161,f4335])).
fof(f4161,plain,(
  ~sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f508,f913,f3989,f412])).
fof(f5890,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4164,f4335])).
fof(f4164,plain,(
  ~sP4('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f911,f508,f890,f3989,f412])).
fof(f5889,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4173,f4441])).
fof(f4441,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f522,f3989,f429])).
fof(f4173,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f508,f504,f3989,f412])).
fof(f5888,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4176,f4443])).
fof(f4443,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1095,f3989,f429])).
fof(f4176,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f1120,f504,f3989,f412])).
fof(f5887,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4179,f4373])).
fof(f4373,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f420])).
fof(f4179,plain,(
  ~sP3('T_0','F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f673,f504,f3989,f412])).
fof(f5886,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4180,f4379])).
fof(f4379,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f426])).
fof(f4180,plain,(
  ~sP2('F_0','T_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f508,f673,f504,f3989,f412])).
fof(f5885,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4182,f4373])).
fof(f4182,plain,(
  ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f1107,f504,f3989,f412])).
fof(f5884,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4183,f4379])).
fof(f4183,plain,(
  ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1120,f1107,f504,f3989,f412])).
fof(f5883,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4188,f4373])).
fof(f4188,plain,(
  ~sP3('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f913,f504,f3989,f412])).
fof(f5882,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4189,f4379])).
fof(f4189,plain,(
  ~sP2('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f508,f913,f504,f3989,f412])).
fof(f5881,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4206,f4373])).
fof(f4206,plain,(
  ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f911,f889,f900,f3989,f412])).
fof(f5880,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4207,f4379])).
fof(f4207,plain,(
  ~sP2('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f508,f889,f900,f3989,f412])).
fof(f5879,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4211,f4373])).
fof(f4211,plain,(
  ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f911,f890,f900,f3989,f412])).
fof(f5878,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4212,f4379])).
fof(f4212,plain,(
  ~sP2('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f508,f890,f900,f3989,f412])).
fof(f5877,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4217,f4441])).
fof(f4217,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f508,f505,f3989,f412])).
fof(f5876,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4220,f673])).
fof(f4220,plain,(
  sP5('T_0','T_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f507,f505,f3989,f412])).
fof(f5875,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4221,f889])).
fof(f4221,plain,(
  sP5('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f911,f507,f505,f3989,f412])).
fof(f5874,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4223,f913])).
fof(f4223,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f909,f505,f3989,f412])).
fof(f5873,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4224,f890])).
fof(f4224,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f911,f909,f505,f3989,f412])).
fof(f5872,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4229,f4373])).
fof(f4229,plain,(
  ~sP3('T_0','F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f673,f505,f3989,f412])).
fof(f5871,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4230,f4379])).
fof(f4230,plain,(
  ~sP2('F_0','T_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f508,f673,f505,f3989,f412])).
fof(f5870,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f4232,f4373])).
fof(f4232,plain,(
  ~sP3(f(s(bool,'T_0')),'F_0','T_0') | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f510,f2688,f505,f3989,f412])).
fof(f5869,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f4233,f4379])).
fof(f4233,plain,(
  ~sP2('F_0',f(s(bool,'T_0')),'T_0') | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f508,f2688,f505,f3989,f412])).
fof(f5868,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4235,f4373])).
fof(f4235,plain,(
  ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f911,f889,f505,f3989,f412])).
fof(f5867,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4236,f4379])).
fof(f4236,plain,(
  ~sP2('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f508,f889,f505,f3989,f412])).
fof(f5866,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4240,f4373])).
fof(f4240,plain,(
  ~sP3('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f913,f505,f3989,f412])).
fof(f5865,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4241,f4379])).
fof(f4241,plain,(
  ~sP2('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f508,f913,f505,f3989,f412])).
fof(f5864,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4243,f4373])).
fof(f4243,plain,(
  ~sP3('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f911,f890,f505,f3989,f412])).
fof(f5863,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4244,f4379])).
fof(f4244,plain,(
  ~sP2('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f508,f890,f505,f3989,f412])).
fof(f5862,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4249,f4443])).
fof(f4249,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f1120,f1115,f3989,f412])).
fof(f5861,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4252,f1107])).
fof(f4252,plain,(
  sP5('T_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f507,f1115,f3989,f412])).
fof(f5860,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4255,f4443])).
fof(f4255,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f909,f1115,f3989,f412])).
fof(f5859,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4261,f4373])).
fof(f4261,plain,(
  ~sP3('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f511,f1107,f1115,f3989,f412])).
fof(f5858,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4262,f4379])).
fof(f4262,plain,(
  ~sP2('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1120,f1107,f1115,f3989,f412])).
fof(f5857,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4287,f4443])).
fof(f4287,plain,(
  sP5('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f511,f1120,f1116,f3989,f412])).
fof(f5856,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4296,f3989])).
fof(f4296,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1115,f498,f3989,f415])).
fof(f5855,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4297,f3989])).
fof(f4297,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f506,f498,f3989,f415])).
fof(f5854,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4298,f3989])).
fof(f4298,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1116,f498,f3989,f415])).
fof(f5853,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4313,f4335])).
fof(f4313,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f3989,f415])).
fof(f5852,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4315,f3989])).
fof(f4315,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f505,f498,f3989,f415])).
fof(f5851,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4316,f3989])).
fof(f4316,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1115,f498,f3989,f415])).
fof(f5850,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4318,f3989])).
fof(f4318,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1116,f498,f3989,f415])).
fof(f5849,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4324,f4335])).
fof(f4324,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f275,f3989,f415])).
fof(f5848,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4333,f4335])).
fof(f4333,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f3989,f415])).
fof(f5847,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4338,f3989])).
fof(f4338,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1120,f498,f3989,f419])).
fof(f5846,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4339,f3989])).
fof(f4339,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f509,f498,f3989,f419])).
fof(f5845,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4340,f3989])).
fof(f4340,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1121,f498,f3989,f419])).
fof(f5844,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4353,f4373])).
fof(f4353,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f3989,f419])).
fof(f5843,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4355,f3989])).
fof(f4355,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f508,f498,f3989,f419])).
fof(f5842,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4356,f3989])).
fof(f4356,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1120,f498,f3989,f419])).
fof(f5841,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4358,f3989])).
fof(f4358,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1121,f498,f3989,f419])).
fof(f5840,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4363,f4373])).
fof(f4363,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f275,f3989,f419])).
fof(f5839,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4371,f4373])).
fof(f4371,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f3989,f419])).
fof(f5834,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4390,f877])).
fof(f4390,plain,(
  s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1090,f3989,f428])).
fof(f5833,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f4391,f2692])).
fof(f4391,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f2934,f3989,f428])).
fof(f2934,plain,(
  ~sP5('F_0','T_0',f(s(bool,'T_0'))) | $spl57),
  inference(unit_resulting_resolution,[],[f498,f2692,f430])).
fof(f5832,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4393,f1243])).
fof(f4393,plain,(
  s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3384,f3989,f428])).
fof(f5831,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f4394,f2692])).
fof(f4394,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f2932,f3989,f428])).
fof(f2932,plain,(
  ~sP5('F_0',f(s(bool,'T_0')),'T_0') | $spl57),
  inference(unit_resulting_resolution,[],[f498,f2692,f430])).
fof(f5830,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4395,f877])).
fof(f4395,plain,(
  s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1085,f3989,f428])).
fof(f5829,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4397,f1243])).
fof(f4397,plain,(
  s(bool,'F_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3412,f3989,f428])).
fof(f5826,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4416,f522])).
fof(f4416,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f663,f3989,f429])).
fof(f5825,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4417,f1095])).
fof(f4417,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3327,f3989,f429])).
fof(f5824,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4418,f522])).
fof(f4418,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f662,f3989,f429])).
fof(f5823,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5822,f522])).
fof(f5822,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(forward_demodulation,[],[f4419,f877])).
fof(f4419,plain,(
  s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1995,f3989,f429])).
fof(f5821,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f5820,f522])).
fof(f5820,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl57 | $spl60)),
  inference(forward_demodulation,[],[f4420,f2692])).
fof(f4420,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f2694,f3989,f429])).
fof(f2694,plain,(
  sP5('F_0','F_0',f(s(bool,'T_0'))) | $spl57),
  inference(unit_resulting_resolution,[],[f498,f2339,f427])).
fof(f5819,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f5818,f522])).
fof(f5818,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl57 | $spl60)),
  inference(forward_demodulation,[],[f4421,f2692])).
fof(f4421,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f2689,f3989,f429])).
fof(f5817,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5816,f522])).
fof(f5816,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(forward_demodulation,[],[f4422,f877])).
fof(f4422,plain,(
  s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1975,f3989,f429])).
fof(f5815,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5814,f522])).
fof(f5814,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(forward_demodulation,[],[f5813,f877])).
fof(f5813,plain,(
  s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(forward_demodulation,[],[f4423,f1243])).
fof(f4423,plain,(
  s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3328,f3989,f429])).
fof(f5812,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5811,f522])).
fof(f5811,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(forward_demodulation,[],[f5810,f877])).
fof(f5810,plain,(
  s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(forward_demodulation,[],[f4424,f1243])).
fof(f4424,plain,(
  s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3355,f3989,f429])).
fof(f5809,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4425,f1095])).
fof(f4425,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3356,f3989,f429])).
fof(f5808,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5807,f522])).
fof(f5807,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(forward_demodulation,[],[f4438,f877])).
fof(f4438,plain,(
  s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3294,f3989,f429])).
fof(f5806,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4439,f522])).
fof(f4439,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1140,f3989,f429])).
fof(f5805,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4440,f522])).
fof(f4440,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1141,f3989,f429])).
fof(f5804,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4463,f4741])).
fof(f4741,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f3989,f3989,f444])).
fof(f4463,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f432])).
fof(f5803,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4479,f4741])).
fof(f4479,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f433])).
fof(f5802,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4482,f3989])).
fof(f4482,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f1142,f3989,f434])).
fof(f5801,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4483,f3989])).
fof(f4483,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f515,f547,f3989,f434])).
fof(f5800,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4484,f3989])).
fof(f4484,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f513,f547,f3989,f434])).
fof(f5799,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4486,f3989])).
fof(f4486,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f547,f3989,f434])).
fof(f5798,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4487,f3989])).
fof(f4487,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f1144,f3989,f434])).
fof(f5797,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4488,f3989])).
fof(f4488,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f513,f1144,f3989,f434])).
fof(f5796,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4490,f3989])).
fof(f4490,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f1144,f3989,f434])).
fof(f5795,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4492,f3989])).
fof(f4492,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f515,f916,f3989,f434])).
fof(f5794,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4493,f3989])).
fof(f4493,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f931,f916,f3989,f434])).
fof(f5793,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4495,f3989])).
fof(f4495,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f916,f3989,f434])).
fof(f5792,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4500,f4619])).
fof(f4619,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f439])).
fof(f4500,plain,(
  ~sP6('F_0','T_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f547,f275,f3989,f434])).
fof(f5791,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4502,f4619])).
fof(f4502,plain,(
  ~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f275,f3989,f434])).
fof(f5790,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4506,f4619])).
fof(f4506,plain,(
  ~sP6('F_0','T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f275,f3989,f434])).
fof(f5789,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4515,f4619])).
fof(f4515,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f434])).
fof(f5788,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4517,f4619])).
fof(f4517,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f3989,f3989,f434])).
fof(f5787,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4519,f4619])).
fof(f4519,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f547,f3989,f3989,f434])).
fof(f5786,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4521,f4619])).
fof(f4521,plain,(
  ( ! [X0] : (~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f3989,f3989,f434])).
fof(f5785,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4527,f4619])).
fof(f4527,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f3989,f3989,f434])).
fof(f5784,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4529,f4619])).
fof(f4529,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f3989,f3989,f434])).
fof(f5783,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4532,f3989])).
fof(f4532,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f514,f546,f3989,f434])).
fof(f5782,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4533,f3989])).
fof(f4533,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f513,f546,f3989,f434])).
fof(f5781,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4535,f3989])).
fof(f4535,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f546,f3989,f434])).
fof(f5780,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4536,f3989])).
fof(f4536,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f1142,f3989,f434])).
fof(f5779,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4537,f3989])).
fof(f4537,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f513,f1142,f3989,f434])).
fof(f5778,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4539,f3989])).
fof(f4539,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f1142,f3989,f434])).
fof(f5777,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4541,f3989])).
fof(f4541,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f1144,f3989,f434])).
fof(f5776,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4542,f3989])).
fof(f4542,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f514,f915,f3989,f434])).
fof(f5775,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4543,f3989])).
fof(f4543,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f931,f915,f3989,f434])).
fof(f5774,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4545,f3989])).
fof(f4545,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f915,f3989,f434])).
fof(f5773,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4551,f4619])).
fof(f4551,plain,(
  ~sP6('T_0','F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f546,f275,f3989,f434])).
fof(f5772,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4553,f4619])).
fof(f4553,plain,(
  ~sP6('T_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f275,f3989,f434])).
fof(f5771,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4557,f4619])).
fof(f4557,plain,(
  ~sP6('T_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f275,f3989,f434])).
fof(f5770,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4566,f4619])).
fof(f4566,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f434])).
fof(f5769,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4568,f4619])).
fof(f4568,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f3989,f3989,f434])).
fof(f5768,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4570,f4619])).
fof(f4570,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f547,f3989,f3989,f434])).
fof(f5767,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4572,f4619])).
fof(f4572,plain,(
  ( ! [X0] : (~sP6('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f3989,f3989,f434])).
fof(f5766,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4578,f4619])).
fof(f4578,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f3989,f3989,f434])).
fof(f5765,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4580,f4619])).
fof(f4580,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f3989,f3989,f434])).
fof(f5764,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4584,f3989])).
fof(f4584,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f498,f3989,f438])).
fof(f5763,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4585,f3989])).
fof(f4585,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f515,f498,f3989,f438])).
fof(f5762,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4586,f3989])).
fof(f4586,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f498,f3989,f438])).
fof(f5761,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4599,f4619])).
fof(f4599,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f3989,f438])).
fof(f5760,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4601,f3989])).
fof(f4601,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f514,f498,f3989,f438])).
fof(f5759,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4602,f3989])).
fof(f4602,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f498,f3989,f438])).
fof(f5758,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4604,f3989])).
fof(f4604,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f498,f3989,f438])).
fof(f5757,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4609,f4619])).
fof(f4609,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f275,f3989,f438])).
fof(f5756,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4617,f4619])).
fof(f4617,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f3989,f438])).
fof(f5755,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4621,f275])).
fof(f4621,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f547,f275,f3989,f444])).
fof(f5754,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4622,f275])).
fof(f4622,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f275,f3989,f444])).
fof(f5753,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4625,f3989])).
fof(f4625,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f275,f3989,f444])).
fof(f5752,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4626,f3989])).
fof(f4626,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f275,f3989,f444])).
fof(f5751,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4627,f275])).
fof(f4627,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f444])).
fof(f5750,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4628,f275])).
fof(f4628,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f3989,f3989,f444])).
fof(f5749,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4629,f275])).
fof(f4629,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f547,f3989,f3989,f444])).
fof(f5748,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4630,f275])).
fof(f4630,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f3989,f3989,f444])).
fof(f5747,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4635,f3989])).
fof(f4635,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f3989,f3989,f444])).
fof(f5746,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4636,f3989])).
fof(f4636,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f3989,f3989,f444])).
fof(f5745,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4637,f3989])).
fof(f4637,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f3989,f3989,f444])).
fof(f5744,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4639,f3989])).
fof(f4639,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f275,f3989,f444])).
fof(f5743,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4640,f3989])).
fof(f4640,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f547,f275,f3989,f444])).
fof(f5742,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4641,f3989])).
fof(f4641,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f275,f3989,f444])).
fof(f5741,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4647,f3989])).
fof(f4647,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f3989,f3989,f444])).
fof(f5740,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4648,f3989])).
fof(f4648,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f547,f3989,f3989,f444])).
fof(f5739,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4649,f3989])).
fof(f4649,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f3989,f3989,f444])).
fof(f5738,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4650,f275])).
fof(f4650,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f727,f3989,f3989,f444])).
fof(f5737,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4651,f275])).
fof(f4651,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f946,f3989,f3989,f444])).
fof(f5736,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4652,f3989])).
fof(f4652,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f949,f3989,f3989,f444])).
fof(f5735,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4653,f3989])).
fof(f4653,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f950,f3989,f3989,f444])).
fof(f5734,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4655,f3989])).
fof(f4655,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f3989,f3989,f444])).
fof(f5733,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4656,f275])).
fof(f4656,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f3989,f3989,f444])).
fof(f5732,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4661,f3989])).
fof(f4661,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f546,f275,f3989,f444])).
fof(f5731,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4662,f3989])).
fof(f4662,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f275,f3989,f444])).
fof(f5730,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4664,f3989])).
fof(f4664,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f275,f3989,f444])).
fof(f5729,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4665,f3989])).
fof(f4665,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f444])).
fof(f5728,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4666,f3989])).
fof(f4666,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f3989,f3989,f444])).
fof(f5727,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4668,f3989])).
fof(f4668,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f3989,f3989,f444])).
fof(f5726,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4669,f275])).
fof(f4669,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f727,f3989,f3989,f444])).
fof(f5725,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4670,f3989])).
fof(f4670,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f946,f3989,f3989,f444])).
fof(f5724,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4671,f275])).
fof(f4671,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f949,f3989,f3989,f444])).
fof(f5723,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4672,f3989])).
fof(f4672,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f950,f3989,f3989,f444])).
fof(f5722,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4673,f3989])).
fof(f4673,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f3989,f3989,f444])).
fof(f5721,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4675,f275])).
fof(f4675,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f3989,f3989,f444])).
fof(f5720,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4676,f275])).
fof(f4676,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f546,f275,f3989,f444])).
fof(f5719,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4677,f275])).
fof(f4677,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f275,f3989,f444])).
fof(f5718,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4680,f3989])).
fof(f4680,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f275,f3989,f444])).
fof(f5717,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4681,f3989])).
fof(f4681,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f275,f3989,f444])).
fof(f5716,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4686,f275])).
fof(f4686,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f444])).
fof(f5715,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4687,f275])).
fof(f4687,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f3989,f3989,f444])).
fof(f5714,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4688,f275])).
fof(f4688,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f547,f3989,f3989,f444])).
fof(f5713,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4689,f275])).
fof(f4689,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f3989,f3989,f444])).
fof(f5712,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4694,f3989])).
fof(f4694,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f3989,f3989,f444])).
fof(f5711,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4695,f3989])).
fof(f4695,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f3989,f3989,f444])).
fof(f5710,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4696,f3989])).
fof(f4696,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f3989,f3989,f444])).
fof(f5709,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4702,f3989])).
fof(f4702,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f275,f3989,f444])).
fof(f5708,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4703,f275])).
fof(f4703,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f727,f275,f3989,f444])).
fof(f5707,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4704,f3989])).
fof(f4704,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f946,f275,f3989,f444])).
fof(f5706,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4706,f275])).
fof(f4706,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f275,f3989,f444])).
fof(f5705,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4707,f3989])).
fof(f4707,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f444])).
fof(f5704,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4708,f3989])).
fof(f4708,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f3989,f3989,f444])).
fof(f5703,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4710,f3989])).
fof(f4710,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f3989,f3989,f444])).
fof(f5702,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4711,f275])).
fof(f4711,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f727,f3989,f3989,f444])).
fof(f5701,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4712,f3989])).
fof(f4712,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f946,f3989,f3989,f444])).
fof(f5700,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4713,f275])).
fof(f4713,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f949,f3989,f3989,f444])).
fof(f5699,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4714,f3989])).
fof(f4714,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f950,f3989,f3989,f444])).
fof(f5698,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4715,f3989])).
fof(f4715,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f3989,f3989,f444])).
fof(f5697,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4717,f275])).
fof(f4717,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f3989,f3989,f444])).
fof(f5696,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4719,f3989])).
fof(f4719,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f275,f3989,f444])).
fof(f5695,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4720,f275])).
fof(f4720,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f727,f275,f3989,f444])).
fof(f5694,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4721,f3989])).
fof(f4721,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f949,f275,f3989,f444])).
fof(f5693,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4723,f275])).
fof(f4723,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f275,f3989,f444])).
fof(f5692,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4729,f3989])).
fof(f4729,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1142,f3989,f3989,f444])).
fof(f5691,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4730,f3989])).
fof(f4730,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f547,f3989,f3989,f444])).
fof(f5690,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4731,f3989])).
fof(f4731,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1144,f3989,f3989,f444])).
fof(f5689,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4732,f275])).
fof(f4732,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f727,f3989,f3989,f444])).
fof(f5688,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4733,f275])).
fof(f4733,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f946,f3989,f3989,f444])).
fof(f5687,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4734,f3989])).
fof(f4734,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f949,f3989,f3989,f444])).
fof(f5686,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4735,f3989])).
fof(f4735,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f950,f3989,f3989,f444])).
fof(f5685,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4737,f3989])).
fof(f4737,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f3989,f3989,f444])).
fof(f5684,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4738,f275])).
fof(f4738,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1161,f3989,f3989,f444])).
fof(f5683,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4744,f3989])).
fof(f4744,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1151,f498,f3989,f447])).
fof(f5682,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4745,f3989])).
fof(f4745,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f556,f498,f3989,f447])).
fof(f5681,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4746,f3989])).
fof(f4746,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1156,f498,f3989,f447])).
fof(f5680,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4772,f3989])).
fof(f4772,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f550,f498,f3989,f447])).
fof(f5679,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4773,f3989])).
fof(f4773,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1151,f498,f3989,f447])).
fof(f5678,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4775,f3989])).
fof(f4775,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1156,f498,f3989,f447])).
fof(f5677,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4817,f5370])).
fof(f5370,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f518,f3989,f483])).
fof(f4817,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f448])).
fof(f5676,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4835,f5370])).
fof(f4835,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f449])).
fof(f5675,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4856,f4908])).
fof(f4908,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f457])).
fof(f4856,plain,(
  ~sP8('F_0','T_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f561,f498,f3989,f452])).
fof(f5674,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4858,f4908])).
fof(f4858,plain,(
  ~sP8('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f954,f498,f3989,f452])).
fof(f5673,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4863,f4908])).
fof(f4863,plain,(
  ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1163,f498,f3989,f452])).
fof(f5672,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4882,f4908])).
fof(f4882,plain,(
  ~sP8('F_0','F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f557,f498,f3989,f453])).
fof(f5671,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4883,f4908])).
fof(f4883,plain,(
  ~sP8('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f952,f498,f3989,f453])).
fof(f5670,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4887,f4908])).
fof(f4887,plain,(
  ~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1162,f498,f3989,f453])).
fof(f5669,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4904,f4908])).
fof(f4904,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f498,f3989,f455])).
fof(f5668,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4914,f275])).
fof(f4914,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f763,f3989,f460])).
fof(f5667,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4915,f275])).
fof(f4915,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1201,f3989,f460])).
fof(f5666,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4916,f275])).
fof(f4916,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1204,f3989,f460])).
fof(f5665,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4917,f3989])).
fof(f4917,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f973,f3989,f460])).
fof(f5664,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4918,f3989])).
fof(f4918,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1162,f3989,f460])).
fof(f5663,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4919,f3989])).
fof(f4919,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1163,f3989,f460])).
fof(f5662,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4924,f275])).
fof(f4924,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f557,f3989,f460])).
fof(f5661,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4925,f3989])).
fof(f4925,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f952,f3989,f460])).
fof(f5660,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4926,f3989])).
fof(f4926,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f561,f3989,f460])).
fof(f5659,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4927,f3989])).
fof(f4927,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f954,f3989,f460])).
fof(f5658,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4930,f3989])).
fof(f4930,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1204,f3989,f460])).
fof(f5657,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4932,f275])).
fof(f4932,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1162,f3989,f460])).
fof(f5656,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4933,f3989])).
fof(f4933,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1163,f3989,f460])).
fof(f5655,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4942,f275])).
fof(f4942,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f763,f3989,f461])).
fof(f5654,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4943,f275])).
fof(f4943,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1201,f3989,f461])).
fof(f5653,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4944,f275])).
fof(f4944,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1204,f3989,f461])).
fof(f5652,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4945,f3989])).
fof(f4945,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f973,f3989,f461])).
fof(f5651,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4946,f3989])).
fof(f4946,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1162,f3989,f461])).
fof(f5650,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4947,f3989])).
fof(f4947,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1163,f3989,f461])).
fof(f5649,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4952,f3989])).
fof(f4952,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f557,f3989,f461])).
fof(f5648,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4953,f3989])).
fof(f4953,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f952,f3989,f461])).
fof(f5647,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4954,f275])).
fof(f4954,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f561,f3989,f461])).
fof(f5646,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4955,f3989])).
fof(f4955,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f954,f3989,f461])).
fof(f5645,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4957,f3989])).
fof(f4957,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1201,f3989,f461])).
fof(f5644,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4960,f3989])).
fof(f4960,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1162,f3989,f461])).
fof(f5643,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f4961,f275])).
fof(f4961,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1163,f3989,f461])).
fof(f5642,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5001,f5348])).
fof(f5348,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f521,f3989,f482])).
fof(f5001,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989,f466])).
fof(f5641,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5004,f3989])).
fof(f5004,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f521,f1209,f3989,f469])).
fof(f5640,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5005,f3989])).
fof(f5005,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f521,f575,f3989,f469])).
fof(f5639,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5006,f3989])).
fof(f5006,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f520,f575,f3989,f469])).
fof(f5638,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5008,f3989])).
fof(f5008,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1222,f575,f3989,f469])).
fof(f5637,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5009,f3989])).
fof(f5009,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f521,f974,f3989,f469])).
fof(f5636,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5010,f3989])).
fof(f5010,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f994,f974,f3989,f469])).
fof(f5635,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5012,f3989])).
fof(f5012,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1222,f974,f3989,f469])).
fof(f5634,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5014,f3989])).
fof(f5014,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1223,f1212,f3989,f469])).
fof(f5633,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5015,f3989])).
fof(f5015,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f520,f1212,f3989,f469])).
fof(f5632,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5017,f3989])).
fof(f5017,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1222,f1212,f3989,f469])).
fof(f5631,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5022,f5204])).
fof(f5204,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f474])).
fof(f5022,plain,(
  ~sP10('F_0','T_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f575,f275,f3989,f469])).
fof(f5630,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5024,f5204])).
fof(f5024,plain,(
  ~sP10('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f974,f275,f3989,f469])).
fof(f5629,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5029,f5204])).
fof(f5029,plain,(
  ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f275,f3989,f469])).
fof(f5628,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5037,f5204])).
fof(f5037,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f574,f3989,f3989,f469])).
fof(f5627,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5039,f5204])).
fof(f5039,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f3989,f3989,f469])).
fof(f5626,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5041,f5204])).
fof(f5041,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f575,f3989,f3989,f469])).
fof(f5625,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5043,f5204])).
fof(f5043,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f974,f3989,f3989,f469])).
fof(f5624,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5050,f5204])).
fof(f5050,plain,(
  ( ! [X0] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f3989,f3989,f469])).
fof(f5623,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5052,f5204])).
fof(f5052,plain,(
  ( ! [X0] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f3989,f3989,f469])).
fof(f5622,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5057,f3989])).
fof(f5057,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f519,f3989,f469])).
fof(f5621,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5058,f3989])).
fof(f5058,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f519,f3989,f469])).
fof(f5620,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5062,f3989])).
fof(f5062,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f1222,f3989,f469])).
fof(f5619,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5066,f3989])).
fof(f5066,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f520,f3989,f469])).
fof(f5618,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5067,f3989])).
fof(f5067,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f520,f3989,f469])).
fof(f5617,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5071,f3989])).
fof(f5071,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f994,f3989,f469])).
fof(f5616,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5073,f3989])).
fof(f5073,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f1223,f3989,f469])).
fof(f5615,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5081,f5204])).
fof(f5081,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f574,f3989,f3989,f469])).
fof(f5614,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5082,f5204])).
fof(f5082,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f3989,f3989,f469])).
fof(f5613,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5083,f5204])).
fof(f5083,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f575,f3989,f3989,f469])).
fof(f5612,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5084,f5204])).
fof(f5084,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f974,f3989,f3989,f469])).
fof(f5611,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5090,f5204])).
fof(f5090,plain,(
  ( ! [X0] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f3989,f3989,f469])).
fof(f5610,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5091,f5204])).
fof(f5091,plain,(
  ( ! [X0] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f3989,f3989,f469])).
fof(f5609,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5121,f5204])).
fof(f5121,plain,(
  ~sP10('F_0','F_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f574,f498,f3989,f470])).
fof(f5608,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5122,f5204])).
fof(f5122,plain,(
  ~sP10('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f498,f3989,f470])).
fof(f5607,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5127,f5204])).
fof(f5127,plain,(
  ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f498,f3989,f470])).
fof(f5606,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5153,f5204])).
fof(f5153,plain,(
  ~sP10('F_0','F_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f574,f498,f3989,f471])).
fof(f5605,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5154,f5204])).
fof(f5154,plain,(
  ~sP10('F_0','F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f498,f3989,f471])).
fof(f5604,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5159,f5204])).
fof(f5159,plain,(
  ~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'F_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f498,f3989,f471])).
fof(f5603,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5179,f3989])).
fof(f5179,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1222,f498,f3989,f472])).
fof(f5602,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5180,f3989])).
fof(f5180,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f521,f498,f3989,f472])).
fof(f5601,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5181,f3989])).
fof(f5181,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1223,f498,f3989,f472])).
fof(f5600,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5184,f5204])).
fof(f5184,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f498,f3989,f472])).
fof(f5599,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5196,f3989])).
fof(f5196,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f519,f498,f3989,f472])).
fof(f5598,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5197,f3989])).
fof(f5197,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1222,f498,f3989,f472])).
fof(f5597,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5199,f3989])).
fof(f5199,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1223,f498,f3989,f472])).
fof(f5596,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5200,f5204])).
fof(f5200,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f275,f498,f3989,f472])).
fof(f5595,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5202,f5204])).
fof(f5202,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f498,f3989,f472])).
fof(f5594,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5236,f275])).
fof(f5236,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f799,f3989,f478])).
fof(f5593,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5237,f275])).
fof(f5237,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1003,f3989,f478])).
fof(f5592,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5238,f275])).
fof(f5238,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1231,f3989,f478])).
fof(f5591,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5239,f3989])).
fof(f5239,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1006,f3989,f478])).
fof(f5590,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5240,f3989])).
fof(f5240,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1007,f3989,f478])).
fof(f5589,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5241,f3989])).
fof(f5241,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f3989,f478])).
fof(f5588,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5242,f3989])).
fof(f5242,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f3989,f478])).
fof(f5587,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5247,f3989])).
fof(f5247,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f574,f3989,f478])).
fof(f5586,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5248,f3989])).
fof(f5248,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f3989,f478])).
fof(f5585,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5249,f275])).
fof(f5249,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f575,f3989,f478])).
fof(f5584,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5250,f3989])).
fof(f5250,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f974,f3989,f478])).
fof(f5583,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5253,f3989])).
fof(f5253,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1231,f3989,f478])).
fof(f5582,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5256,f3989])).
fof(f5256,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f3989,f478])).
fof(f5581,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5257,f275])).
fof(f5257,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f3989,f478])).
fof(f5580,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5262,f275])).
fof(f5262,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f577,f498,f3989,f479])).
fof(f5579,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5263,f275])).
fof(f5263,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1229,f498,f3989,f479])).
fof(f5578,plain,(
  $false | ($spl57 | $spl60)),
  inference(subsumption_resolution,[],[f5266,f3989])).
fof(f5266,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f3239,f498,f3989,f479])).
fof(f3239,plain,(
  ( ! [X0] : (sP11(f(s(bool,'T_0')),X0,'F_0')) ) | $spl57),
  inference(unit_resulting_resolution,[],[f498,f3012])).
fof(f3012,plain,(
  ( ! [X30,X29] : (sP11(f(s(bool,'T_0')),X30,X29) | p(s(bool,X29))) ) | $spl57),
  inference(subsumption_resolution,[],[f2968,f275])).
fof(f2968,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP11(f(s(bool,'T_0')),X30,X29)) ) | $spl57),
  inference(superposition,[],[f477,f2692])).
fof(f5577,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5267,f3989])).
fof(f5267,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1008,f498,f3989,f479])).
fof(f5576,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5268,f3989])).
fof(f5268,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3680,f498,f3989,f479])).
fof(f5575,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5269,f3989])).
fof(f5269,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1227,f498,f3989,f479])).
fof(f5574,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5294,f3989])).
fof(f5294,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f1229,f498,f3989,f479])).
fof(f5573,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5295,f275])).
fof(f5295,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f786,f498,f3989,f479])).
fof(f5572,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5296,f3989])).
fof(f5296,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) | $spl60),
  inference(unit_resulting_resolution,[],[f995,f498,f3989,f479])).
fof(f5571,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5299,f3989])).
fof(f5299,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f3680,f498,f3989,f479])).
fof(f5570,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5300,f275])).
fof(f5300,plain,(
  ~p(s(bool,'T_0')) | $spl60),
  inference(unit_resulting_resolution,[],[f1227,f498,f3989,f479])).
fof(f5569,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5329,f519])).
fof(f5329,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f574,f3989,f482])).
fof(f5568,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5330,f1222])).
fof(f5330,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f3989,f482])).
fof(f5567,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5331,f520])).
fof(f5331,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f575,f3989,f482])).
fof(f5566,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5332,f994])).
fof(f5332,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f974,f3989,f482])).
fof(f5565,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5338,f519])).
fof(f5338,plain,(
  ( ! [X0] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f3989,f482])).
fof(f5564,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5339,f520])).
fof(f5339,plain,(
  ( ! [X0] : (~sP10('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1212,f3989,f482])).
fof(f5563,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5352,f516])).
fof(f5352,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f557,f3989,f483])).
fof(f5562,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5353,f961])).
fof(f5353,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f952,f3989,f483])).
fof(f5561,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5354,f517])).
fof(f5354,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f561,f3989,f483])).
fof(f5560,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5355,f962])).
fof(f5355,plain,(
  ( ! [X0] : (~sP8('F_0','const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14)))),X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f954,f3989,f483])).
fof(f5559,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5360,f516])).
fof(f5360,plain,(
  ( ! [X0] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),X0,'T_0')) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1162,f3989,f483])).
fof(f5558,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5361,f517])).
fof(f5361,plain,(
  ( ! [X0] : (~sP8('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),'T_0',X0)) ) | $spl60),
  inference(unit_resulting_resolution,[],[f1163,f3989,f483])).
fof(f5557,plain,(
  $false | ($spl16 | $spl60)),
  inference(subsumption_resolution,[],[f5394,f3989])).
fof(f5394,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f1142,f3989,f1663])).
fof(f5556,plain,(
  $false | ($spl16 | $spl60)),
  inference(subsumption_resolution,[],[f5395,f3989])).
fof(f5395,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f547,f3989,f1663])).
fof(f5555,plain,(
  $false | ($spl16 | $spl60)),
  inference(subsumption_resolution,[],[f5396,f3989])).
fof(f5396,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f1144,f3989,f1663])).
fof(f5554,plain,(
  $false | ($spl16 | $spl60)),
  inference(subsumption_resolution,[],[f5401,f3989])).
fof(f5401,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f546,f3989,f1663])).
fof(f5553,plain,(
  $false | ($spl16 | $spl60)),
  inference(subsumption_resolution,[],[f5402,f3989])).
fof(f5402,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f1142,f3989,f1663])).
fof(f5552,plain,(
  $false | ($spl16 | $spl60)),
  inference(subsumption_resolution,[],[f5404,f3989])).
fof(f5404,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f1144,f3989,f1663])).
fof(f5551,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5461,f3989])).
fof(f5461,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f916,f3989,f2568])).
fof(f5550,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f5465,f3989])).
fof(f5465,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f915,f3989,f2568])).
fof(f5512,plain,(
  $false | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f498,f1121,f3989,f3002])).
fof(f3002,plain,(
  ( ! [X10,X9] : (~sP3(X9,f(s(bool,'T_0')),X10) | p(s(bool,X10)) | ~p(s(bool,X9))) ) | $spl57),
  inference(subsumption_resolution,[],[f2958,f275])).
fof(f2958,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP3(X9,f(s(bool,'T_0')),X10)) ) | $spl57),
  inference(superposition,[],[f419,f2692])).
fof(f5510,plain,(
  $false | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f498,f509,f3989,f3002])).
fof(f5505,plain,(
  $false | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f498,f1116,f3989,f3001])).
fof(f3001,plain,(
  ( ! [X8,X7] : (~sP4(X7,f(s(bool,'T_0')),X8) | p(s(bool,X8)) | ~p(s(bool,X7))) ) | $spl57),
  inference(subsumption_resolution,[],[f2957,f275])).
fof(f2957,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,f(s(bool,'T_0')),X8)) ) | $spl57),
  inference(superposition,[],[f415,f2692])).
fof(f5503,plain,(
  $false | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f498,f506,f3989,f3001])).
fof(f5496,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3680,f3989,f2578])).
fof(f5494,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f1008,f3989,f2578])).
fof(f5491,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1006,f3989,f2577])).
fof(f5489,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f1223,f3989,f2574])).
fof(f5487,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f521,f3989,f2574])).
fof(f5482,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f973,f3989,f2572])).
fof(f5480,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f973,f3989,f2571])).
fof(f5473,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f1156,f3989,f2569])).
fof(f5471,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f556,f3989,f2569])).
fof(f5468,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f916,f3989,f2568])).
fof(f5466,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f275,f915,f3989,f2568])).
fof(f5462,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f275,f916,f3989,f2568])).
fof(f5460,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f915,f3989,f2568])).
fof(f5456,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f1150,f3989,f2567])).
fof(f5454,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f515,f3989,f2567])).
fof(f5449,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f1121,f3989,f2566])).
fof(f5447,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f509,f3989,f2566])).
fof(f5442,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f1116,f3989,f2565])).
fof(f5440,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f506,f3989,f2565])).
fof(f5435,plain,(
  $false | ($spl36 | $spl60)),
  inference(unit_resulting_resolution,[],[f577,f498,f3989,f1693])).
fof(f5429,plain,(
  $false | ($spl34 | $spl60)),
  inference(unit_resulting_resolution,[],[f799,f3989,f1690])).
fof(f5427,plain,(
  $false | ($spl28 | $spl60)),
  inference(unit_resulting_resolution,[],[f521,f498,f3989,f1681])).
fof(f5421,plain,(
  $false | ($spl24 | $spl60)),
  inference(unit_resulting_resolution,[],[f763,f3989,f1675])).
fof(f5419,plain,(
  $false | ($spl22 | $spl60)),
  inference(unit_resulting_resolution,[],[f763,f3989,f1672])).
fof(f5417,plain,(
  $false | ($spl18 | $spl60)),
  inference(unit_resulting_resolution,[],[f556,f498,f3989,f1666])).
fof(f5408,plain,(
  $false | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f1663])).
fof(f5406,plain,(
  $false | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f546,f275,f3989,f1663])).
fof(f5400,plain,(
  $false | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f1663])).
fof(f5398,plain,(
  $false | ($spl16 | $spl60)),
  inference(unit_resulting_resolution,[],[f547,f275,f3989,f1663])).
fof(f5392,plain,(
  $false | ($spl14 | $spl60)),
  inference(unit_resulting_resolution,[],[f515,f498,f3989,f1660])).
fof(f5386,plain,(
  $false | ($spl12 | $spl60)),
  inference(unit_resulting_resolution,[],[f509,f498,f3989,f1657])).
fof(f5380,plain,(
  $false | ($spl10 | $spl60)),
  inference(unit_resulting_resolution,[],[f506,f498,f3989,f1654])).
fof(f5373,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1162,f1186,f3989,f483])).
fof(f5371,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f518,f3989,f483])).
fof(f5369,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f954,f962,f3989,f483])).
fof(f5367,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f561,f517,f3989,f483])).
fof(f5365,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f952,f961,f3989,f483])).
fof(f5363,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f516,f3989,f483])).
fof(f5351,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f1223,f3989,f482])).
fof(f5349,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f521,f3989,f482])).
fof(f5347,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f974,f994,f3989,f482])).
fof(f5345,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f575,f520,f3989,f482])).
fof(f5343,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f1222,f3989,f482])).
fof(f5341,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f519,f3989,f482])).
fof(f5328,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f577,f498,f3989,f3989,f479])).
fof(f5311,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f786,f498,f275,f3989,f479])).
fof(f5292,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f577,f498,f3989,f3989,f479])).
fof(f5275,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f577,f498,f275,f3989,f479])).
fof(f5261,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f3989,f3989,f478])).
fof(f5259,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f575,f275,f3989,f478])).
fof(f5246,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f3989,f3989,f478])).
fof(f5244,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f799,f275,f3989,f478])).
fof(f5231,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f498,f3989,f477])).
fof(f5218,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f498,f498,f3989,f476])).
fof(f5205,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f793,f3989,f474])).
fof(f5203,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f519,f3989,f498,f3989,f472])).
fof(f5201,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f519,f275,f498,f3989,f472])).
fof(f5185,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f519,f3989,f498,f3989,f472])).
fof(f5183,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f521,f275,f498,f3989,f472])).
fof(f5167,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f1223,f498,f3989,f471])).
fof(f5165,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f521,f498,f3989,f471])).
fof(f5163,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f1222,f498,f3989,f471])).
fof(f5161,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f519,f498,f3989,f471])).
fof(f5135,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f1223,f498,f3989,f470])).
fof(f5133,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f521,f498,f3989,f470])).
fof(f5131,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f1222,f498,f3989,f470])).
fof(f5129,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f519,f498,f3989,f470])).
fof(f5103,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1207,f1223,f3989,f3989,f469])).
fof(f5101,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f521,f3989,f3989,f469])).
fof(f5099,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f974,f994,f3989,f3989,f469])).
fof(f5097,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f575,f520,f3989,f3989,f469])).
fof(f5095,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1209,f1222,f3989,f3989,f469])).
fof(f5093,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f574,f519,f3989,f3989,f469])).
fof(f5053,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1223,f1212,f3989,f3989,f469])).
fof(f5051,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1223,f1207,f3989,f3989,f469])).
fof(f5044,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f521,f974,f3989,f3989,f469])).
fof(f5042,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f521,f575,f3989,f3989,f469])).
fof(f5040,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f521,f1209,f3989,f3989,f469])).
fof(f5038,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f521,f574,f3989,f3989,f469])).
fof(f5030,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1223,f1212,f275,f3989,f469])).
fof(f5025,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f521,f974,f275,f3989,f469])).
fof(f5023,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f521,f575,f275,f3989,f469])).
fof(f5002,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f786,f498,f3989,f466])).
fof(f4985,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f750,f498,f498,f3989,f462])).
fof(f4965,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f3989,f3989,f461])).
fof(f4963,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f561,f275,f3989,f461])).
fof(f4951,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f3989,f3989,f461])).
fof(f4949,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f763,f275,f3989,f461])).
fof(f4937,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f3989,f3989,f460])).
fof(f4935,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f275,f3989,f460])).
fof(f4923,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f3989,f3989,f460])).
fof(f4921,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f763,f275,f3989,f460])).
fof(f4909,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f757,f3989,f457])).
fof(f4907,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f757,f3989,f456])).
fof(f4905,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f518,f498,f498,f3989,f455])).
fof(f4895,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1162,f1186,f498,f3989,f453])).
fof(f4893,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f518,f498,f3989,f453])).
fof(f4891,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f952,f961,f498,f3989,f453])).
fof(f4889,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f557,f516,f498,f3989,f453])).
fof(f4864,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1186,f1163,f498,f3989,f452])).
fof(f4859,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f518,f954,f498,f3989,f452])).
fof(f4857,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f518,f561,f498,f3989,f452])).
fof(f4836,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f750,f498,f3989,f449])).
fof(f4818,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f750,f498,f3989,f448])).
fof(f4800,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f550,f498,f3989,f3989,f447])).
fof(f4784,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f550,f498,f275,f3989,f447])).
fof(f4771,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f550,f498,f3989,f3989,f447])).
fof(f4755,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f556,f498,f275,f3989,f447])).
fof(f4742,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f3989,f444])).
fof(f4740,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f547,f275,f3989,f3989,f444])).
fof(f4727,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f275,f3989,f444])).
fof(f4725,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f727,f275,f275,f3989,f444])).
fof(f4700,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f3989,f444])).
fof(f4698,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f546,f275,f3989,f3989,f444])).
fof(f4685,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f275,f3989,f444])).
fof(f4683,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f546,f275,f275,f3989,f444])).
fof(f4660,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f3989,f3989,f444])).
fof(f4658,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f547,f275,f3989,f3989,f444])).
fof(f4645,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f546,f3989,f275,f3989,f444])).
fof(f4643,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f547,f275,f275,f3989,f444])).
fof(f4620,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f721,f3989,f439])).
fof(f4618,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f498,f3989,f3989,f438])).
fof(f4610,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f498,f275,f3989,f438])).
fof(f4600,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f498,f3989,f3989,f438])).
fof(f4592,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f515,f498,f275,f3989,f438])).
fof(f4581,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f515,f916,f3989,f3989,f434])).
fof(f4579,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f915,f3989,f3989,f434])).
fof(f4573,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f1144,f3989,f3989,f434])).
fof(f4571,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f515,f547,f3989,f3989,f434])).
fof(f4569,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f1142,f3989,f3989,f434])).
fof(f4567,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f546,f3989,f3989,f434])).
fof(f4558,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f915,f275,f3989,f434])).
fof(f4554,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f1142,f275,f3989,f434])).
fof(f4552,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f546,f275,f3989,f434])).
fof(f4530,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f515,f916,f3989,f3989,f434])).
fof(f4528,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f915,f3989,f3989,f434])).
fof(f4522,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f1144,f3989,f3989,f434])).
fof(f4520,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f515,f547,f3989,f3989,f434])).
fof(f4518,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1149,f1142,f3989,f3989,f434])).
fof(f4516,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f514,f546,f3989,f3989,f434])).
fof(f4507,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f515,f916,f275,f3989,f434])).
fof(f4503,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1150,f1144,f275,f3989,f434])).
fof(f4501,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f515,f547,f275,f3989,f434])).
fof(f4480,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f556,f498,f3989,f433])).
fof(f4464,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f550,f498,f3989,f432])).
fof(f4448,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3356,f1095,f3989,f429])).
fof(f4446,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f662,f522,f3989,f429])).
fof(f4444,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3327,f1095,f3989,f429])).
fof(f4442,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f663,f522,f3989,f429])).
fof(f4415,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3384,f1243,f3989,f428])).
fof(f4413,plain,(
  $false | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f2934,f2692,f3989,f428])).
fof(f4411,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1090,f877,f3989,f428])).
fof(f4409,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3412,f1243,f3989,f428])).
fof(f4407,plain,(
  $false | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f2932,f2692,f3989,f428])).
fof(f4405,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1085,f877,f3989,f428])).
fof(f5547,plain,(
  $false | $spl60),
  inference(trivial_inequality_removal,[],[f4403])).
fof(f4403,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f3451,f3989,f428])).
fof(f5546,plain,(
  $false | $spl60),
  inference(trivial_inequality_removal,[],[f4402])).
fof(f4402,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f3311,f3989,f428])).
fof(f5545,plain,(
  $false | $spl60),
  inference(trivial_inequality_removal,[],[f4396])).
fof(f4396,plain,(
  s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK14))))) | $spl60),
  inference(unit_resulting_resolution,[],[f2645,f3989,f428])).
fof(f5544,plain,(
  $false | $spl60),
  inference(trivial_inequality_removal,[],[f4392])).
fof(f4392,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f709,f3989,f428])).
fof(f5543,plain,(
  $false | $spl60),
  inference(trivial_inequality_removal,[],[f4389])).
fof(f4389,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f701,f3989,f428])).
fof(f4380,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f704,f3989,f426])).
fof(f4378,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f704,f3989,f425])).
fof(f4376,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f704,f3989,f424])).
fof(f4374,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f691,f3989,f420])).
fof(f4372,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f508,f498,f3989,f3989,f419])).
fof(f4364,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f508,f498,f275,f3989,f419])).
fof(f4354,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f508,f498,f3989,f3989,f419])).
fof(f4346,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f509,f498,f275,f3989,f419])).
fof(f4336,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f685,f3989,f416])).
fof(f4334,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f505,f498,f3989,f3989,f415])).
fof(f4325,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f505,f498,f275,f3989,f415])).
fof(f4314,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f505,f498,f3989,f3989,f415])).
fof(f4305,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f506,f498,f275,f3989,f415])).
fof(f4263,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f511,f1120,f1107,f1115,f3989,f412])).
fof(f4245,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f911,f508,f890,f505,f3989,f412])).
fof(f4242,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f511,f508,f913,f505,f3989,f412])).
fof(f4237,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f911,f508,f889,f505,f3989,f412])).
fof(f4234,plain,(
  $false | ($spl57 | $spl60)),
  inference(unit_resulting_resolution,[],[f510,f508,f2688,f505,f3989,f412])).
fof(f4231,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f511,f508,f673,f505,f3989,f412])).
fof(f4213,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f911,f508,f890,f900,f3989,f412])).
fof(f4208,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f911,f508,f889,f900,f3989,f412])).
fof(f4190,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f511,f508,f913,f504,f3989,f412])).
fof(f4184,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f511,f1120,f1107,f504,f3989,f412])).
fof(f4181,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f511,f508,f673,f504,f3989,f412])).
fof(f4152,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f663,f3989,f498,f3989,f411])).
fof(f4150,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f847,f275,f498,f3989,f411])).
fof(f4108,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f663,f498,f3989,f3989,f411])).
fof(f4081,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f663,f498,f275,f3989,f411])).
fof(f4065,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f529,f498,f3989,f385])).
fof(f4056,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f529,f498,f3989,f382])).
fof(f4047,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f533,f3989,f3989,f381])).
fof(f4045,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f533,f275,f3989,f381])).
fof(f4035,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f533,f3989,f3989,f381])).
fof(f4033,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f536,f275,f3989,f381])).
fof(f4023,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f543,f3989,f367])).
fof(f4021,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f543,f3989,f366])).
fof(f4019,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f1095,f3989,f360])).
fof(f4017,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f522,f3989,f360])).
fof(f4015,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f275,f1095,f3989,f360])).
fof(f4013,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f275,f522,f3989,f360])).
fof(f4009,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f1095,f3989,f360])).
fof(f4007,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f3989,f522,f3989,f360])).
fof(f4005,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f275,f1095,f3989,f360])).
fof(f4003,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f275,f522,f3989,f360])).
fof(f3999,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f1243,f3989,f298])).
fof(f3997,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f522,f3989,f295])).
fof(f3995,plain,(
  $false | $spl60),
  inference(unit_resulting_resolution,[],[f498,f3989])).
fof(f3993,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl62),
  inference(cnf_transformation,[],[f3993_D])).
fof(f3993_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3992,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f3979,f294])).
fof(f3979,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1704,f283])).
fof(f3988,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl58),
  inference(cnf_transformation,[],[f3988_D])).
fof(f3988_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3990,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f3975,f294])).
fof(f3975,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1704,f608])).
fof(f3987,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f3974,f294])).
fof(f3974,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1704,f608])).
fof(f2327,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK17(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1054,f608])).
fof(f2326,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1054,f608])).
fof(f2306,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK17(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1053,f283])).
fof(f2305,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK17(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1053,f608])).
fof(f2304,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK17(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1053,f608])).
fof(f2303,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK17(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1053,f608])).
fof(f2302,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1053,f608])).
fof(f1792,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f274,f273,f272,f275,f498,f283,f295,f294,f299,f298,f348,f347,f355,f356,f360,f357,f367,f366,f499,f385,f382,f381,f380,f400,f399,f398,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f868,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1018,f1020,f1026,f1053,f1054,f1067,f1071,f1076,f1085,f1090,f1243,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1094,f1249,f1252,f1251,f1257,f1281,f1338,f1339,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1722,f1448,f1449,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1648,f1784,f1787,f1789,f1791])).
fof(f1791,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f1777,f275])).
fof(f1777,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) ) | $spl6),
  inference(backward_demodulation,[],[f1648,f398])).
fof(f1789,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) ) | $spl6),
  inference(subsumption_resolution,[],[f1774,f275])).
fof(f1774,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) ) | $spl6),
  inference(backward_demodulation,[],[f1648,f1018])).
fof(f1787,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f1773,f275])).
fof(f1773,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f1648,f348])).
fof(f1784,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f1763,f1648])).
fof(f1763,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f1648,f298])).
fof(f1790,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f274,f273,f272,f275,f498,f283,f295,f294,f299,f298,f348,f347,f355,f356,f360,f357,f367,f366,f499,f385,f382,f381,f380,f400,f399,f398,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f868,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1018,f1020,f1026,f1053,f1054,f1067,f1071,f1076,f1085,f1090,f1243,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1094,f1249,f1252,f1251,f1257,f1281,f1338,f1339,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1722,f1448,f1449,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1648,f1784,f1787,f1789])).
fof(f1788,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f274,f273,f272,f275,f498,f283,f295,f294,f299,f298,f348,f347,f355,f356,f360,f357,f367,f366,f499,f385,f382,f381,f380,f400,f399,f398,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f868,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1018,f1020,f1026,f1053,f1054,f1067,f1071,f1076,f1085,f1090,f1243,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1094,f1249,f1252,f1251,f1257,f1281,f1338,f1339,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1722,f1448,f1449,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1648,f1784,f1787])).
fof(f1786,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1765,f275])).
fof(f1765,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f1648,f498])).
fof(f1785,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f274,f273,f272,f275,f498,f283,f295,f294,f299,f298,f348,f347,f355,f356,f360,f357,f367,f366,f499,f385,f382,f381,f380,f400,f399,f398,f401,f414,f413,f412,f411,f410,f409,f408,f418,f417,f416,f415,f422,f421,f420,f419,f426,f425,f424,f423,f430,f429,f428,f427,f484,f434,f433,f432,f431,f441,f440,f439,f438,f447,f444,f443,f442,f454,f453,f452,f483,f450,f449,f448,f458,f457,f456,f455,f462,f461,f460,f459,f471,f470,f469,f482,f467,f466,f465,f475,f474,f473,f472,f479,f478,f477,f476,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f529,f530,f533,f536,f543,f546,f547,f550,f556,f557,f561,f564,f567,f574,f575,f577,f584,f608,f654,f663,f662,f673,f671,f685,f691,f701,f704,f709,f713,f721,f727,f750,f757,f763,f786,f793,f799,f833,f847,f866,f868,f872,f1011,f1010,f1009,f1008,f1007,f1006,f1003,f1002,f998,f995,f994,f991,f990,f986,f974,f973,f972,f968,f967,f963,f962,f961,f958,f954,f952,f950,f949,f946,f945,f944,f943,f942,f938,f937,f932,f931,f928,f927,f923,f916,f915,f914,f913,f912,f911,f910,f909,f906,f905,f901,f900,f897,f896,f892,f891,f890,f889,f886,f885,f884,f883,f881,f880,f879,f877,f1018,f1020,f1026,f1053,f1054,f1067,f1071,f1076,f1085,f1090,f1243,f1240,f1239,f1235,f1234,f1233,f1231,f1229,f1228,f1227,f1224,f1223,f1222,f1219,f1212,f1210,f1209,f1207,f1205,f1204,f1201,f1200,f1199,f1198,f1197,f1194,f1192,f1187,f1186,f1183,f1182,f1178,f1163,f1162,f1161,f1160,f1156,f1155,f1151,f1150,f1149,f1146,f1144,f1142,f1141,f1140,f1137,f1136,f1135,f1134,f1129,f1128,f1122,f1121,f1120,f1117,f1116,f1115,f1112,f1109,f1107,f1106,f1105,f1104,f1103,f1101,f1095,f1094,f1249,f1252,f1251,f1257,f1281,f1338,f1339,f1340,f1381,f1646,f1650,f1653,f1656,f1659,f1662,f1665,f1668,f1671,f1674,f1677,f1680,f1683,f1686,f1689,f1692,f1695,f1698,f1700,f1701,f1702,f1703,f1704,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1712,f1713,f1714,f1715,f1716,f1717,f1718,f1719,f1720,f1721,f1722,f1448,f1449,f1450,f1728,f1730,f1732,f1734,f1736,f1738,f1740,f1742,f1744,f1746,f1748,f1750,f1752,f1754,f1756,f1758,f1760,f1762,f1648,f1784])).
fof(f1699,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl40),
  inference(cnf_transformation,[],[f1699_D])).
fof(f1699_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1762,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1761])).
fof(f1761,plain,(
  ( ! [X97,X95,X98,X96] : (~sP8('T_0',X96,X97) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f1479,f654])).
fof(f1479,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~sP8('T_0',X96,X97) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f483,f608])).
fof(f1696,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f1696_D])).
fof(f1696_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1760,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1759])).
fof(f1759,plain,(
  ( ! [X94,X92,X93,X91] : (~sP10('T_0',X92,X93) | sP11('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1478,f654])).
fof(f1478,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP10('T_0',X92,X93) | sP11('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f482,f608])).
fof(f1693,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f1693_D])).
fof(f1693_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1758,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1757])).
fof(f1757,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | ~sP11('T_0',X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1477,f654])).
fof(f1477,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~p(s(bool,X88)) | p(s(bool,X89)) | ~sP11('T_0',X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f479,f608])).
fof(f1690,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f1690_D])).
fof(f1690_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1756,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1755])).
fof(f1755,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1476,f654])).
fof(f1476,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f478,f608])).
fof(f1687,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f1687_D])).
fof(f1687_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1754,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1753])).
fof(f1753,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1475,f654])).
fof(f1475,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f477,f608])).
fof(f1684,plain,(
  ( ! [X2,X1] : (sP11(X2,X1,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1684_D])).
fof(f1684_D,plain,(
  ( ! [X2,X1] : (sP11(X2,X1,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1752,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1751])).
fof(f1751,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11(X77,X76,'T_0') | p(s(bool,X77)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1474,f654])).
fof(f1474,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11(X77,X76,'T_0') | p(s(bool,X77)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f476,f608])).
fof(f1681,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f1681_D])).
fof(f1681_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1750,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1749])).
fof(f1749,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1473,f654])).
fof(f1473,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f472,f608])).
fof(f1678,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f1678_D])).
fof(f1678_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1748,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1747])).
fof(f1747,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1472,f654])).
fof(f1472,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f462,f608])).
fof(f1675,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f1675_D])).
fof(f1675_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1746,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1745])).
fof(f1745,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1471,f654])).
fof(f1471,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f461,f608])).
fof(f1672,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f1672_D])).
fof(f1672_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1744,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1743])).
fof(f1743,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1470,f654])).
fof(f1470,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f460,f608])).
fof(f1669,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f1669_D])).
fof(f1669_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1742,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1741])).
fof(f1741,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1469,f654])).
fof(f1469,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f455,f608])).
fof(f1666,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f1666_D])).
fof(f1666_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1740,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1739])).
fof(f1739,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1468,f654])).
fof(f1468,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f447,f608])).
fof(f1663,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f1663_D])).
fof(f1663_D,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1738,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1737])).
fof(f1737,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP7('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1467,f654])).
fof(f1467,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | sP7('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f444,f608])).
fof(f1660,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f1660_D])).
fof(f1660_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1736,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1735])).
fof(f1735,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1466,f654])).
fof(f1466,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f438,f608])).
fof(f1657,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f1657_D])).
fof(f1657_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1734,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1733])).
fof(f1733,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1465,f654])).
fof(f1465,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f419,f608])).
fof(f1654,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl10),
  inference(cnf_transformation,[],[f1654_D])).
fof(f1654_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1732,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1731])).
fof(f1731,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1464,f654])).
fof(f1464,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f415,f608])).
fof(f1651,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl8),
  inference(cnf_transformation,[],[f1651_D])).
fof(f1651_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1730,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1729])).
fof(f1729,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1463,f654])).
fof(f1463,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f385,f608])).
fof(f1647,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl4),
  inference(cnf_transformation,[],[f1647_D])).
fof(f1647_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1728,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1727])).
fof(f1727,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1462,f654])).
fof(f1462,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f382,f608])).
fof(f1698,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1697])).
fof(f1697,plain,(
  ( ! [X97,X95,X98,X96] : (~sP8('T_0',X96,X97) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f1409,f654])).
fof(f1409,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~sP8('T_0',X96,X97) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f483,f608])).
fof(f1695,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1694])).
fof(f1694,plain,(
  ( ! [X94,X92,X93,X91] : (~sP10('T_0',X92,X93) | sP11('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1408,f654])).
fof(f1408,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP10('T_0',X92,X93) | sP11('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f482,f608])).
fof(f1692,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1691])).
fof(f1691,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | ~sP11('T_0',X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1407,f654])).
fof(f1407,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~p(s(bool,X88)) | p(s(bool,X89)) | ~sP11('T_0',X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f479,f608])).
fof(f1689,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1688])).
fof(f1688,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1406,f654])).
fof(f1406,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f478,f608])).
fof(f1686,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1685])).
fof(f1685,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1405,f654])).
fof(f1405,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f477,f608])).
fof(f1683,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1682])).
fof(f1682,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11(X77,X76,'T_0') | p(s(bool,X77)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1404,f654])).
fof(f1404,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11(X77,X76,'T_0') | p(s(bool,X77)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f476,f608])).
fof(f1680,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1679])).
fof(f1679,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1403,f654])).
fof(f1403,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f472,f608])).
fof(f1677,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1676])).
fof(f1676,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1402,f654])).
fof(f1402,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f462,f608])).
fof(f1674,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1673])).
fof(f1673,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1401,f654])).
fof(f1401,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f461,f608])).
fof(f1671,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1670])).
fof(f1670,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1400,f654])).
fof(f1400,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f460,f608])).
fof(f1668,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1667])).
fof(f1667,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1399,f654])).
fof(f1399,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f455,f608])).
fof(f1665,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1664])).
fof(f1664,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1398,f654])).
fof(f1398,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f447,f608])).
fof(f1662,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1661])).
fof(f1661,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP7('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1397,f654])).
fof(f1397,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | sP7('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f444,f608])).
fof(f1659,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1658])).
fof(f1658,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1396,f654])).
fof(f1396,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f438,f608])).
fof(f1656,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1655])).
fof(f1655,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1395,f654])).
fof(f1395,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f419,f608])).
fof(f1653,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1652])).
fof(f1652,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1394,f654])).
fof(f1394,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f415,f608])).
fof(f1650,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1649])).
fof(f1649,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1393,f654])).
fof(f1393,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f385,f608])).
fof(f1646,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1645])).
fof(f1645,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1392,f654])).
fof(f1392,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f382,f608])).
fof(f501,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f500])).
fof(f500,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f495,f496])).
fof(f494,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | $spl0),
  inference(cnf_transformation,[],[f494_D])).
fof(f494_D,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f496,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | $spl2),
  inference(cnf_transformation,[],[f496_D])).
fof(f496_D,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f273,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14)))),
  inference(cnf_transformation,[],[f174])).
fof(f495,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) | $spl1),
  inference(cnf_transformation,[],[f495_D])).
fof(f495_D,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f497,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) | $spl3),
  inference(cnf_transformation,[],[f497_D])).
fof(f497_D,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f274,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK14))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK13),s('type.rat.rat',sK14)))),
  inference(cnf_transformation,[],[f174])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 776
% Passive clauses: 1977
% Generated clauses: 26417
% Final active clauses: 503
% Final passive clauses: 525
% Input formulas: 53
% Initial clauses: 210
% 
% Duplicate literals: 1452
% Trivial inequalities: 114
% Fw subsumption resolutions: 1969
% Fw demodulations: 2608
% Bw demodulations: 326
% Condensations: 2198
% Global subsumptions: 49
% 
% Simple tautologies: 184
% Equational tautologies: 710
% Forward subsumptions: 14450
% Bw demodulations to eq. taut.: 156
% 
% Unit resulting resolution: 7639
% Factoring: 1
% Forward superposition: 5084
% Backward superposition: 3882
% Self superposition: 154
% Equality factoring: 134
% Equality resolution: 9
% 
% Split clauses: 124
% Split components: 264
% SAT solver clauses: 2908
% SAT solver unit clauses: 780
% SAT solver binary clauses: 471
% 
% Sat splits: 124
% Sat splitting refutations: 632
% 
% Memory used [KB]: 10490
% Time elapsed: 1.410 s
% ------------------------------
% ------------------------------
% Success in time 2.728 s
