Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.bool
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bool
% SZS output start Proof for thm.bool
fof(f8434,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8409,f55,f82,f80,f78,f8105,f53,f79,f54,f81,f77,f312,f317,f315,f410,f8163,f8141])).
fof(f8141,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f8134,f320])).
fof(f320,plain,(
  s(bool,'T_0') != s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl17),
  inference(cnf_transformation,[],[f320_D])).
fof(f320_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f8134,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl0),
  inference(resolution,[],[f77,f85])).
fof(f85,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f57,f58])).
fof(f58,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bool.RES_EXISTS_UNIQUE_THM',aHL_BOOL_CASES)).
fof(f57,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bool.RES_EXISTS_UNIQUE_THM',aHL_FALSITY)).
fof(f8163,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f8153,f77])).
fof(f8153,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl3),
  inference(resolution,[],[f80,f60])).
fof(f60,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))))) & (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))) | p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))))) & ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0'))))))))) | p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0'))))))))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0'))))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X5,X9] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X6,bool),X5),s(fun(X6,bool),X9)))) <=> (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X5),s(fun(X6,bool),X9)))) & p(s(bool,happ(s(fun(fun(X6,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X6,bool),X5))),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(fun(X6,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X6,bool),X5))),s(fun(X6,fun(X6,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X6,fun(X6,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X6,bool),X9),s(fun(X6,bool),X9))),s(fun(X6,fun(X6,bool)),'=_0'))))))))))),
  file('thm.bool.RES_EXISTS_UNIQUE_THM','thm.bool.RES_EXISTS_UNIQUE_DEF')).
fof(f410,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f362,f56])).
fof(f56,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bool.RES_EXISTS_UNIQUE_THM',aHL_TRUTH)).
fof(f362,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f315,f57])).
fof(f315,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(cnf_transformation,[],[f315_D])).
fof(f315_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f317,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl12),
  inference(cnf_transformation,[],[f317_D])).
fof(f317_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f312,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(duplicate_literal_removal,[],[f276])).
fof(f276,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(superposition,[],[f256,f58])).
fof(f256,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(subsumption_resolution,[],[f182,f56])).
fof(f182,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(superposition,[],[f78,f83])).
fof(f83,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f58,f58])).
fof(f77,plain,(
  p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f77_D])).
fof(f77_D,plain,(
  p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f81,plain,(
  p(s(bool,happ(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) | $spl4),
  inference(cnf_transformation,[],[f81_D])).
fof(f81_D,plain,(
  p(s(bool,happ(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f54,plain,(
  p(s(bool,happ(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) | p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) & p(s(bool,happ(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))))) & (~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,happ(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f43])).
fof(f43,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))))) & (~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0'))))))))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))))) & (~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))))))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <~> (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0'))))))))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0'))))))))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0'))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,happ(s(fun(fun(X6,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(fun(X6,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X6,bool),X7))),s(fun(X6,fun(X6,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X6,fun(X6,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X6,bool),X8),s(fun(X6,bool),X8))),s(fun(X6,fun(X6,bool)),'=_0'))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,happ(s(fun(fun(X6,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),'_dst_x0x1_2'(s(fun(fun(X6,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X6,bool),X7))),s(fun(X6,fun(X6,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X6,fun(X6,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X6,bool),X8),s(fun(X6,bool),X8))),s(fun(X6,fun(X6,bool)),'=_0'))))))))))),
  file('thm.bool.RES_EXISTS_UNIQUE_THM','thm.bool.RES_EXISTS_UNIQUE_THM')).
fof(f79,plain,(
  p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl2),
  inference(cnf_transformation,[],[f79_D])).
fof(f79_D,plain,(
  p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f53,plain,(
  p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f44])).
fof(f8105,plain,(
  $false | ($spl2 | $spl4 | $spl12)),
  inference(subsumption_resolution,[],[f8104,f56])).
fof(f8104,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl12)),
  inference(forward_demodulation,[],[f8103,f91])).
fof(f91,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(subsumption_resolution,[],[f88,f57])).
fof(f88,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(superposition,[],[f79,f58])).
fof(f8103,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ($spl4 | $spl12)),
  inference(subsumption_resolution,[],[f8102,f57])).
fof(f8102,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ($spl4 | $spl12)),
  inference(forward_demodulation,[],[f8072,f317])).
fof(f8072,plain,(
  p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl4),
  inference(resolution,[],[f7723,f5787])).
fof(f5787,plain,(
  p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) | $spl4),
  inference(forward_demodulation,[],[f81,f59])).
fof(f59,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X10,X11] : s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X10),s(fun(X6,bool),X11))) = s(bool,happ(s(fun(fun(X6,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X6,bool),X10))),s(fun(X6,bool),X11)))),
  file('thm.bool.RES_EXISTS_UNIQUE_THM','ARITY_1_2_const.bool.RES_FORALL')).
fof(f7723,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))) | p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f62,f59])).
fof(f62,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f46])).
fof(f78,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f80,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl3),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f82,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) | $spl5),
  inference(cnf_transformation,[],[f82_D])).
fof(f82_D,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f55,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f44])).
fof(f8409,plain,(
  $false | ($spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8408,f8167])).
fof(f8167,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) | $spl5),
  inference(forward_demodulation,[],[f82,f59])).
fof(f8408,plain,(
  p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) | $spl16),
  inference(subsumption_resolution,[],[f8328,f56])).
fof(f8328,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'_dst_x0x1_2'(s(fun(fun(sK0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,fun(sK0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(sK0,fun(sK0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))),s(fun(sK0,fun(sK0,bool)),'=_0')))))))) | $spl16),
  inference(superposition,[],[f6672,f319])).
fof(f319,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl16),
  inference(cnf_transformation,[],[f319_D])).
fof(f319_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f6672,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0'))))))))) )),
  inference(forward_demodulation,[],[f61,f59])).
fof(f61,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'const.bool.RES_FORALL_1'(s(fun(X0,bool),X1))),s(fun(X0,fun(X0,bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,fun(X0,bool)),'_dst_x00x10x01_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(X0,bool),X2),s(fun(X0,bool),X2))),s(fun(X0,fun(X0,bool)),'=_0')))))))) | ~p(s(bool,'const.bool.RES_EXISTS_UNIQUE_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f46])).
% SZS output end Proof for thm.bool
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 109
% Passive clauses: 1372
% Generated clauses: 8267
% Final active clauses: 74
% Final passive clauses: 1112
% Input formulas: 15
% Initial clauses: 24
% 
% Duplicate literals: 363
% Trivial inequalities: 25
% Fw subsumption resolutions: 577
% Fw demodulations: 904
% Bw demodulations: 26
% 
% Simple tautologies: 88
% Equational tautologies: 24
% Forward subsumptions: 4582
% 
% Binary resolution: 179
% Forward superposition: 3628
% Backward superposition: 2270
% Self superposition: 110
% Equality factoring: 36
% Equality resolution: 2
% 
% Split clauses: 211
% Split components: 430
% SAT solver clauses: 405
% SAT solver unit clauses: 25
% SAT solver binary clauses: 92
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 211
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 402
% TWLsolver calls for satisfiability: 81
% 
% Memory used [KB]: 10362
% Time elapsed: 0.260 s
% ------------------------------
% ------------------------------
% Success in time 0.262 s
