Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f11490,plain,(
  $false),
  inference(subsumption_resolution,[],[f11489,f214])).
fof(f214,plain,(
  s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14)))))) & s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f98])).
fof(f98,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2)))))) & s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X2))))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))))))) & s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X2))))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))))))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))))))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X2))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))))))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X2))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X6,'type.num.num'),X7),s(fun(X6,bool),X8)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.num.num',happ(s(fun(X6,'type.num.num'),X7),s(X6,X9))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X6,'type.num.num'),X7),s(fun(X6,bool),X8))))))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X6,'type.num.num'),X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X8))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X6,'type.num.num'),X7),s(X6,X9))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X6,'type.num.num'),X7),s(fun(X6,bool),X8))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X6,'type.num.num'),X7),s(fun(X6,bool),X8)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.num.num',happ(s(fun(X6,'type.num.num'),X7),s(X6,X9))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X6,'type.num.num'),X7),s(fun(X6,bool),X8))))))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X6,'type.num.num'),X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X8))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X6,'type.num.num'),X7),s(X6,X9))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(X6,'type.num.num'),X7),s(fun(X6,bool),X8))))))))),
  file('thm.pred_set.PROD_SET_IMAGE_REDUCTION','thm.pred_set.PROD_SET_IMAGE_REDUCTION')).
fof(f11489,plain,(
  s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(forward_demodulation,[],[f11488,f294])).
fof(f294,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X4)))))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X4)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X4)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13,X6] : ! [X7,X9,X8] : s(fun(X13,bool),'const.pred_set.IMAGE_2'(s(fun(X6,X13),X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X8))))) = s(fun(X13,bool),'const.pred_set.INSERT_2'(s(X13,happ(s(fun(X6,X13),X7),s(X6,X9))),s(fun(X13,bool),'const.pred_set.IMAGE_2'(s(fun(X6,X13),X7),s(fun(X6,bool),X8)))))),
  file('thm.pred_set.PROD_SET_IMAGE_REDUCTION','thm.pred_set.IMAGE_INSERT')).
fof(f11488,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14))))))) = s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(subsumption_resolution,[],[f11477,f3417])).
fof(f3417,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14)))))),
  inference(resolution,[],[f212,f234])).
fof(f234,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.PROD_SET_IMAGE_REDUCTION','thm.bool.EQ_CLAUSES|split|2')).
fof(f212,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14))))))),
  inference(cnf_transformation,[],[f142])).
fof(f11477,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14))))))) = s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14)))))),
  inference(superposition,[],[f4180,f4048])).
fof(f4048,plain,(
  s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14))) = s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14))),s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15)))))),
  inference(resolution,[],[f261,f213])).
fof(f213,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.num.num',happ(s(fun(sK12,'type.num.num'),sK13),s(sK12,sK15))),s(fun('type.num.num',bool),'const.pred_set.IMAGE_2'(s(fun(sK12,'type.num.num'),sK13),s(fun(sK12,bool),sK14))))))),
  inference(cnf_transformation,[],[f142])).
fof(f261,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1)))) & (s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X9,X8] : (~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X9))))),
  file('thm.pred_set.PROD_SET_IMAGE_REDUCTION','thm.pred_set.DELETE_NON_ELEMENT')).
fof(f4180,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X1),s(fun('type.num.num',bool),X2))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X2),s('type.num.num',X1))))))) | s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X2)))) )),
  inference(resolution,[],[f263,f235])).
fof(f235,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f147])).
fof(f263,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) | s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X1),s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) | s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X1),s('type.num.num',X0))))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X1),s('type.num.num',X0))))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9,X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X8)))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X9),s(fun('type.num.num',bool),X8))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X8),s('type.num.num',X9))))))))),
  file('thm.pred_set.PROD_SET_IMAGE_REDUCTION','thm.pred_set.PROD_SET_THM|split|1')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 565
% Passive clauses: 4202
% Generated clauses: 11192
% Final active clauses: 531
% Final passive clauses: 3265
% Input formulas: 38
% Initial clauses: 157
% 
% Duplicate literals: 320
% Trivial inequalities: 74
% Fw subsumption resolutions: 1191
% Fw demodulations: 178
% Bw demodulations: 33
% 
% Simple tautologies: 182
% Equational tautologies: 59
% Forward subsumptions: 4726
% 
% Binary resolution: 5498
% Factoring: 1017
% Forward superposition: 1369
% Backward superposition: 1142
% Self superposition: 28
% Equality factoring: 26
% Equality resolution: 22
% 
% Split clauses: 344
% Split components: 694
% SAT solver clauses: 587
% SAT solver unit clauses: 16
% SAT solver binary clauses: 94
% 
% Sat splits: 344
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 585
% TWLsolver calls for satisfiability: 321
% 
% Memory used [KB]: 5117
% Time elapsed: 0.302 s
% ------------------------------
% ------------------------------
% Success in time 0.306 s
