Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.semi_ring
% (90047)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 166
% Passive clauses: 3799
% Generated clauses: 21558
% Final active clauses: 159
% Final passive clauses: 3527
% Input formulas: 69
% Initial clauses: 107
% 
% Duplicate literals: 559
% Trivial inequalities: 311
% Fw subsumption resolutions: 1540
% Fw demodulations: 2722
% Bw demodulations: 1
% 
% Simple tautologies: 355
% Equational tautologies: 40
% Forward subsumptions: 11866
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 520
% Factoring: 2
% Forward superposition: 9076
% Backward superposition: 6392
% Self superposition: 160
% Equality factoring: 65
% Equality resolution: 28
% 
% Split clauses: 248
% Split components: 509
% SAT solver clauses: 472
% SAT solver unit clauses: 17
% SAT solver binary clauses: 97
% 
% Sat splits: 248
% 
% TWLsolver clauses: 472
% TWLsolver calls for satisfiability: 130
% 
% Memory used [KB]: 18677
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.semi_ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.semi_ring
% SZS output start Proof for thm.semi_ring
fof(f20526,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f20512,f453,f231,f342,f232,f344,f233,f346,f234,f348,f340,f20271,f17264,f14560,f14590,f13368,f13370,f15991,f3851,f3864,f3669,f3671,f13364,f3598,f3602,f455,f3584,f3662,f235,f349,f347,f345,f343,f341,f311,f402,f313,f404,f314,f405,f15995,f15997,f312,f403,f15994,f17276])).
fof(f17276,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) | $spl2038),
  inference(cnf_transformation,[],[f17276_D])).
fof(f17276_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) <=> ~$spl2038),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2038])])).
fof(f15994,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f340,f346])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) ) | $spl104),
  inference(cnf_transformation,[],[f403_D])).
fof(f403_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f312,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X29,X7,X30,X39] : s(X6,X7) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X29),s(X6,X7),s(fun(X6,fun(X6,X6)),X30),s(fun(X6,fun(X6,X6)),X39)))))),
  file('thm.semi_ring.semi_ring_|49||49|','thm.semi_ring.semi_ring_SR|49|')).
fof(f15997,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) | $spl1918),
  inference(cnf_transformation,[],[f15997_D])).
fof(f15997_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) <=> ~$spl1918),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1918])])).
fof(f15995,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) | ($spl0 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f15994,f342])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) ) | $spl108),
  inference(cnf_transformation,[],[f405_D])).
fof(f405_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f314,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X29,X7,X30,X39] : s(fun(X6,fun(X6,X6)),X39) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X29),s(X6,X7),s(fun(X6,fun(X6,X6)),X30),s(fun(X6,fun(X6,X6)),X39)))))),
  file('thm.semi_ring.semi_ring_|49||49|','thm.semi_ring.semi_ring_SRM')).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) ) | $spl106),
  inference(cnf_transformation,[],[f404_D])).
fof(f404_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f313,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X29,X7,X30,X39] : s(fun(X6,fun(X6,X6)),X30) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X29),s(X6,X7),s(fun(X6,fun(X6,X6)),X30),s(fun(X6,fun(X6,X6)),X39)))))),
  file('thm.semi_ring.semi_ring_|49||49|','thm.semi_ring.semi_ring_SRP')).
fof(f402,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) ) | $spl102),
  inference(cnf_transformation,[],[f402_D])).
fof(f402_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f311,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X29,X7,X30,X39] : s(X6,X29) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X29),s(X6,X7),s(fun(X6,fun(X6,X6)),X30),s(fun(X6,fun(X6,X6)),X39)))))),
  file('thm.semi_ring.semi_ring_|49||49|','thm.semi_ring.semi_ring_SR|48|')).
fof(f341,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) | $spl1),
  inference(cnf_transformation,[],[f341_D])).
fof(f341_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f343,plain,(
  s(sK0,sK1) != s(sK0,sK5) | $spl3),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  s(sK0,sK1) != s(sK0,sK5) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f345,plain,(
  s(sK0,sK2) != s(sK0,sK6) | $spl5),
  inference(cnf_transformation,[],[f345_D])).
fof(f345_D,plain,(
  s(sK0,sK2) != s(sK0,sK6) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f347,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK7) | $spl7),
  inference(cnf_transformation,[],[f347_D])).
fof(f347_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK7) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f349,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK8) | $spl9),
  inference(cnf_transformation,[],[f349_D])).
fof(f349_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK8) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f235,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK8) | s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK7) | s(sK0,sK2) != s(sK0,sK6) | s(sK0,sK1) != s(sK0,sK5) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8)))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  (s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) | (s(sK0,sK1) = s(sK0,sK5) & s(sK0,sK2) = s(sK0,sK6) & s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) & s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8))) & (s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) | s(sK0,sK1) != s(sK0,sK5) | s(sK0,sK2) != s(sK0,sK6) | s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK7) | s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK8))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f200])).
fof(f200,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X5),s(X0,X6),s(fun(X0,fun(X0,X0)),X7),s(fun(X0,fun(X0,X0)),X8))) | (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8))) & (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4))) != s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X5),s(X0,X6),s(fun(X0,fun(X0,X0)),X7),s(fun(X0,fun(X0,X0)),X8))) | s(X0,X1) != s(X0,X5) | s(X0,X2) != s(X0,X6) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X7) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X8)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X5),s(X0,X6),s(fun(X0,fun(X0,X0)),X7),s(fun(X0,fun(X0,X0)),X8))) | (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8))) & (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4))) != s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X5),s(X0,X6),s(fun(X0,fun(X0,X0)),X7),s(fun(X0,fun(X0,X0)),X8))) | (s(X0,X1) != s(X0,X5) | s(X0,X2) != s(X0,X6) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X7) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X8))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X5),s(X0,X6),s(fun(X0,fun(X0,X0)),X7),s(fun(X0,fun(X0,X0)),X8))) <~> (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X5),s(X0,X6),s(fun(X0,fun(X0,X0)),X7),s(fun(X0,fun(X0,X0)),X8))) <=> (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X5),s(X0,X6),s(fun(X0,fun(X0,X0)),X7),s(fun(X0,fun(X0,X0)),X8))) <=> (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X7),s(X6,X8),s(fun(X6,fun(X6,X6)),X9),s(fun(X6,fun(X6,X6)),X10))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X11),s(X6,X12),s(fun(X6,fun(X6,X6)),X13),s(fun(X6,fun(X6,X6)),X14))) <=> (s(X6,X7) = s(X6,X11) & s(X6,X8) = s(X6,X12) & s(fun(X6,fun(X6,X6)),X9) = s(fun(X6,fun(X6,X6)),X13) & s(fun(X6,fun(X6,X6)),X10) = s(fun(X6,fun(X6,X6)),X14)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X7),s(X6,X8),s(fun(X6,fun(X6,X6)),X9),s(fun(X6,fun(X6,X6)),X10))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X11),s(X6,X12),s(fun(X6,fun(X6,X6)),X13),s(fun(X6,fun(X6,X6)),X14))) <=> (s(X6,X7) = s(X6,X11) & s(X6,X8) = s(X6,X12) & s(fun(X6,fun(X6,X6)),X9) = s(fun(X6,fun(X6,X6)),X13) & s(fun(X6,fun(X6,X6)),X10) = s(fun(X6,fun(X6,X6)),X14)))),
  file('thm.semi_ring.semi_ring_|49||49|','thm.semi_ring.semi_ring_|49||49|')).
fof(f3662,plain,(
  $false | ($spl3 | $spl102 | $spl770)),
  inference(subsumption_resolution,[],[f3661,f343])).
fof(f3661,plain,(
  s(sK0,sK1) = s(sK0,sK5) | ($spl102 | $spl770)),
  inference(forward_demodulation,[],[f3631,f402])).
fof(f3631,plain,(
  s(sK0,sK5) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))))) | ($spl102 | $spl770)),
  inference(superposition,[],[f402,f3602])).
fof(f3584,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) | $spl742),
  inference(cnf_transformation,[],[f3584_D])).
fof(f3584_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) <=> ~$spl742),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl742])])).
fof(f455,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl0 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f454,f346])).
fof(f454,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f340,f348])).
fof(f3602,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) | $spl770),
  inference(cnf_transformation,[],[f3602_D])).
fof(f3602_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) <=> ~$spl770),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl770])])).
fof(f3598,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl4 | $spl742)),
  inference(backward_demodulation,[],[f344,f3584])).
fof(f13364,plain,(
  $false | ($spl7 | $spl106 | $spl836)),
  inference(subsumption_resolution,[],[f13363,f347])).
fof(f13363,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) | ($spl106 | $spl836)),
  inference(forward_demodulation,[],[f13308,f404])).
fof(f13308,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK7) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))))) | ($spl106 | $spl836)),
  inference(superposition,[],[f404,f3864])).
fof(f3671,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK4))) | $spl776),
  inference(cnf_transformation,[],[f3671_D])).
fof(f3671_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK4))) <=> ~$spl776),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl776])])).
fof(f3669,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl0 | $spl4 | $spl8)),
  inference(forward_demodulation,[],[f454,f344])).
fof(f3864,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK4))) | $spl836),
  inference(cnf_transformation,[],[f3864_D])).
fof(f3864_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK4))) <=> ~$spl836),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl836])])).
fof(f3851,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl2 | $spl776)),
  inference(backward_demodulation,[],[f342,f3671])).
fof(f15991,plain,(
  $false | ($spl9 | $spl108 | $spl1796)),
  inference(subsumption_resolution,[],[f15990,f349])).
fof(f15990,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl108 | $spl1796)),
  inference(forward_demodulation,[],[f15931,f405])).
fof(f15931,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))))) | ($spl108 | $spl1796)),
  inference(superposition,[],[f405,f14590])).
fof(f13370,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) | $spl1730),
  inference(cnf_transformation,[],[f13370_D])).
fof(f13370_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) <=> ~$spl1730),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1730])])).
fof(f13368,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f13367,f342])).
fof(f13367,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f340,f344])).
fof(f14590,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) | $spl1796),
  inference(cnf_transformation,[],[f14590_D])).
fof(f14590_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) <=> ~$spl1796),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1796])])).
fof(f14560,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK8))) | ($spl6 | $spl1730)),
  inference(backward_demodulation,[],[f346,f13370])).
fof(f17264,plain,(
  $false | ($spl5 | $spl104 | $spl1918)),
  inference(subsumption_resolution,[],[f17263,f345])).
fof(f17263,plain,(
  s(sK0,sK2) = s(sK0,sK6) | ($spl104 | $spl1918)),
  inference(forward_demodulation,[],[f17190,f403])).
fof(f17190,plain,(
  s(sK0,sK6) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))))) | ($spl104 | $spl1918)),
  inference(superposition,[],[f403,f15997])).
fof(f20271,plain,(
  $false | ($spl3 | $spl102 | $spl2038)),
  inference(subsumption_resolution,[],[f20270,f343])).
fof(f20270,plain,(
  s(sK0,sK1) = s(sK0,sK5) | ($spl102 | $spl2038)),
  inference(forward_demodulation,[],[f20196,f402])).
fof(f20196,plain,(
  s(sK0,sK5) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))))) | ($spl102 | $spl2038)),
  inference(superposition,[],[f402,f17276])).
fof(f340,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) | $spl0),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f348,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | $spl8),
  inference(cnf_transformation,[],[f348_D])).
fof(f348_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f234,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8)))),
  inference(cnf_transformation,[],[f201])).
fof(f346,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) | $spl6),
  inference(cnf_transformation,[],[f346_D])).
fof(f346_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f233,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8)))),
  inference(cnf_transformation,[],[f201])).
fof(f344,plain,(
  s(sK0,sK2) = s(sK0,sK6) | $spl4),
  inference(cnf_transformation,[],[f344_D])).
fof(f344_D,plain,(
  s(sK0,sK2) = s(sK0,sK6) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f232,plain,(
  s(sK0,sK2) = s(sK0,sK6) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8)))),
  inference(cnf_transformation,[],[f201])).
fof(f342,plain,(
  s(sK0,sK1) = s(sK0,sK5) | $spl2),
  inference(cnf_transformation,[],[f342_D])).
fof(f342_D,plain,(
  s(sK0,sK1) = s(sK0,sK5) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f231,plain,(
  s(sK0,sK1) = s(sK0,sK5) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK8)))),
  inference(cnf_transformation,[],[f201])).
fof(f453,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(trivial_inequality_removal,[],[f452])).
fof(f452,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f451,f342])).
fof(f451,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f450,f344])).
fof(f450,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl1 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f449,f346])).
fof(f449,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK5),s(sK0,sK6),s(fun(sK0,fun(sK0,sK0)),sK7),s(fun(sK0,fun(sK0,sK0)),sK4))) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f341,f348])).
fof(f20512,plain,(
  $false | ($spl0 | $spl7 | $spl106)),
  inference(subsumption_resolution,[],[f20511,f347])).
fof(f20511,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) | ($spl0 | $spl106)),
  inference(forward_demodulation,[],[f20437,f404])).
fof(f20437,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK7) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.recordtype|46|semi_ring_4'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4))))) | ($spl0 | $spl106)),
  inference(superposition,[],[f404,f340])).
% SZS output end Proof for thm.semi_ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 155
% Passive clauses: 1095
% Generated clauses: 20325
% Final active clauses: 123
% Final passive clauses: 909
% Input formulas: 69
% Initial clauses: 107
% 
% Duplicate literals: 273
% Trivial inequalities: 164
% Fw subsumption resolutions: 983
% Fw demodulations: 3040
% Bw demodulations: 33
% Condensations: 806
% Global subsumptions: 4
% 
% Simple tautologies: 110
% Equational tautologies: 21
% Forward subsumptions: 12193
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 80
% Factoring: 1
% Forward superposition: 5768
% Backward superposition: 7644
% Self superposition: 174
% Equality factoring: 29
% Equality resolution: 18
% 
% Split clauses: 64
% Split components: 131
% SAT solver clauses: 1591
% SAT solver unit clauses: 150
% SAT solver binary clauses: 56
% SAT solver learnt clauses: 15
% SAT solver learnt literals: 13
% 
% Sat splits: 64
% Sat splitting refutations: 56
% 
% TWLsolver clauses: 1567
% TWLsolver calls for satisfiability: 565
% 
% Memory used [KB]: 18166
% Time elapsed: 0.673 s
% ------------------------------
% ------------------------------
% Success in time 1.3 s
