Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (89812)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 850
% Passive clauses: 8317
% Generated clauses: 21277
% Final active clauses: 837
% Final passive clauses: 7059
% Input formulas: 41
% Initial clauses: 154
% 
% Duplicate literals: 592
% Trivial inequalities: 163
% Fw subsumption resolutions: 2391
% Fw demodulations: 509
% Bw demodulations: 12
% 
% Simple tautologies: 218
% Equational tautologies: 59
% Forward subsumptions: 8425
% 
% Binary resolution: 9875
% Factoring: 1444
% Forward superposition: 3807
% Backward superposition: 4935
% Self superposition: 108
% Equality factoring: 30
% Equality resolution: 22
% 
% Split clauses: 613
% Split components: 1232
% SAT solver clauses: 1020
% SAT solver unit clauses: 19
% SAT solver binary clauses: 130
% 
% Sat splits: 613
% 
% TWLsolver clauses: 1020
% TWLsolver calls for satisfiability: 494
% 
% Memory used [KB]: 10874
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (89824)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 452
% Passive clauses: 2906
% Generated clauses: 11636
% Final active clauses: 443
% Final passive clauses: 2453
% Input formulas: 41
% Initial clauses: 154
% 
% Duplicate literals: 240
% Trivial inequalities: 44
% Fw subsumption resolutions: 974
% Fw demodulations: 181
% Bw demodulations: 9
% Condensations: 215
% Global subsumptions: 13
% 
% Simple tautologies: 258
% Equational tautologies: 17
% Forward subsumptions: 3410
% 
% Binary resolution: 4216
% Factoring: 259
% Forward superposition: 1370
% Backward superposition: 996
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 8
% 
% Split clauses: 79
% Split components: 163
% SAT solver clauses: 3440
% SAT solver unit clauses: 170
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 79
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3388
% TWLsolver calls for satisfiability: 1731
% 
% Memory used [KB]: 13048
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (89856)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2685
% Passive clauses: 11687
% Generated clauses: 149517
% Final active clauses: 961
% Final passive clauses: 1563
% Input formulas: 41
% Initial clauses: 154
% 
% Duplicate literals: 3580
% Trivial inequalities: 234
% Fw subsumption resolutions: 5619
% Fw demodulations: 25501
% Bw demodulations: 5340
% Condensations: 5551
% Global subsumptions: 1592
% 
% Simple tautologies: 213
% Equational tautologies: 4453
% Forward subsumptions: 84137
% Bw demodulations to eq. taut.: 232
% 
% Unit resulting resolution: 62754
% Forward superposition: 14391
% Backward superposition: 10346
% Self superposition: 226
% Equality factoring: 172
% Equality resolution: 28
% 
% Split clauses: 253
% Split components: 522
% SAT solver clauses: 18064
% SAT solver unit clauses: 5900
% SAT solver binary clauses: 4908
% 
% Sat splits: 253
% Sat splitting refutations: 4276
% 
% Memory used [KB]: 114113
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% (90138)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2154
% Passive clauses: 43928
% Generated clauses: 187071
% Final active clauses: 2020
% Final passive clauses: 37072
% Input formulas: 41
% Initial clauses: 154
% 
% Selected by SInE selection: 41
% SInE iterations: 1
% 
% Duplicate literals: 5604
% Trivial inequalities: 231
% Fw subsumption resolutions: 20496
% Fw demodulations: 6547
% Bw demodulations: 16
% 
% Simple tautologies: 3753
% Equational tautologies: 281
% Forward subsumptions: 96739
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 131889
% Factoring: 678
% Forward superposition: 11184
% Backward superposition: 8580
% Self superposition: 7
% Equality factoring: 2
% Equality resolution: 17
% 
% Split clauses: 10888
% Split components: 25377
% SAT solver clauses: 12048
% SAT solver unit clauses: 38
% SAT solver binary clauses: 495
% 
% Sat splits: 10888
% Sat splitting refutations: 71
% 
% Memory used [KB]: 58591
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.arithmetic
% (90249)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 378
% Passive clauses: 2246
% Generated clauses: 14281
% Final active clauses: 371
% Final passive clauses: 1798
% Input formulas: 41
% Initial clauses: 154
% 
% Duplicate literals: 548
% Trivial inequalities: 230
% Fw subsumption resolutions: 2154
% Fw demodulations: 104
% Bw demodulations: 21
% Global subsumptions: 16
% 
% Simple tautologies: 570
% Equational tautologies: 49
% Forward subsumptions: 7924
% 
% Unit resulting resolution: 1700
% Factoring: 18
% Forward superposition: 4354
% Backward superposition: 4820
% Self superposition: 92
% Equality factoring: 6
% Equality resolution: 22
% 
% Split clauses: 21
% Split components: 42
% SAT solver clauses: 2336
% SAT solver unit clauses: 228
% SAT solver binary clauses: 152
% 
% Sat splits: 21
% Sat splitting refutations: 51
% 
% Memory used [KB]: 7036
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.arithmetic
% (90287)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1059
% Passive clauses: 6579
% Generated clauses: 19370
% Final active clauses: 723
% Final passive clauses: 4631
% Input formulas: 41
% Initial clauses: 154
% 
% Duplicate literals: 626
% Trivial inequalities: 64
% Fw subsumption resolutions: 1201
% Bw subsumption resolutions: 552
% Fw demodulations: 1008
% Bw demodulations: 99
% Condensations: 1344
% Global subsumptions: 2
% 
% Simple tautologies: 1172
% Equational tautologies: 56
% Forward subsumptions: 7147
% Backward subsumptions: 678
% 
% Binary resolution: 9492
% Factoring: 41
% Forward superposition: 1959
% Backward superposition: 1804
% Self superposition: 31
% Equality factoring: 4
% Equality resolution: 15
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 6675
% SAT solver unit clauses: 301
% SAT solver binary clauses: 1105
% 
% Sat splits: 1
% Sat splitting refutations: 90
% 
% Memory used [KB]: 11257
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.arithmetic
% (90346)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1072
% Passive clauses: 3340
% Generated clauses: 20985
% Final active clauses: 1045
% Final passive clauses: 2317
% Input formulas: 41
% Initial clauses: 154
% 
% Duplicate literals: 248
% Trivial inequalities: 64
% Fw demodulations: 4320
% Bw demodulations: 19
% Global subsumptions: 556
% 
% Simple tautologies: 647
% Equational tautologies: 11
% Forward subsumptions: 10908
% Backward subsumptions: 80
% 
% Binary resolution: 4226
% Unit resulting resolution: 8579
% Forward superposition: 2397
% Backward superposition: 442
% Equality resolution: 8
% 
% Split clauses: 34
% Split components: 68
% SAT solver clauses: 5056
% SAT solver unit clauses: 1480
% SAT solver binary clauses: 2724
% 
% Sat splits: 34
% Sat splitting refutations: 681
% 
% Memory used [KB]: 23539
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 109 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f46145,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f45959])).
fof(f45959,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.|48|_0'))) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3563,f45773])).
fof(f45773,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.|48|_0')),sK1))),
  inference(global_subsumption,[],[f45442,f691])).
fof(f691,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f637,f635])).
fof(f635,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X7)))))) )),
  inference(resolution,[],[f175,f196])).
fof(f196,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.bool.EQ_CLAUSES|split|1')).
fof(f175,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f102])).
fof(f102,plain,(
  ? [X0] : ! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ~! [X0] : ~! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))))))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.arithmetic.NOT_STRICTLY_DECREASING')).
fof(f637,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X9)))))) )),
  inference(resolution,[],[f175,f199])).
fof(f199,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.bool.EQ_CLAUSES|split|3')).
fof(f45442,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.|48|_0')),sK1))),
  inference(superposition,[],[f11840,f4667])).
fof(f4667,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f3707,f185])).
fof(f185,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f48])).
fof(f48,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) | ? [X7] : s('type.num.num',X8) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.arithmetic.num_CASES')).
fof(f3707,plain,(
  ( ! [X180] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X180)))))) )),
  inference(resolution,[],[f1076,f176])).
fof(f176,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING',aHL_TRUTH)).
fof(f1076,plain,(
  ( ! [X177,X178] : (~p(s(bool,X178)) | s(bool,X178) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X177)))))) )),
  inference(resolution,[],[f696,f183])).
fof(f183,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.prim_rec.LESS_|48|')).
fof(f696,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | s(bool,X44) = s(bool,X45) | ~p(s(bool,X45))) )),
  inference(global_subsumption,[],[f652,f641])).
fof(f641,plain,(
  ( ! [X12,X13,X11] : (sP0(X11,X12,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X13)))))) )),
  inference(resolution,[],[f175,f258])).
fof(f258,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f652,plain,(
  ( ! [X45,X46,X44] : (s(bool,X44) = s(bool,X45) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | ~sP0(X44,X45,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X46)))))) )),
  inference(resolution,[],[f175,f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP0(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f269])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | (sP0(X0,X1,X2) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | (~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | (sP0(X0,X1,X2) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | (~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> (sP0(X0,X1,X2) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(definition_folding,[],[f95,f132])).
fof(f95,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> s(bool,X14) = s(bool,X15)) <=> ((p(s(bool,X16)) | p(s(bool,X15)) | p(s(bool,X14))) & (p(s(bool,X16)) | ~p(s(bool,X14)) | ~p(s(bool,X15))) & (p(s(bool,X15)) | ~p(s(bool,X14)) | ~p(s(bool,X16))) & (p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X16)))))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.sat.sat_2|49|_prove_ltb')).
fof(f11840,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK3(happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X5)),sK1))))) )),
  inference(resolution,[],[f3561,f2068])).
fof(f2068,plain,(
  ( ! [X12,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X11))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X12)))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',X11)))) )),
  inference(resolution,[],[f630,f200])).
fof(f200,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f144])).
fof(f630,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X0))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',X1))))))) )),
  inference(resolution,[],[f175,f194])).
fof(f194,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X0)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK5(X0)))))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X0)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK5(X0))))))) | ! [X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f140])).
fof(f140,plain,(
  ! [X0] : (? [X3] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3)))))) | ! [X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))))))))),
  inference(rectify,[],[f110])).
fof(f110,plain,(
  ! [X0] : (? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1)))))) | ! [X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1)))))) => ! [X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2))))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7)))))) => ! [X8,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X8))))))))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.arithmetic.arithmetic_357_prove_ltb')).
fof(f3561,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',sK3(X1,sK1)))),s('type.num.num',X1))))) )),
  inference(resolution,[],[f735,f255])).
fof(f255,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X7] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.arithmetic.LESS_ANTISYM')).
fof(f735,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',sK3(X0,sK1)))))))) )),
  inference(resolution,[],[f629,f187])).
fof(f187,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK3(X1,X0))))))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK3(X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f105])).
fof(f105,plain,(
  ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) | ! [X1] : ? [X2] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2)))))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) => ! [X1] : ? [X2] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X2)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X6)))) => ! [X9] : ? [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7)))))))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.arithmetic.ONE_ONE_UNBOUNDED')).
fof(f629,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),sK1))))),
  inference(resolution,[],[f175,f193])).
fof(f193,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X0)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK4(X0)))))))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X0)))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',sK4(X0))))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f109])).
fof(f109,plain,(
  ! [X0] : (? [X1] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1)))))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (! [X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1)))))) => p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X0)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : (! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X6),s('type.num.num',X7)))))) => p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.num.num','type.num.num'),X6)))))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.arithmetic.arithmetic_358_prove_ltb')).
fof(f3563,plain,(
  ( ! [X3] : (s('type.num.num',X3) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK1),s('type.num.num',sK3(X3,sK1))))) )),
  inference(resolution,[],[f735,f223])).
fof(f223,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => s('type.num.num',X0) != s('type.num.num',X1))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => ~s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X8,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => ~s('type.num.num',X7) = s('type.num.num',X8))),
  file('thm.arithmetic.NOT_STRICTLY_DECREASING','thm.prim_rec.LESS_NOT_EQ')).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 421
% Passive clauses: 3122
% Generated clauses: 45074
% Final active clauses: 405
% Final passive clauses: 2570
% Input formulas: 41
% Initial clauses: 283
% 
% Duplicate literals: 3283
% Trivial inequalities: 890
% Fw demodulations: 2647
% Bw demodulations: 74
% Condensations: 5035
% Global subsumptions: 70
% 
% Simple tautologies: 2318
% Equational tautologies: 115
% Forward subsumptions: 27959
% Bw demodulations to eq. taut.: 128
% 
% Binary resolution: 12744
% Factoring: 13
% Forward superposition: 10903
% Backward superposition: 9479
% Self superposition: 105
% Equality factoring: 98
% Equality resolution: 65
% 
% SAT solver clauses: 3408
% SAT solver unit clauses: 163
% SAT solver binary clauses: 356
% SAT solver learnt clauses: 140
% SAT solver learnt literals: 312
% 
% TWLsolver clauses: 3192
% TWLsolver calls for satisfiability: 95926
% 
% Memory used [KB]: 22259
% Time elapsed: 2.242 s
% ------------------------------
% ------------------------------
% Success in time 21.309 s
