Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (82317)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 52
% Passive clauses: 1972
% Generated clauses: 9631
% Final active clauses: 50
% Final passive clauses: 1869
% Input formulas: 23
% Initial clauses: 24
% 
% Duplicate literals: 601
% Trivial inequalities: 100
% Fw subsumption resolutions: 997
% Fw demodulations: 100
% 
% Simple tautologies: 192
% Equational tautologies: 109
% Forward subsumptions: 4422
% 
% Binary resolution: 25
% Forward superposition: 3891
% Backward superposition: 3174
% Self superposition: 400
% Equality factoring: 326
% Equality resolution: 10
% 
% Split clauses: 375
% Split components: 754
% SAT solver clauses: 1190
% SAT solver unit clauses: 251
% SAT solver binary clauses: 403
% 
% Sat splits: 375
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1186
% TWLsolver calls for satisfiability: 34
% 
% Memory used [KB]: 6780
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (82329)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 38
% Passive clauses: 249
% Generated clauses: 5978
% Final active clauses: 35
% Final passive clauses: 209
% Input formulas: 23
% Initial clauses: 24
% 
% Duplicate literals: 835
% Trivial inequalities: 53
% Fw subsumption resolutions: 154
% Fw demodulations: 49
% Bw demodulations: 1
% Condensations: 1686
% 
% Simple tautologies: 96
% Equational tautologies: 57
% Forward subsumptions: 2192
% 
% Binary resolution: 16
% Forward superposition: 1564
% Backward superposition: 1093
% Self superposition: 214
% Equality factoring: 174
% Equality resolution: 5
% 
% Split clauses: 29
% Split components: 58
% SAT solver clauses: 650
% SAT solver unit clauses: 42
% SAT solver binary clauses: 66
% 
% Sat splits: 29
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 650
% TWLsolver calls for satisfiability: 56
% 
% Memory used [KB]: 5756
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 285 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f13000,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f12999])).
fof(f12999,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1)))),
  inference(superposition,[],[f71,f12788])).
fof(f12788,plain,(
  ( ! [X6,X4,X5] : (s(X4,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X4),X5))) = s(X4,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X4),X5),s('type.quote.varmap'(X4),X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f12629,f10396])).
fof(f10396,plain,(
  ( ! [X4,X5] : (s(X4,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X4),X5))) = s(X4,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5)))))) )),
  inference(superposition,[],[f88,f77])).
fof(f77,plain,(
  ( ! [X0,X1] : (s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7] : s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))),s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_vl_def|split|0','thm.ring.semi_ring_of_def')).
fof(f88,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X13,X14,X15,X16] : s(X6,X14) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X13),s(X6,X14),s(fun(X6,fun(X6,X6)),X15),s(fun(X6,fun(X6,X6)),X16)))))),
  file('thm.ringNorm.interp_vl_def|split|0','thm.semi_ring.semi_ring_SR|49|')).
fof(f12629,plain,(
  ( ! [X6,X4,X5] : (s(X4,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))))) = s(X4,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X4),X5),s('type.quote.varmap'(X4),X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) )),
  inference(superposition,[],[f8582,f78])).
fof(f78,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11,X8] : s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X11))) = s(X6,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X6),X11),s('type.quote.varmap'(X6),X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  file('thm.ringNorm.interp_vl_def|split|0','thm.canonical.interp_vl_def|split|0')).
fof(f8582,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(forward_demodulation,[],[f8454,f86])).
fof(f86,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X17,X18,X19] : s(X6,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X6),X17),s('type.quote.varmap'(X6),X18),s('type.list.list'('type.quote.index'),X19))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.list.list'('type.quote.index'),X6)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X6),X17))),s('type.quote.varmap'(X6),X18))),s('type.list.list'('type.quote.index'),X19)))),
  file('thm.ringNorm.interp_vl_def|split|0','ARITY_1_3_const.ringNorm.r_interp_vl')).
fof(f8454,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(superposition,[],[f85,f75])).
fof(f75,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.list.list'('type.quote.index'),X6)),'const.ringNorm.r_interp_vl_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.list.list'('type.quote.index'),X6)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_vl_def|split|0','thm.ringNorm.r_interp_vl_def')).
fof(f85,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X17,X18,X19] : s(X6,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X6),X17),s('type.quote.varmap'(X6),X18),s('type.list.list'('type.quote.index'),X19))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.list.list'('type.quote.index'),X6)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X6),X17))),s('type.quote.varmap'(X6),X18))),s('type.list.list'('type.quote.index'),X19)))),
  file('thm.ringNorm.interp_vl_def|split|0','ARITY_1_3_const.canonical.interp_vl')).
fof(f71,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),sK1))) != s(sK0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ~! [X0,X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ~! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))) = s(X6,'const.ringNorm.r_interp_vl_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  file('thm.ringNorm.interp_vl_def|split|0','thm.ringNorm.interp_vl_def|split|0')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 57
% Passive clauses: 740
% Generated clauses: 12669
% Final active clauses: 54
% Final passive clauses: 581
% Input formulas: 23
% Initial clauses: 24
% 
% Duplicate literals: 1557
% Trivial inequalities: 69
% Fw subsumption resolutions: 325
% Fw demodulations: 219
% Bw demodulations: 95
% Condensations: 2934
% 
% Simple tautologies: 152
% Equational tautologies: 131
% Forward subsumptions: 6646
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 76
% Forward superposition: 3362
% Backward superposition: 3635
% Self superposition: 333
% Equality factoring: 246
% Equality resolution: 6
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 795
% SAT solver unit clauses: 34
% SAT solver binary clauses: 58
% 
% Sat splits: 26
% Sat splitting refutations: 1
% 
% Memory used [KB]: 10746
% Time elapsed: 1.718 s
% ------------------------------
% ------------------------------
% Success in time 3.125 s
