Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (101732)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 88
% Passive clauses: 1539
% Generated clauses: 13563
% Final active clauses: 85
% Final passive clauses: 1351
% Input formulas: 19
% Initial clauses: 20
% 
% Duplicate literals: 938
% Trivial inequalities: 186
% Fw subsumption resolutions: 1854
% Fw demodulations: 350
% 
% Simple tautologies: 296
% Equational tautologies: 98
% Forward subsumptions: 6396
% 
% Binary resolution: 51
% Forward superposition: 6255
% Backward superposition: 3400
% Self superposition: 390
% Equality factoring: 283
% Equality resolution: 14
% 
% Split clauses: 750
% Split components: 1514
% SAT solver clauses: 1976
% SAT solver unit clauses: 244
% SAT solver binary clauses: 594
% 
% Sat splits: 750
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1968
% TWLsolver calls for satisfiability: 74
% 
% Memory used [KB]: 6780
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (101750)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 93
% Passive clauses: 295
% Generated clauses: 10998
% Final active clauses: 88
% Final passive clauses: 196
% Input formulas: 19
% Initial clauses: 20
% 
% Duplicate literals: 1244
% Trivial inequalities: 83
% Fw subsumption resolutions: 680
% Fw demodulations: 671
% Bw demodulations: 7
% Condensations: 1953
% 
% Simple tautologies: 172
% Equational tautologies: 43
% Forward subsumptions: 5026
% 
% Binary resolution: 14
% Forward superposition: 2931
% Backward superposition: 2813
% Self superposition: 209
% Equality factoring: 132
% Equality resolution: 4
% 
% Split clauses: 89
% Split components: 183
% SAT solver clauses: 844
% SAT solver unit clauses: 34
% SAT solver binary clauses: 67
% 
% Sat splits: 89
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 844
% TWLsolver calls for satisfiability: 96
% 
% Memory used [KB]: 8955
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f14840,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f14818])).
fof(f14818,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3)))))),
  inference(superposition,[],[f59,f1170])).
fof(f1170,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X4),X6),s('type.canonical.spolynom'(X4),X7))))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X6))),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X7)))))) )),
  inference(forward_demodulation,[],[f1169,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(forward_demodulation,[],[f287,f68])).
fof(f68,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X15,X16] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X15),s('type.canonical.spolynom'(X6),X16))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X6),X15))),s('type.canonical.spolynom'(X6),X16)))),
  file('thm.ringNorm.spolynom_normalize_def|split|2','ARITY_1_2_const.ringNorm.r_spolynom_normalize')).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(superposition,[],[f67,f63])).
fof(f63,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|2','thm.ringNorm.r_spolynom_normalize_def')).
fof(f67,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X15,X16] : s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X15),s('type.canonical.spolynom'(X6),X16))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X6),X15))),s('type.canonical.spolynom'(X6),X16)))),
  file('thm.ringNorm.spolynom_normalize_def|split|2','ARITY_1_2_const.canonical.spolynom_normalize')).
fof(f1169,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X4),X6),s('type.canonical.spolynom'(X4),X7))))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X6))),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X7)))))) )),
  inference(forward_demodulation,[],[f1168,f289])).
fof(f1168,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X4),X6),s('type.canonical.spolynom'(X4),X7))))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),X6))),s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X4),X5),s('type.canonical.spolynom'(X4),X7)))))) )),
  inference(forward_demodulation,[],[f1153,f289])).
fof(f1153,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X4),X6),s('type.canonical.spolynom'(X4),X7))))) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),X6))),s('type.canonical.canonical_sum'(X4),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X4),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X4),X5))),s('type.canonical.spolynom'(X4),X7)))))) )),
  inference(superposition,[],[f542,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X12,X8,X7] : s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7))))) = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.spolynom'(X6),X8))),s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X12),s('type.canonical.spolynom'(X6),X7)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|2','thm.canonical.spolynom_normalize_def|split|2')).
fof(f542,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f500,f74])).
fof(f74,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X15),s('type.canonical.canonical_sum'(X6),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X15))),s('type.canonical.canonical_sum'(X6),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.spolynom_normalize_def|split|2','ARITY_1_3_const.ringNorm.r_canonical_sum_merge')).
fof(f500,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f72,f64])).
fof(f64,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|2','thm.ringNorm.r_canonical_sum_merge_def')).
fof(f72,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X15),s('type.canonical.canonical_sum'(X6),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),X15))),s('type.canonical.canonical_sum'(X6),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.spolynom_normalize_def|split|2','ARITY_1_3_const.canonical.canonical_sum_merge')).
fof(f59,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK3)))))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3))))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f54])).
fof(f54,plain,(
  ? [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X3))))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X9))))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X8))),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X9))))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X8))),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X9)))))),
  file('thm.ringNorm.spolynom_normalize_def|split|2','thm.ringNorm.spolynom_normalize_def|split|2')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 116
% Passive clauses: 403
% Generated clauses: 14536
% Final active clauses: 87
% Final passive clauses: 216
% Input formulas: 19
% Initial clauses: 20
% 
% Duplicate literals: 1738
% Trivial inequalities: 147
% Fw subsumption resolutions: 690
% Fw demodulations: 435
% Bw demodulations: 71
% Condensations: 2862
% Global subsumptions: 4
% 
% Simple tautologies: 195
% Equational tautologies: 115
% Forward subsumptions: 7945
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 138
% Forward superposition: 4078
% Backward superposition: 3816
% Self superposition: 302
% Equality factoring: 196
% Equality resolution: 11
% 
% Split clauses: 70
% Split components: 146
% SAT solver clauses: 560
% SAT solver unit clauses: 115
% SAT solver binary clauses: 28
% 
% Sat splits: 70
% Sat splitting refutations: 11
% 
% Memory used [KB]: 9466
% Time elapsed: 0.867 s
% ------------------------------
% ------------------------------
% Success in time 2.184 s
