Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (93551)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 104
% Passive clauses: 2199
% Generated clauses: 20525
% Final active clauses: 90
% Final passive clauses: 1853
% Input formulas: 23
% Initial clauses: 51
% 
% Duplicate literals: 1048
% Trivial inequalities: 288
% Fw subsumption resolutions: 2901
% Fw demodulations: 2486
% Bw demodulations: 7
% 
% Simple tautologies: 351
% Equational tautologies: 73
% Forward subsumptions: 8639
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 288
% Factoring: 5
% Forward superposition: 5904
% Backward superposition: 6859
% Self superposition: 303
% Equality factoring: 219
% Equality resolution: 20
% 
% Split clauses: 901
% Split components: 1817
% SAT solver clauses: 1639
% SAT solver unit clauses: 9
% SAT solver binary clauses: 118
% 
% Sat splits: 901
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 1635
% TWLsolver calls for satisfiability: 144
% 
% Memory used [KB]: 9083
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (93573)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 68
% Passive clauses: 553
% Generated clauses: 8722
% Final active clauses: 61
% Final passive clauses: 476
% Input formulas: 23
% Initial clauses: 51
% 
% Duplicate literals: 732
% Trivial inequalities: 111
% Fw subsumption resolutions: 783
% Fw demodulations: 462
% Bw demodulations: 10
% Condensations: 1225
% Global subsumptions: 9
% 
% Simple tautologies: 181
% Equational tautologies: 51
% Forward subsumptions: 3745
% 
% Binary resolution: 175
% Factoring: 3
% Forward superposition: 2342
% Backward superposition: 1981
% Self superposition: 169
% Equality factoring: 131
% Equality resolution: 10
% 
% Split clauses: 72
% Split components: 145
% SAT solver clauses: 960
% SAT solver unit clauses: 57
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 18
% 
% Sat splits: 72
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 924
% TWLsolver calls for satisfiability: 184
% 
% Memory used [KB]: 6780
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f11828,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f148,f151,f149,f174,f178,f176,f187,f230,f234,f232,f423,f427,f425,f1056,f1066,f1057,f1069,f1067,f1107,f1110,f1111,f1113,f1115,f1116,f1118,f1120,f1121,f1122,f1163,f1173,f1164,f1176,f1174,f1216,f1219,f1220,f1222,f1224,f1226,f1227,f1229,f1231,f1233,f1234,f1235,f1166,f1165,f1237,f1171,f1281,f1290,f1282,f1283,f1284,f1294,f1287,f1285,f1293,f1288,f1291,f1498,f1501,f1515,f1513,f2775,f2776,f2777,f2778,f2779,f2782,f2785,f2787,f2790,f4295,f4323,f4324,f4325,f4328,f4331,f4334,f4337,f4340,f4343,f1058,f1059,f4344,f1064,f7098,f7129,f7130,f7131,f7134,f7137,f7140,f7143,f7146,f7149,f7152,f9991,f10051,f10182,f10213,f10223,f9992,f175,f11569,f11592,f11593,f11594,f11595,f11605,f11606,f11607,f11608,f11609,f11610,f11611,f11612,f11613,f11614,f11615,f11616,f11617,f11618,f11619,f11620,f11621,f11622,f11623,f11624,f11625,f11626,f11627,f11628,f11629,f11630,f11631,f11632,f11633,f11634,f11635,f11636,f11637,f11638,f11639,f11640,f11641,f11642,f11643,f11644,f11645,f11646,f11647,f11648,f11649,f11650,f11651,f11652,f11653,f11654,f11655,f11656,f11657,f11658,f11659,f11660,f11661,f11662,f11663,f11664,f11665,f11666,f11667,f11668,f11669,f11670,f11671,f11672,f11673,f11674,f11675,f11676,f11677,f11678,f11679,f11680,f11681,f11682,f11683,f11684,f11685,f11686,f11687,f11688,f11689,f11690,f11691,f11692,f11693,f11694,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11708,f11709,f11710,f11711,f11712,f11713,f11714,f11715,f11716,f11717,f11718,f11719,f11720,f11721,f11722,f11723,f11724,f11725,f11726,f11727,f11728,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11741,f11742,f11743,f11744,f11745,f11746,f11747,f11748,f11749,f11750,f11751,f11752,f11753,f11754,f11755,f11756,f11757,f11758,f11759,f11760,f11761,f11762,f11763,f11764,f11765,f11766,f11767,f11768,f11769,f11770,f11771,f11772,f11773,f11774,f11775,f11776,f11777,f11778,f11779,f11780,f11781,f11782,f11783,f11784,f11785,f11786,f11787,f11788,f11789,f11790,f11791,f11792,f11793,f11794,f11795,f11796,f11797,f11798,f11799,f11800,f11801,f11802,f11803,f11804,f11805,f11806,f11807,f11808,f11809,f11810,f11811,f11812,f11813,f11814,f11815,f11816,f11817,f11818,f11819,f11820,f11821,f11822,f11823,f11824,f11825,f11826,f11827])).
fof(f11827,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228,f163,f89,f4663,f4660,f1060,f1069])).
fof(f1060,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f409,f140])).
fof(f4660,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f409,f140])).
fof(f4663,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f409,f140])).
fof(f89,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.UNION_EMPTY|split|0',aHL_FALSITY)).
fof(f163,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f89,f126,f132])).
fof(f228,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) )),
  inference(superposition,[],[f136,f90])).
fof(f229,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f136,f90])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f130,f90])).
fof(f277,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f130,f90])).
fof(f353,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X8)))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f131,f90])).
fof(f411,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f128,f163])).
fof(f416,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK2('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f128,f163])).
fof(f444,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X8)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f129,f90])).
fof(f1372,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f283,f90])).
fof(f1380,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),X31)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X31),s(fun(bool,bool),X32))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f283,f90])).
fof(f1647,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,X5) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f1635,f87])).
fof(f1635,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,X5) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f276,f140])).
fof(f1746,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f1741,f87])).
fof(f1741,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f277,f140])).
fof(f1838,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f1833,f87])).
fof(f1833,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f353,f140])).
fof(f2194,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2179,f87])).
fof(f2179,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f444,f140])).
fof(f1511,plain,(
  ( ! [X4,X2,X3] : (s(fun(X2,bool),X3) = s(fun(X2,bool),X4) | s(bool,'T_0') = s(bool,g(s(X2,sK3(X3,X4,bool,X2)))) | s(bool,'T_0') = s(bool,f(s(X2,sK3(X3,X4,bool,X2))))) )),
  inference(trivial_inequality_removal,[],[f1505])).
fof(f1505,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X2,bool),X3) = s(fun(X2,bool),X4) | s(bool,'T_0') = s(bool,g(s(X2,sK3(X3,X4,bool,X2)))) | s(bool,'T_0') = s(bool,f(s(X2,sK3(X3,X4,bool,X2))))) )),
  inference(superposition,[],[f229,f90])).
fof(f692,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,'T_0') != s(bool,f(s(X87,sK3(X88,X89,bool,X87)))) | s(fun(X87,bool),X88) = s(fun(X87,bool),X89) | s(bool,X90) = s(bool,g(s(X87,sK3(X88,X89,bool,X87)))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f136,f140])).
fof(f611,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X95) != s(bool,f(s(X92,sK3(X93,X94,bool,X92)))) | s(fun(X92,bool),X93) = s(fun(X92,bool),X94) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,g(s(X92,sK3(X93,X94,bool,X92))))) )),
  inference(superposition,[],[f136,f140])).
fof(f128,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X7,X10] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.pred_set.UNION_EMPTY|split|0','thm.pred_set.EXTENSION')).
fof(f127,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f82])).
fof(f129,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X7,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.UNION_EMPTY|split|0','thm.pred_set.IN_UNION')).
fof(f130,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f131,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f283,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) )),
  inference(subsumption_resolution,[],[f280,f89])).
fof(f280,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) )),
  inference(superposition,[],[f130,f90])).
fof(f890,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),X40)))) | p(s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),X41)))) | s(bool,X42) = s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),'const.pred_set.UNION_2'(s(fun(X38,bool),X40),s(fun(X38,bool),X41))))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f681,f87])).
fof(f681,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),X40)))) | p(s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),X41)))) | s(bool,X42) = s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),'const.pred_set.UNION_2'(s(fun(X38,bool),X40),s(fun(X38,bool),X41))))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f129,f140])).
fof(f1399,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.pred_set.UNION_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) | s(bool,X11) = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X9))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f1388,f87])).
fof(f1388,plain,(
  ( ! [X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.pred_set.UNION_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) | s(bool,X11) = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X9))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f283,f140])).
fof(f656,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X23),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X24),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f129,f140])).
fof(f658,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,X31),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X32),s(fun(bool,bool),X33)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X33)))) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f131,f140])).
fof(f659,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,'const.bool.IN_2'(s(bool,X35),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X36),s(fun(bool,bool),X37)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X36)))) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f130,f140])).
fof(f693,plain,(
  ( ! [X92,X93,X91] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X91))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X92))) | s(fun(bool,bool),X91) = s(fun(bool,bool),X92) | s(bool,X93) = s(bool,sK2(X91,X92,bool)) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f128,f140])).
fof(f1373,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f283,f140])).
fof(f606,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'const.bool.IN_2'(s(bool,X71),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X72),s(fun(bool,bool),X73)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X70),s(fun(bool,bool),X72)))) | p(s(bool,'const.bool.IN_2'(s(bool,X70),s(fun(bool,bool),X73)))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f129,f140])).
fof(f608,plain,(
  ( ! [X80,X78,X81,X79] : (p(s(bool,'const.bool.IN_2'(s(bool,X79),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X80),s(fun(bool,bool),X81)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X78),s(fun(bool,bool),X81)))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f131,f140])).
fof(f609,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(bool,X83),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X84),s(fun(bool,bool),X85)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),X84)))) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f130,f140])).
fof(f612,plain,(
  ( ! [X97,X98,X96] : (s(bool,'const.bool.IN_2'(s(bool,X98),s(fun(bool,bool),X96))) != s(bool,'const.bool.IN_2'(s(bool,X98),s(fun(bool,bool),X97))) | s(fun(bool,bool),X96) = s(fun(bool,bool),X97) | s(bool,'T_0') = s(bool,X98) | s(bool,'T_0') = s(bool,sK2(X96,X97,bool))) )),
  inference(superposition,[],[f128,f140])).
fof(f687,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X67),s(fun(bool,bool),X68)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),X67)))) | p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),X68)))) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f129,f140])).
fof(f689,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X75),s(fun(bool,bool),X76)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X74),s(fun(bool,bool),X76)))) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f131,f140])).
fof(f690,plain,(
  ( ! [X80,X78,X81,X79] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X79),s(fun(bool,bool),X80)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X78),s(fun(bool,bool),X79)))) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f130,f140])).
fof(f1374,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X9)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X9),s(fun(bool,bool),X10))))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f283,f140])).
fof(f1378,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X23)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f283,f140])).
fof(f86,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f51])).
fof(f51,plain,(
  ? [X0,X1] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ~! [X0,X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X7)))),
  file('thm.pred_set.UNION_EMPTY|split|0','thm.pred_set.UNION_EMPTY|split|0')).
fof(f8010,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))))) ) | $spl50),
  inference(subsumption_resolution,[],[f7949,f87])).
fof(f7949,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))))) ) | $spl50),
  inference(superposition,[],[f283,f5345])).
fof(f5345,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl50),
  inference(condensation,[],[f5344])).
fof(f5344,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl50),
  inference(subsumption_resolution,[],[f5320,f90])).
fof(f5320,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl50),
  inference(superposition,[],[f4344,f90])).
fof(f9811,plain,(
  ( ! [X12,X10,X11] : (s(bool,X12) != s(bool,'const.bool.IN_2'(s(X10,sK2('const.pred_set.EMPTY_0',X11,X10)),s(fun(X10,bool),X11))) | s(fun(X10,bool),'const.pred_set.EMPTY_0') = s(fun(X10,bool),X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f128,f225])).
fof(f9810,plain,(
  ( ! [X8,X7,X9] : (s(bool,X9) != s(bool,'const.bool.IN_2'(s(X7,sK2(X8,'const.pred_set.EMPTY_0',X7)),s(fun(X7,bool),X8))) | s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f128,f225])).
fof(f8418,plain,(
  ( ! [X10] : (s(bool,sK2('const.pred_set.EMPTY_0',X10,bool)) != s(bool,'const.bool.IN_2'(s(bool,sK2('const.pred_set.EMPTY_0',X10,bool)),s(fun(bool,bool),X10))) | s(fun(bool,bool),'const.pred_set.EMPTY_0') = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK2('const.pred_set.EMPTY_0',X10,bool))) )),
  inference(superposition,[],[f128,f207])).
fof(f8417,plain,(
  ( ! [X9] : (s(bool,sK2(X9,'const.pred_set.EMPTY_0',bool)) != s(bool,'const.bool.IN_2'(s(bool,sK2(X9,'const.pred_set.EMPTY_0',bool)),s(fun(bool,bool),X9))) | s(fun(bool,bool),'const.pred_set.EMPTY_0') = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,sK2(X9,'const.pred_set.EMPTY_0',bool))) )),
  inference(superposition,[],[f128,f207])).
fof(f410,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))))),
  inference(unit_resulting_resolution,[],[f86,f128])).
fof(f409,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))))),
  inference(unit_resulting_resolution,[],[f86,f128])).
fof(f2456,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(trivial_inequality_removal,[],[f2449])).
fof(f2449,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(superposition,[],[f1069,f140])).
fof(f2447,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl23),
  inference(unit_resulting_resolution,[],[f87,f1069,f135])).
fof(f226,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f86,f136])).
fof(f694,plain,(
  ( ! [X94,X97,X95,X96] : (s(X96,f(s(bool,'T_0'))) != s(X96,g(s(bool,'T_0'))) | s(fun(bool,X96),X94) = s(fun(bool,X96),X95) | s(bool,X97) = s(bool,sK3(X94,X95,X96,bool)) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f136,f140])).
fof(f686,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X62,sK2(X63,X64,X62)),s(fun(X62,bool),X64))) | s(fun(X62,bool),X63) = s(fun(X62,bool),X64) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(X62,sK2(X63,X64,X62)),s(fun(X62,bool),X63))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f128,f140])).
fof(f685,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X58,sK2(X59,X60,X58)),s(fun(X58,bool),X59))) | s(fun(X58,bool),X59) = s(fun(X58,bool),X60) | s(bool,X61) = s(bool,'const.bool.IN_2'(s(X58,sK2(X59,X60,X58)),s(fun(X58,bool),X60))) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f128,f140])).
fof(f889,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f664,f87])).
fof(f664,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f135,f140])).
fof(f613,plain,(
  ( ! [X101,X99,X102,X100] : (s(X101,f(s(bool,X102))) != s(X101,g(s(bool,X102))) | s(fun(bool,X101),X99) = s(fun(bool,X101),X100) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,sK3(X99,X100,X101,bool))) )),
  inference(superposition,[],[f136,f140])).
fof(f605,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,X69) != s(bool,'const.bool.IN_2'(s(X66,sK2(X67,X68,X66)),s(fun(X66,bool),X68))) | s(fun(X66,bool),X67) = s(fun(X66,bool),X68) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X66,sK2(X67,X68,X66)),s(fun(X66,bool),X67)))) )),
  inference(superposition,[],[f128,f140])).
fof(f604,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X65) != s(bool,'const.bool.IN_2'(s(X62,sK2(X63,X64,X62)),s(fun(X62,bool),X63))) | s(fun(X62,bool),X63) = s(fun(X62,bool),X64) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X62,sK2(X63,X64,X62)),s(fun(X62,bool),X64)))) )),
  inference(superposition,[],[f128,f140])).
fof(f549,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f140,f140])).
fof(f227,plain,(
  s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0))))),
  inference(unit_resulting_resolution,[],[f86,f136])).
fof(f225,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f224,f89])).
fof(f224,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f177,f163])).
fof(f207,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f163,f90])).
fof(f177,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f169,f89])).
fof(f169,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f132,f90])).
fof(f164,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f126,f126,f132])).
fof(f142,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f89,f90])).
fof(f140,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f90,f90])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.UNION_EMPTY|split|0',aHL_EXT)).
fof(f132,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.UNION_EMPTY|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f135,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f54])).
fof(f126,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X11] : ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.UNION_EMPTY|split|0','thm.pred_set.NOT_IN_EMPTY')).
fof(f90,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.UNION_EMPTY|split|0',aHL_BOOL_CASES)).
fof(f87,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.UNION_EMPTY|split|0',aHL_TRUTH)).
fof(f11826,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228,f163,f89,f4663,f4660,f1060])).
fof(f11825,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228,f163,f89,f4663,f4660])).
fof(f11824,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228,f163,f89,f4663])).
fof(f11823,plain,(
  $false | ($spl17 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f6505,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228,f163,f89])).
fof(f6505,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(trivial_inequality_removal,[],[f6496])).
fof(f6496,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(superposition,[],[f1063,f140])).
fof(f1063,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | $spl17),
  inference(cnf_transformation,[],[f1063_D])).
fof(f1063_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f11822,plain,(
  $false | ($spl17 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f6504,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228,f163,f89])).
fof(f6504,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl17),
  inference(trivial_inequality_removal,[],[f6499])).
fof(f6499,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl17),
  inference(superposition,[],[f1063,f140])).
fof(f11821,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228,f163,f89])).
fof(f11820,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228,f163])).
fof(f11819,plain,(
  $false | ($spl6 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f176,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228])).
fof(f11818,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229,f228])).
fof(f11817,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276,f229])).
fof(f11816,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277,f276])).
fof(f11815,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353,f277])).
fof(f11814,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411,f353])).
fof(f11813,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416,f411])).
fof(f11812,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444,f416])).
fof(f11811,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372,f444])).
fof(f11810,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f468,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f468,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl15),
  inference(unit_resulting_resolution,[],[f438,f90])).
fof(f438,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl15),
  inference(backward_demodulation,[],[f432,f227])).
fof(f432,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl15),
  inference(unit_resulting_resolution,[],[f89,f427,f177])).
fof(f11809,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1035,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f1035,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X40,sK2(X41,X42,X40)),s(fun(X40,bool),X41))) | s(fun(X40,bool),X41) = s(fun(X40,bool),X42) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X40,sK2(X41,X42,X40)),s(fun(X40,bool),X42)))) ) | $spl15),
  inference(forward_demodulation,[],[f968,f468])).
fof(f968,plain,(
  ( ! [X41,X42,X40] : (s(bool,'const.bool.IN_2'(s(X40,sK2(X41,X42,X40)),s(fun(X40,bool),X41))) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(fun(X40,bool),X41) = s(fun(X40,bool),X42) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X40,sK2(X41,X42,X40)),s(fun(X40,bool),X42)))) ) | $spl15),
  inference(superposition,[],[f128,f496])).
fof(f496,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl15),
  inference(subsumption_resolution,[],[f484,f89])).
fof(f484,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl15),
  inference(superposition,[],[f177,f468])).
fof(f11808,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1036,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f1036,plain,(
  ( ! [X45,X43,X44] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X43,sK2(X44,X45,X43)),s(fun(X43,bool),X45))) | s(fun(X43,bool),X44) = s(fun(X43,bool),X45) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X43,sK2(X44,X45,X43)),s(fun(X43,bool),X44)))) ) | $spl15),
  inference(forward_demodulation,[],[f969,f468])).
fof(f969,plain,(
  ( ! [X45,X43,X44] : (s(bool,'const.bool.IN_2'(s(X43,sK2(X44,X45,X43)),s(fun(X43,bool),X45))) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(fun(X43,bool),X44) = s(fun(X43,bool),X45) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X43,sK2(X44,X45,X43)),s(fun(X43,bool),X44)))) ) | $spl15),
  inference(superposition,[],[f128,f496])).
fof(f11807,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1037,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f1037,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X47),s(fun(bool,bool),X48)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),X47)))) | p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),X48)))) | s(bool,'T_0') = s(bool,X46)) ) | $spl15),
  inference(forward_demodulation,[],[f970,f468])).
fof(f970,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X47),s(fun(bool,bool),X48)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),X47)))) | p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),X48)))) | s(bool,'T_0') = s(bool,X46)) ) | $spl15),
  inference(superposition,[],[f129,f496])).
fof(f11806,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1039,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f1039,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X53),s(fun(bool,bool),X54)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X52),s(fun(bool,bool),X54)))) | s(bool,'T_0') = s(bool,X52)) ) | $spl15),
  inference(forward_demodulation,[],[f972,f468])).
fof(f972,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X53),s(fun(bool,bool),X54)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X52),s(fun(bool,bool),X54)))) | s(bool,'T_0') = s(bool,X52)) ) | $spl15),
  inference(superposition,[],[f131,f496])).
fof(f11805,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1042,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f1042,plain,(
  ( ! [X66,X65] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X65))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X66))) | s(fun(bool,bool),X65) = s(fun(bool,bool),X66) | s(bool,'T_0') = s(bool,sK2(X65,X66,bool))) ) | $spl15),
  inference(forward_demodulation,[],[f976,f468])).
fof(f976,plain,(
  ( ! [X66,X65] : (s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))),s(fun(bool,bool),X65))) != s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))),s(fun(bool,bool),X66))) | s(fun(bool,bool),X65) = s(fun(bool,bool),X66) | s(bool,'T_0') = s(bool,sK2(X65,X66,bool))) ) | $spl15),
  inference(superposition,[],[f128,f496])).
fof(f11804,plain,(
  $false | ($spl8 | $spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1270,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f1270,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ($spl8 | $spl15)),
  inference(forward_demodulation,[],[f1257,f468])).
fof(f1257,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl8 | $spl15)),
  inference(unit_resulting_resolution,[],[f1243,f496])).
fof(f1243,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl8),
  inference(backward_demodulation,[],[f232,f226])).
fof(f11803,plain,(
  $false | ($spl23 | $spl34 | $spl50)),
  inference(global_subsumption,[],[f1287,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f11802,plain,(
  $false | ($spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1293,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f11801,plain,(
  $false | ($spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1364,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f1364,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl43),
  inference(trivial_inequality_removal,[],[f1359])).
fof(f1359,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl43),
  inference(superposition,[],[f1293,f140])).
fof(f11800,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1365,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f1365,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl43)),
  inference(forward_demodulation,[],[f1346,f468])).
fof(f1346,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl15 | $spl43)),
  inference(unit_resulting_resolution,[],[f1293,f496])).
fof(f11799,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380,f1372])).
fof(f11798,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647,f1380])).
fof(f11797,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1495,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647])).
fof(f1495,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f1472,f142])).
fof(f1472,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f176,f1401])).
fof(f1401,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f1365,f90])).
fof(f11796,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1562,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647])).
fof(f1562,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X26))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X27),s(fun(bool,bool),X28))))) | s(bool,'T_0') = s(bool,X26)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f1495])).
fof(f11795,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1563,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647])).
fof(f1563,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X30),s(fun(bool,bool),X31)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X29))),s(fun(bool,bool),X30)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X29))),s(fun(bool,bool),X31)))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f1495])).
fof(f11794,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1564,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647])).
fof(f1564,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X33),s(fun(bool,bool),X34)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X32))),s(fun(bool,bool),X34)))) | s(bool,'T_0') = s(bool,X32)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f1495])).
fof(f11793,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1565,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647])).
fof(f1565,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X36),s(fun(bool,bool),X37)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X35))),s(fun(bool,bool),X36)))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f1495])).
fof(f11792,plain,(
  $false | ($spl23 | $spl49 | $spl50)),
  inference(global_subsumption,[],[f1589,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647])).
fof(f1589,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl49),
  inference(superposition,[],[f1515,f140])).
fof(f11791,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746,f1647])).
fof(f11790,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838,f1746])).
fof(f11789,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194,f1838])).
fof(f11788,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1999,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194])).
fof(f1999,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f1998])).
fof(f1998,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X42))))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42)) ) | ($spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f1925])).
fof(f1925,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X42))))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f1495,f1483])).
fof(f1483,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f1457])).
fof(f1457,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f1401,f140])).
fof(f11787,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2186,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194])).
fof(f2186,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X0))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))))) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f2185])).
fof(f2185,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X23))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X24),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X23)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2162])).
fof(f2162,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X23))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X24),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f444,f1483])).
fof(f11786,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511,f2194])).
fof(f11785,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2364,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2364,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X44)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X43))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X44),s(fun(bool,bool),X45))))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f1999])).
fof(f11784,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2365,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2365,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X47),s(fun(bool,bool),X48)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X46))))),s(fun(bool,bool),X47)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X46))))),s(fun(bool,bool),X48)))) | s(bool,'T_0') = s(bool,X46)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f1999])).
fof(f11783,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2366,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2366,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X50),s(fun(bool,bool),X51)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X49))))),s(fun(bool,bool),X51)))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f1999])).
fof(f11782,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2367,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2367,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X53),s(fun(bool,bool),X54)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X52))))),s(fun(bool,bool),X53)))) | s(bool,'T_0') = s(bool,X52)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f1999])).
fof(f11781,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2375,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2375,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f2374])).
fof(f2374,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2301])).
fof(f2301,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f1999,f1483])).
fof(f11780,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3038,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3038,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X56))))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X57),s(fun(bool,bool),X58))))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f2375])).
fof(f11779,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3039,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3039,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X60),s(fun(bool,bool),X61)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X59))))))),s(fun(bool,bool),X60)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X59))))))),s(fun(bool,bool),X61)))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f2375])).
fof(f11778,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3040,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3040,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X63),s(fun(bool,bool),X64)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X62))))))),s(fun(bool,bool),X64)))) | s(bool,'T_0') = s(bool,X62)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f2375])).
fof(f11777,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3041,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3041,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X66),s(fun(bool,bool),X67)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X65))))))),s(fun(bool,bool),X66)))) | s(bool,'T_0') = s(bool,X65)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f2375])).
fof(f11776,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3052,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3052,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f3051])).
fof(f3051,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2941])).
fof(f2941,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2375,f1483])).
fof(f11775,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4507,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4507,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X77)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X76))))))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X77),s(fun(bool,bool),X78))))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f3052])).
fof(f11774,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4508,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4508,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X80),s(fun(bool,bool),X81)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X79))))))))),s(fun(bool,bool),X80)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X79))))))))),s(fun(bool,bool),X81)))) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f3052])).
fof(f11773,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4509,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4509,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X83),s(fun(bool,bool),X84)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X82))))))))),s(fun(bool,bool),X84)))) | s(bool,'T_0') = s(bool,X82)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f3052])).
fof(f11772,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4510,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4510,plain,(
  ( ! [X87,X85,X86] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X86),s(fun(bool,bool),X87)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X85))))))))),s(fun(bool,bool),X86)))) | s(bool,'T_0') = s(bool,X85)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f3052])).
fof(f11771,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4524,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4524,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f4523])).
fof(f4523,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f4359])).
fof(f4359,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3052,f1483])).
fof(f11770,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6224,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6224,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))))) ) | ($spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f6223])).
fof(f6223,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X83))))))))))))) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X83)) ) | ($spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f6024])).
fof(f6024,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X83))))))))))))) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X83)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3052,f2024])).
fof(f2024,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,f(s(bool,X1))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f2023])).
fof(f2023,plain,(
  ( ! [X17,X18,X16] : (s(bool,X18) = s(bool,f(s(bool,f(s(bool,X17))))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f1882])).
fof(f1882,plain,(
  ( ! [X17,X18,X16] : (s(bool,X18) = s(bool,f(s(bool,f(s(bool,X17))))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f1483,f1483])).
fof(f11769,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6260,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6260,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X0))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))))) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f6259])).
fof(f6259,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X35))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X36),s(fun(bool,bool),X37)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X36)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X37)))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X35)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f6004])).
fof(f6004,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X35))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X36),s(fun(bool,bool),X37)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X36)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X37)))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X35)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f444,f2024])).
fof(f11768,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl17 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6515,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6515,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl43)),
  inference(backward_demodulation,[],[f6514,f6508])).
fof(f6508,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl43)),
  inference(backward_demodulation,[],[f6486,f6489])).
fof(f6489,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl43)),
  inference(unit_resulting_resolution,[],[f1063,f4562])).
fof(f4562,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X15))))))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f4362,f151])).
fof(f4362,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X15))))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3052,f90])).
fof(f6486,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl43)),
  inference(unit_resulting_resolution,[],[f1063,f3078])).
fof(f3078,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f2944,f151])).
fof(f2944,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X14)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2375,f90])).
fof(f6514,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl43)),
  inference(forward_demodulation,[],[f6513,f6508])).
fof(f6513,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl43)),
  inference(forward_demodulation,[],[f6512,f6508])).
fof(f6512,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))))) | ($spl3 | $spl6 | $spl15 | $spl17 | $spl43)),
  inference(forward_demodulation,[],[f6485,f6508])).
fof(f6485,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))))))) | ($spl6 | $spl15 | $spl17 | $spl43)),
  inference(unit_resulting_resolution,[],[f1063,f3052])).
fof(f11767,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl30 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f7577,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f7577,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | ($spl3 | $spl6 | $spl15 | $spl30 | $spl43)),
  inference(backward_demodulation,[],[f7576,f7566])).
fof(f7566,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))) | ($spl3 | $spl6 | $spl15 | $spl30 | $spl43)),
  inference(backward_demodulation,[],[f7541,f7544])).
fof(f7544,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))))))) | ($spl3 | $spl6 | $spl15 | $spl30 | $spl43)),
  inference(unit_resulting_resolution,[],[f7154,f4562])).
fof(f7154,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | $spl30),
  inference(backward_demodulation,[],[f1175,f410])).
fof(f1175,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl30),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f7541,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))))) | ($spl3 | $spl6 | $spl15 | $spl30 | $spl43)),
  inference(unit_resulting_resolution,[],[f7154,f3078])).
fof(f7576,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))) | ($spl3 | $spl6 | $spl15 | $spl30 | $spl43)),
  inference(forward_demodulation,[],[f7575,f7566])).
fof(f7575,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))) | ($spl3 | $spl6 | $spl15 | $spl30 | $spl43)),
  inference(forward_demodulation,[],[f7574,f7566])).
fof(f7574,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))))) | ($spl3 | $spl6 | $spl15 | $spl30 | $spl43)),
  inference(forward_demodulation,[],[f7540,f7566])).
fof(f7540,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))))))))))) | ($spl6 | $spl15 | $spl30 | $spl43)),
  inference(unit_resulting_resolution,[],[f7154,f3052])).
fof(f11766,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9214,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9214,plain,(
  ( ! [X134,X136,X135] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X135)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X134))))))))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X135),s(fun(bool,bool),X136))))) | s(bool,'T_0') = s(bool,X134)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f4524])).
fof(f11765,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9215,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9215,plain,(
  ( ! [X138,X139,X137] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X138),s(fun(bool,bool),X139)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X137))))))))))),s(fun(bool,bool),X138)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X137))))))))))),s(fun(bool,bool),X139)))) | s(bool,'T_0') = s(bool,X137)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f4524])).
fof(f11764,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9216,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9216,plain,(
  ( ! [X142,X140,X141] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X141),s(fun(bool,bool),X142)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X140))))))))))),s(fun(bool,bool),X142)))) | s(bool,'T_0') = s(bool,X140)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f4524])).
fof(f11763,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9217,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9217,plain,(
  ( ! [X144,X145,X143] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X144),s(fun(bool,bool),X145)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X143))))))))))),s(fun(bool,bool),X144)))) | s(bool,'T_0') = s(bool,X143)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f4524])).
fof(f11762,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9241,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9241,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))))))) ) | ($spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f9240])).
fof(f9240,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f8951])).
fof(f8951,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f4524,f2024])).
fof(f11761,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1475,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1475,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X23)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X22))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24))))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f283,f1401])).
fof(f11760,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1476,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1476,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X26),s(fun(bool,bool),X27)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X25))),s(fun(bool,bool),X26)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X25))),s(fun(bool,bool),X27)))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f129,f1401])).
fof(f11759,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1477,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1477,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X28))),s(fun(bool,bool),X30)))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f131,f1401])).
fof(f11758,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1478,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1478,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,X31),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X32),s(fun(bool,bool),X33)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X31))),s(fun(bool,bool),X32)))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f130,f1401])).
fof(f11757,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1494,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1494,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl43)),
  inference(subsumption_resolution,[],[f1464,f142])).
fof(f1464,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f132,f1401])).
fof(f11756,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1723,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1723,plain,(
  ( ! [X24,X25] : (s(bool,f(s(bool,X24))) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f1711,f89])).
fof(f1711,plain,(
  ( ! [X24,X25] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X24))) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f1494,f1495])).
fof(f11755,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1945,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1945,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,f(s(bool,X107))) != s(bool,'const.bool.IN_2'(s(X104,sK2(X105,X106,X104)),s(fun(X104,bool),X105))) | s(fun(X104,bool),X105) = s(fun(X104,bool),X106) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X104,sK2(X105,X106,X104)),s(fun(X104,bool),X106)))) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f128,f1483])).
fof(f11754,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1946,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1946,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,f(s(bool,X111))) != s(bool,'const.bool.IN_2'(s(X108,sK2(X109,X110,X108)),s(fun(X108,bool),X110))) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X108,sK2(X109,X110,X108)),s(fun(X108,bool),X109)))) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f128,f1483])).
fof(f11753,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1950,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1950,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X125))),s(fun(bool,bool),X126)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X124),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X126),s(fun(bool,bool),X127))))) | s(bool,'T_0') = s(bool,X125) | s(bool,'T_0') = s(bool,X124)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f283,f1483])).
fof(f11752,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1951,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1951,plain,(
  ( ! [X130,X128,X131,X129] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X129))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X130),s(fun(bool,bool),X131)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X128),s(fun(bool,bool),X130)))) | p(s(bool,'const.bool.IN_2'(s(bool,X128),s(fun(bool,bool),X131)))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,X128)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f129,f1483])).
fof(f11751,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1952,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1952,plain,(
  ( ! [X134,X132,X135,X133] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X133))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X134),s(fun(bool,bool),X135)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X132),s(fun(bool,bool),X135)))) | s(bool,'T_0') = s(bool,X133) | s(bool,'T_0') = s(bool,X132)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f131,f1483])).
fof(f11750,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1953,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1953,plain,(
  ( ! [X138,X136,X139,X137] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X137))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X138),s(fun(bool,bool),X139)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X136),s(fun(bool,bool),X138)))) | s(bool,'T_0') = s(bool,X137) | s(bool,'T_0') = s(bool,X136)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f130,f1483])).
fof(f11749,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1955,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1955,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,f(s(bool,X150))) != s(bool,f(s(X147,sK3(X148,X149,bool,X147)))) | s(fun(X147,bool),X148) = s(fun(X147,bool),X149) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,g(s(X147,sK3(X148,X149,bool,X147))))) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f136,f1483])).
fof(f11748,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1957,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1957,plain,(
  ( ! [X156,X154,X155] : (s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X156))),s(fun(bool,bool),X154))) != s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X156))),s(fun(bool,bool),X155))) | s(fun(bool,bool),X154) = s(fun(bool,bool),X155) | s(bool,'T_0') = s(bool,X156) | s(bool,'T_0') = s(bool,sK2(X154,X155,bool))) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f128,f1483])).
fof(f11747,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1958,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1958,plain,(
  ( ! [X158,X159,X157,X160] : (s(X159,f(s(bool,f(s(bool,X160))))) != s(X159,g(s(bool,f(s(bool,X160))))) | s(fun(bool,X159),X157) = s(fun(bool,X159),X158) | s(bool,'T_0') = s(bool,X160) | s(bool,'T_0') = s(bool,sK3(X157,X158,X159,bool))) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f136,f1483])).
fof(f11746,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1982,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1982,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'const.bool.IN_2'(s(bool,X58),s(fun(bool,bool),X59)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X57))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X59),s(fun(bool,bool),X60))))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X58)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f283,f1483])).
fof(f11745,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1983,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1983,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,X62),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X63),s(fun(bool,bool),X64)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X61))),s(fun(bool,bool),X63)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X61))),s(fun(bool,bool),X64)))) | s(bool,'T_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X62)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f129,f1483])).
fof(f11744,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1984,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1984,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X67),s(fun(bool,bool),X68)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X65))),s(fun(bool,bool),X68)))) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X66)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f131,f1483])).
fof(f11743,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f1985,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f1985,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,'const.bool.IN_2'(s(bool,X70),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X71),s(fun(bool,bool),X72)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X69))),s(fun(bool,bool),X71)))) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X70)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f130,f1483])).
fof(f11742,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2001,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2001,plain,(
  ( ! [X0] : (s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f2000])).
fof(f2000,plain,(
  ( ! [X37,X36] : (s(bool,f(s(bool,X37))) = s(bool,f(s(bool,f(s(bool,X37))))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X37)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f1923])).
fof(f1923,plain,(
  ( ! [X37,X36] : (s(bool,f(s(bool,X37))) = s(bool,f(s(bool,f(s(bool,X37))))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f1401,f1483])).
fof(f11741,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2024,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f11740,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2071,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2071,plain,(
  ( ! [X1] : (~p(s(bool,f(s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2070])).
fof(f2070,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | ~p(s(bool,f(s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f2069])).
fof(f2069,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,f(s(bool,X9)))) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | ($spl15 | $spl43)),
  inference(subsumption_resolution,[],[f1913,f177])).
fof(f1913,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,f(s(bool,X9)))) | ~p(s(bool,X10)) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f135,f1483])).
fof(f11739,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2130,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2130,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f2129])).
fof(f2129,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,X11)))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2097])).
fof(f2097,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,X11)))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2071,f1483])).
fof(f11738,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2250,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2250,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f2249])).
fof(f2249,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2207])).
fof(f2207,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2130,f1483])).
fof(f11737,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2392,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2392,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,f(s(bool,f(s(bool,X14))))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f2304,f151])).
fof(f2304,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,f(s(bool,f(s(bool,X14))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X14)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f1999,f90])).
fof(f11736,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2405,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2405,plain,(
  ( ! [X8,X9] : (p(s(bool,X9)) | s(bool,X9) = s(bool,f(s(bool,f(s(bool,X8))))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f2344,f89])).
fof(f2344,plain,(
  ( ! [X8,X9] : (p(s(bool,'F_0')) | p(s(bool,X9)) | s(bool,X9) = s(bool,f(s(bool,f(s(bool,X8))))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f132,f1999])).
fof(f11735,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2406,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2406,plain,(
  ( ! [X39,X38] : (s(bool,f(s(bool,X39))) = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f2360,f89])).
fof(f2360,plain,(
  ( ! [X39,X38] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X39))) = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f1494,f1999])).
fof(f11734,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2568,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2568,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),X47)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X46))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X47),s(fun(bool,bool),X48))))) | s(bool,'T_0') = s(bool,X46)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f2392])).
fof(f11733,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2569,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2569,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X49),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X50),s(fun(bool,bool),X51)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X49))))),s(fun(bool,bool),X50)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X49))))),s(fun(bool,bool),X51)))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f2392])).
fof(f11732,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2570,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2570,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,X52),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X53),s(fun(bool,bool),X54)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X52))))),s(fun(bool,bool),X54)))) | s(bool,'T_0') = s(bool,X52)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f2392])).
fof(f11731,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2571,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2571,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X56),s(fun(bool,bool),X57)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X55))))),s(fun(bool,bool),X56)))) | s(bool,'T_0') = s(bool,X55)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f2392])).
fof(f11730,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2584,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2584,plain,(
  ( ! [X0] : (s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f2583])).
fof(f2583,plain,(
  ( ! [X10,X11] : (s(bool,f(s(bool,X11))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2496])).
fof(f2496,plain,(
  ( ! [X10,X11] : (s(bool,f(s(bool,X11))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2392,f1483])).
fof(f11729,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2688,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2688,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f2687])).
fof(f2687,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2625])).
fof(f2625,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2250,f1483])).
fof(f11728,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2757,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2757,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,X8) = s(bool,f(s(bool,f(s(bool,f(s(bool,X7))))))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl15 | $spl43)),
  inference(subsumption_resolution,[],[f2675,f87])).
fof(f2675,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X7) | s(bool,X8) = s(bool,f(s(bool,f(s(bool,f(s(bool,X7))))))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2250,f140])).
fof(f11727,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2883,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2883,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X53))),s(fun(bool,bool),X54)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X53))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X54),s(fun(bool,bool),X55))))) | s(bool,'T_0') = s(bool,X53)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f283,f2001])).
fof(f11726,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2884,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2884,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X56))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X57),s(fun(bool,bool),X58)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X56))))),s(fun(bool,bool),X57)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X56))))),s(fun(bool,bool),X58)))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f129,f2001])).
fof(f11725,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2885,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2885,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X59))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X60),s(fun(bool,bool),X61)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X59))))),s(fun(bool,bool),X61)))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f131,f2001])).
fof(f11724,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2886,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2886,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X62))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X63),s(fun(bool,bool),X64)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X62))))),s(fun(bool,bool),X63)))) | s(bool,'T_0') = s(bool,X62)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f130,f2001])).
fof(f11723,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f2899,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f2899,plain,(
  ( ! [X0] : (s(bool,f(s(bool,f(s(bool,X0))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f2898])).
fof(f2898,plain,(
  ( ! [X10,X11] : (s(bool,f(s(bool,f(s(bool,X11))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f2802])).
fof(f2802,plain,(
  ( ! [X10,X11] : (s(bool,f(s(bool,f(s(bool,X11))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2001,f1483])).
fof(f11722,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3078,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f11721,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3109,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3109,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,X18) = s(bool,f(s(bool,f(s(bool,f(s(bool,X17))))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f3014,f89])).
fof(f3014,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | p(s(bool,X18)) | s(bool,X18) = s(bool,f(s(bool,f(s(bool,f(s(bool,X17))))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f132,f2375])).
fof(f11720,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3110,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3110,plain,(
  ( ! [X47,X48] : (s(bool,f(s(bool,X48))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X47))))))) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X47)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f3030,f89])).
fof(f3030,plain,(
  ( ! [X47,X48] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X48))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X47))))))) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X47)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f1494,f2375])).
fof(f11719,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3345,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3345,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,X60),s(fun(bool,bool),X61)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X60))))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X61),s(fun(bool,bool),X62))))) | s(bool,'T_0') = s(bool,X60)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f3078])).
fof(f11718,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3346,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3346,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,X63),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X64),s(fun(bool,bool),X65)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X63))))))),s(fun(bool,bool),X64)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X63))))))),s(fun(bool,bool),X65)))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f3078])).
fof(f11717,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3347,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3347,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X67),s(fun(bool,bool),X68)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X66))))))),s(fun(bool,bool),X68)))) | s(bool,'T_0') = s(bool,X66)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f3078])).
fof(f11716,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3348,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3348,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,'const.bool.IN_2'(s(bool,X69),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X70),s(fun(bool,bool),X71)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X69))))))),s(fun(bool,bool),X70)))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f3078])).
fof(f11715,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3364,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3364,plain,(
  ( ! [X0] : (s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f3363])).
fof(f3363,plain,(
  ( ! [X10,X11] : (s(bool,f(s(bool,X11))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f3236])).
fof(f3236,plain,(
  ( ! [X10,X11] : (s(bool,f(s(bool,X11))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3078,f1483])).
fof(f11714,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3524,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3524,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f3523])).
fof(f3523,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f3431])).
fof(f3431,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2688,f1483])).
fof(f11713,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3635,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3635,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,X10) | s(bool,X11) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X10))))))))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43)),
  inference(subsumption_resolution,[],[f3508,f87])).
fof(f3508,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X10) | s(bool,X11) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X10))))))))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2688,f140])).
fof(f11712,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3994,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3994,plain,(
  ( ! [X39,X38] : (s(bool,f(s(bool,f(s(bool,X38))))) = s(bool,f(s(bool,f(s(bool,X39))))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f3968,f2071])).
fof(f3968,plain,(
  ( ! [X39,X38] : (p(s(bool,f(s(bool,X38)))) | s(bool,f(s(bool,f(s(bool,X38))))) = s(bool,f(s(bool,f(s(bool,X39))))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2405,f2001])).
fof(f11711,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f3995,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f3995,plain,(
  ( ! [X45,X44] : (s(bool,f(s(bool,f(s(bool,X45))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X44))))))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X44)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f3971,f142])).
fof(f3971,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | s(bool,f(s(bool,f(s(bool,X45))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X44))))))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X44)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2405,f3078])).
fof(f11710,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4146,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4146,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X72))),s(fun(bool,bool),X73)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X72))))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X73),s(fun(bool,bool),X74))))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f2584])).
fof(f11709,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4147,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4147,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X75))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X76),s(fun(bool,bool),X77)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X75))))))),s(fun(bool,bool),X76)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X75))))))),s(fun(bool,bool),X77)))) | s(bool,'T_0') = s(bool,X75)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f2584])).
fof(f11708,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4148,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4148,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X78))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X79),s(fun(bool,bool),X80)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X78))))))),s(fun(bool,bool),X80)))) | s(bool,'T_0') = s(bool,X78)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f2584])).
fof(f11707,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4149,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4149,plain,(
  ( ! [X83,X81,X82] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X81))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X82),s(fun(bool,bool),X83)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X81))))))),s(fun(bool,bool),X82)))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f2584])).
fof(f11706,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4164,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4164,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,X0))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f4163])).
fof(f4163,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,f(s(bool,X12))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f4021])).
fof(f4021,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,f(s(bool,X12))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2584,f1483])).
fof(f11705,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4562,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f11704,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4635,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4635,plain,(
  ( ! [X30,X31] : (p(s(bool,X31)) | s(bool,X31) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X30))))))))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f4477,f89])).
fof(f4477,plain,(
  ( ! [X30,X31] : (p(s(bool,'F_0')) | p(s(bool,X31)) | s(bool,X31) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X30))))))))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f132,f3052])).
fof(f11703,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4636,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4636,plain,(
  ( ! [X61,X62] : (s(bool,f(s(bool,X62))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X61))))))))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f4494,f89])).
fof(f4494,plain,(
  ( ! [X61,X62] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X62))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X61))))))))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f1494,f3052])).
fof(f11702,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4648,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4648,plain,(
  ( ! [X72,X71] : (s(bool,f(s(bool,f(s(bool,X72))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X71))))))))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X71)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f4503,f89])).
fof(f4503,plain,(
  ( ! [X72,X71] : (p(s(bool,'F_0')) | s(bool,f(s(bool,f(s(bool,X72))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X71))))))))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X71)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2405,f3052])).
fof(f11701,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4847,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4847,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,X81),s(fun(bool,bool),X82)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X81))))))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X82),s(fun(bool,bool),X83))))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f4562])).
fof(f11700,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4848,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4848,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,'const.bool.IN_2'(s(bool,X84),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X85),s(fun(bool,bool),X86)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X84))))))))),s(fun(bool,bool),X85)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X84))))))))),s(fun(bool,bool),X86)))) | s(bool,'T_0') = s(bool,X84)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f4562])).
fof(f11699,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4849,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4849,plain,(
  ( ! [X88,X87,X89] : (p(s(bool,'const.bool.IN_2'(s(bool,X87),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X88),s(fun(bool,bool),X89)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X87))))))))),s(fun(bool,bool),X89)))) | s(bool,'T_0') = s(bool,X87)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f4562])).
fof(f11698,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4850,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4850,plain,(
  ( ! [X92,X90,X91] : (p(s(bool,'const.bool.IN_2'(s(bool,X90),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X91),s(fun(bool,bool),X92)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X90))))))))),s(fun(bool,bool),X91)))) | s(bool,'T_0') = s(bool,X90)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f4562])).
fof(f11697,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f4869,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f4869,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f4868])).
fof(f4868,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,X12))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f4684])).
fof(f4684,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,X12))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f4562,f1483])).
fof(f11696,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f5113,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f5113,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))))))))) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f5112])).
fof(f5112,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12)))))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f4978])).
fof(f4978,plain,(
  ( ! [X12,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12)))))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f3524,f1483])).
fof(f11695,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f5292,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f5292,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X13))))))))))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43)),
  inference(subsumption_resolution,[],[f5093,f87])).
fof(f5093,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X13))))))))))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f3524,f140])).
fof(f11694,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f5581,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f5581,plain,(
  ( ! [X182,X180,X181,X179] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X180))),s(fun(bool,bool),X181)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X179))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X181),s(fun(bool,bool),X182))))) | s(bool,'T_0') = s(bool,X180) | s(bool,'T_0') = s(bool,X179)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f1723])).
fof(f11693,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f5582,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f5582,plain,(
  ( ! [X185,X186,X184,X183] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X184))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X185),s(fun(bool,bool),X186)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X183))),s(fun(bool,bool),X185)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X183))),s(fun(bool,bool),X186)))) | s(bool,'T_0') = s(bool,X184) | s(bool,'T_0') = s(bool,X183)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f1723])).
fof(f11692,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f5583,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f5583,plain,(
  ( ! [X189,X187,X190,X188] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X188))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X189),s(fun(bool,bool),X190)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X187))),s(fun(bool,bool),X190)))) | s(bool,'T_0') = s(bool,X188) | s(bool,'T_0') = s(bool,X187)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f1723])).
fof(f11691,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f5584,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f5584,plain,(
  ( ! [X191,X194,X192,X193] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X192))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X193),s(fun(bool,bool),X194)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X191))),s(fun(bool,bool),X193)))) | s(bool,'T_0') = s(bool,X192) | s(bool,'T_0') = s(bool,X191)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f1723])).
fof(f11690,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6086,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6086,plain,(
  ( ! [X243,X241,X244,X242] : (s(bool,f(s(bool,f(s(bool,X244))))) != s(bool,'const.bool.IN_2'(s(X241,sK2(X242,X243,X241)),s(fun(X241,bool),X242))) | s(fun(X241,bool),X242) = s(fun(X241,bool),X243) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X241,sK2(X242,X243,X241)),s(fun(X241,bool),X243))) | s(bool,'T_0') = s(bool,X244)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f128,f2024])).
fof(f11689,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6087,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6087,plain,(
  ( ! [X246,X248,X247,X245] : (s(bool,f(s(bool,f(s(bool,X248))))) != s(bool,'const.bool.IN_2'(s(X245,sK2(X246,X247,X245)),s(fun(X245,bool),X247))) | s(fun(X245,bool),X246) = s(fun(X245,bool),X247) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X245,sK2(X246,X247,X245)),s(fun(X245,bool),X246))) | s(bool,'T_0') = s(bool,X248)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f128,f2024])).
fof(f11688,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6094,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6094,plain,(
  ( ! [X269,X271,X268,X270] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X269))))),s(fun(bool,bool),X270)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X268),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X270),s(fun(bool,bool),X271))))) | s(bool,'T_0') = s(bool,X268) | s(bool,'T_0') = s(bool,X269)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f283,f2024])).
fof(f11687,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6095,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6095,plain,(
  ( ! [X273,X275,X272,X274] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X273))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X274),s(fun(bool,bool),X275)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X272),s(fun(bool,bool),X274)))) | p(s(bool,'const.bool.IN_2'(s(bool,X272),s(fun(bool,bool),X275)))) | s(bool,'T_0') = s(bool,X272) | s(bool,'T_0') = s(bool,X273)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f129,f2024])).
fof(f11686,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6096,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6096,plain,(
  ( ! [X277,X279,X276,X278] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X277))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X278),s(fun(bool,bool),X279)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X276),s(fun(bool,bool),X279)))) | s(bool,'T_0') = s(bool,X276) | s(bool,'T_0') = s(bool,X277)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f131,f2024])).
fof(f11685,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6097,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6097,plain,(
  ( ! [X280,X282,X281,X283] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X281))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X282),s(fun(bool,bool),X283)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X280),s(fun(bool,bool),X282)))) | s(bool,'T_0') = s(bool,X280) | s(bool,'T_0') = s(bool,X281)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f130,f2024])).
fof(f11684,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6099,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6099,plain,(
  ( ! [X292,X294,X293,X291] : (s(bool,f(s(bool,f(s(bool,X294))))) != s(bool,f(s(X291,sK3(X292,X293,bool,X291)))) | s(fun(X291,bool),X292) = s(fun(X291,bool),X293) | s(bool,'T_0') = s(bool,g(s(X291,sK3(X292,X293,bool,X291)))) | s(bool,'T_0') = s(bool,X294)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f136,f2024])).
fof(f11683,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6100,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6100,plain,(
  ( ! [X296,X298,X297] : (s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X298))))),s(fun(bool,bool),X296))) != s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X298))))),s(fun(bool,bool),X297))) | s(fun(bool,bool),X296) = s(fun(bool,bool),X297) | s(bool,'T_0') = s(bool,sK2(X296,X297,bool)) | s(bool,'T_0') = s(bool,X298)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f128,f2024])).
fof(f11682,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6101,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6101,plain,(
  ( ! [X300,X302,X301,X299] : (s(X301,f(s(bool,f(s(bool,f(s(bool,X302))))))) != s(X301,g(s(bool,f(s(bool,f(s(bool,X302))))))) | s(fun(bool,X301),X299) = s(fun(bool,X301),X300) | s(bool,'T_0') = s(bool,sK3(X299,X300,X301,bool)) | s(bool,'T_0') = s(bool,X302)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f136,f2024])).
fof(f11681,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6183,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6183,plain,(
  ( ! [X182,X184,X183,X181] : (~p(s(bool,'const.bool.IN_2'(s(bool,X182),s(fun(bool,bool),X183)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X181))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X183),s(fun(bool,bool),X184))))) | s(bool,'T_0') = s(bool,X182) | s(bool,'T_0') = s(bool,X181)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f283,f2024])).
fof(f11680,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6184,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6184,plain,(
  ( ! [X187,X185,X188,X186] : (~p(s(bool,'const.bool.IN_2'(s(bool,X186),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X187),s(fun(bool,bool),X188)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X185))))),s(fun(bool,bool),X187)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X185))))),s(fun(bool,bool),X188)))) | s(bool,'T_0') = s(bool,X186) | s(bool,'T_0') = s(bool,X185)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f129,f2024])).
fof(f11679,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6185,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6185,plain,(
  ( ! [X191,X189,X192,X190] : (p(s(bool,'const.bool.IN_2'(s(bool,X190),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X191),s(fun(bool,bool),X192)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X189))))),s(fun(bool,bool),X192)))) | s(bool,'T_0') = s(bool,X190) | s(bool,'T_0') = s(bool,X189)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f131,f2024])).
fof(f11678,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6186,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6186,plain,(
  ( ! [X196,X194,X195,X193] : (p(s(bool,'const.bool.IN_2'(s(bool,X194),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X195),s(fun(bool,bool),X196)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X193))))),s(fun(bool,bool),X195)))) | s(bool,'T_0') = s(bool,X194) | s(bool,'T_0') = s(bool,X193)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f130,f2024])).
fof(f11677,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6216,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6216,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,X0))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))))) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f6215])).
fof(f6215,plain,(
  ( ! [X90,X91] : (s(bool,f(s(bool,f(s(bool,X91))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X91))))))))))))) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X91)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f6028])).
fof(f6028,plain,(
  ( ! [X90,X91] : (s(bool,f(s(bool,f(s(bool,X91))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X91))))))))))))) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X91)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f4562,f2024])).
fof(f11676,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6220,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6220,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))))))))))) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f6219])).
fof(f6219,plain,(
  ( ! [X87,X86] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X87)))))))))))))))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f6026])).
fof(f6026,plain,(
  ( ! [X87,X86] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X87)))))))))))))))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f3524,f2024])).
fof(f11675,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6222,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6222,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,X0))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f6221])).
fof(f6221,plain,(
  ( ! [X85,X84] : (s(bool,f(s(bool,f(s(bool,X85))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X85))))))))))) | s(bool,'T_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X85)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f6025])).
fof(f6025,plain,(
  ( ! [X85,X84] : (s(bool,f(s(bool,f(s(bool,X85))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X85))))))))))) | s(bool,'T_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X85)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3078,f2024])).
fof(f11674,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6228,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6228,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f6227])).
fof(f6227,plain,(
  ( ! [X78,X79] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X79))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X79))))))))))) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X79)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f6022])).
fof(f6022,plain,(
  ( ! [X78,X79] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X79))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X79))))))))))) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X79)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2584,f2024])).
fof(f11673,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6240,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6240,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f6239])).
fof(f6239,plain,(
  ( ! [X64,X63] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X64))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X64))))))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f6015])).
fof(f6015,plain,(
  ( ! [X64,X63] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X64))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X64))))))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2001,f2024])).
fof(f11672,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6401,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6401,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43 | $spl50)),
  inference(forward_demodulation,[],[f6218,f5345])).
fof(f6218,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,X0))))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl43 | $spl50)),
  inference(condensation,[],[f6217])).
fof(f6217,plain,(
  ( ! [X88,X89] : (s(bool,f(s(bool,f(s(bool,X89))))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X89)) ) | ($spl15 | $spl43 | $spl50)),
  inference(duplicate_literal_removal,[],[f6027])).
fof(f6027,plain,(
  ( ! [X88,X89] : (s(bool,f(s(bool,f(s(bool,X89))))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X89)) ) | ($spl15 | $spl43 | $spl50)),
  inference(superposition,[],[f4344,f2024])).
fof(f11671,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6725,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6725,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43 | $spl50)),
  inference(condensation,[],[f6724])).
fof(f6724,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43 | $spl50)),
  inference(duplicate_literal_removal,[],[f6634])).
fof(f6634,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43 | $spl50)),
  inference(superposition,[],[f6401,f2024])).
fof(f11670,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6727,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6727,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl43 | $spl50)),
  inference(condensation,[],[f6726])).
fof(f6726,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl15 | $spl43 | $spl50)),
  inference(duplicate_literal_removal,[],[f6633])).
fof(f6633,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl43 | $spl50)),
  inference(superposition,[],[f6401,f1483])).
fof(f11669,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f6920,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f6920,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))) ) | ($spl15 | $spl43 | $spl50)),
  inference(condensation,[],[f6919])).
fof(f6919,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43 | $spl50)),
  inference(duplicate_literal_removal,[],[f6802])).
fof(f6802,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43 | $spl50)),
  inference(superposition,[],[f6727,f2024])).
fof(f11668,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f7313,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f7313,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))))) ) | ($spl15 | $spl43 | $spl50)),
  inference(condensation,[],[f7312])).
fof(f7312,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43 | $spl50)),
  inference(duplicate_literal_removal,[],[f7170])).
fof(f7170,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43 | $spl50)),
  inference(superposition,[],[f6725,f2024])).
fof(f11667,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f7744,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f7744,plain,(
  ( ! [X57,X56] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X56))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X57))))))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X56)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f7712,f2071])).
fof(f7712,plain,(
  ( ! [X57,X56] : (p(s(bool,f(s(bool,X56)))) | s(bool,f(s(bool,f(s(bool,f(s(bool,X56))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X57))))))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X56)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3109,f2584])).
fof(f11666,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f7745,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f7745,plain,(
  ( ! [X62,X63] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X63))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X62))))))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f7715,f142])).
fof(f7715,plain,(
  ( ! [X62,X63] : (p(s(bool,X62)) | s(bool,f(s(bool,f(s(bool,f(s(bool,X63))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X62))))))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3109,f4562])).
fof(f11665,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8234,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8234,plain,(
  ( ! [X125,X126,X124] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X124))))),s(fun(bool,bool),X125)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X124))))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X125),s(fun(bool,bool),X126))))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f283,f2899])).
fof(f11664,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8235,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8235,plain,(
  ( ! [X127,X128,X129] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X127))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X128),s(fun(bool,bool),X129)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X127))))))),s(fun(bool,bool),X128)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X127))))))),s(fun(bool,bool),X129)))) | s(bool,'T_0') = s(bool,X127)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f129,f2899])).
fof(f11663,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8236,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8236,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X130))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X131),s(fun(bool,bool),X132)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X130))))))),s(fun(bool,bool),X132)))) | s(bool,'T_0') = s(bool,X130)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f131,f2899])).
fof(f11662,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8237,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8237,plain,(
  ( ! [X134,X135,X133] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,X133))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X134),s(fun(bool,bool),X135)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,X133))))))),s(fun(bool,bool),X134)))) | s(bool,'T_0') = s(bool,X133)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f130,f2899])).
fof(f11661,plain,(
  $false | ($spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8266,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8266,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))) ) | ($spl15 | $spl43)),
  inference(condensation,[],[f8265])).
fof(f8265,plain,(
  ( ! [X14,X13] : (s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f8036])).
fof(f8036,plain,(
  ( ! [X14,X13] : (s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl43)),
  inference(superposition,[],[f2899,f2024])).
fof(f11660,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8767,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8767,plain,(
  ( ! [X130,X131,X129] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X129))),s(fun(bool,bool),X130)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X129))))))))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X130),s(fun(bool,bool),X131))))) | s(bool,'T_0') = s(bool,X129)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f283,f3364])).
fof(f11659,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8768,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8768,plain,(
  ( ! [X134,X132,X133] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X132))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X133),s(fun(bool,bool),X134)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X132))))))))),s(fun(bool,bool),X133)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X132))))))))),s(fun(bool,bool),X134)))) | s(bool,'T_0') = s(bool,X132)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f129,f3364])).
fof(f11658,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8769,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8769,plain,(
  ( ! [X136,X135,X137] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X135))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X136),s(fun(bool,bool),X137)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X135))))))))),s(fun(bool,bool),X137)))) | s(bool,'T_0') = s(bool,X135)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f131,f3364])).
fof(f11657,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8770,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8770,plain,(
  ( ! [X140,X138,X139] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X138))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X139),s(fun(bool,bool),X140)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X138))))))))),s(fun(bool,bool),X139)))) | s(bool,'T_0') = s(bool,X138)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f130,f3364])).
fof(f11656,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f8795,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f8795,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))))))) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(condensation,[],[f8794])).
fof(f8794,plain,(
  ( ! [X14,X13] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(duplicate_literal_removal,[],[f8538])).
fof(f8538,plain,(
  ( ! [X14,X13] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3364,f2024])).
fof(f11655,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9294,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9294,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X17))))))))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl3 | $spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f8954,f151])).
fof(f8954,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X17))))))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f4524,f90])).
fof(f11654,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9481,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9481,plain,(
  ( ! [X68,X67] : (p(s(bool,X68)) | s(bool,X68) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X67))))))))))) | s(bool,'T_0') = s(bool,X67)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f9168,f89])).
fof(f9168,plain,(
  ( ! [X68,X67] : (p(s(bool,'F_0')) | p(s(bool,X68)) | s(bool,X68) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X67))))))))))) | s(bool,'T_0') = s(bool,X67)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f132,f4524])).
fof(f11653,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9484,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9484,plain,(
  ( ! [X101,X100] : (s(bool,f(s(bool,X101))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X100))))))))))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X100)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f9187,f89])).
fof(f9187,plain,(
  ( ! [X101,X100] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X101))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X100))))))))))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X100)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f1494,f4524])).
fof(f11652,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9512,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9512,plain,(
  ( ! [X118,X119] : (s(bool,f(s(bool,f(s(bool,X119))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X118))))))))))) | s(bool,'T_0') = s(bool,X119) | s(bool,'T_0') = s(bool,X118)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f9200,f89])).
fof(f9200,plain,(
  ( ! [X118,X119] : (p(s(bool,'F_0')) | s(bool,f(s(bool,f(s(bool,X119))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X118))))))))))) | s(bool,'T_0') = s(bool,X119) | s(bool,'T_0') = s(bool,X118)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f2405,f4524])).
fof(f11651,plain,(
  $false | ($spl6 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f9531,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f9531,plain,(
  ( ! [X125,X126] : (s(bool,f(s(bool,f(s(bool,f(s(bool,X126))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X125))))))))))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125)) ) | ($spl6 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f9206,f89])).
fof(f9206,plain,(
  ( ! [X125,X126] : (p(s(bool,'F_0')) | s(bool,f(s(bool,f(s(bool,f(s(bool,X126))))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X125))))))))))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125)) ) | ($spl6 | $spl15 | $spl43)),
  inference(superposition,[],[f3109,f4524])).
fof(f11650,plain,(
  $false | ($spl8 | $spl15 | $spl23 | $spl43 | $spl50)),
  inference(global_subsumption,[],[f10694,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f10694,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) != s(bool,f(s(bool,f(s(bool,X33))))) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X33)) ) | ($spl8 | $spl15 | $spl43)),
  inference(subsumption_resolution,[],[f10676,f10226])).
fof(f10226,plain,(
  ( ! [X428,X427] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X427,X428),s(fun(X427,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(subsumption_resolution,[],[f9775,f1243])).
fof(f9775,plain,(
  ( ! [X428,X427] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X427,X428),s(fun(X427,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f1243,f225])).
fof(f10676,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) != s(bool,f(s(bool,f(s(bool,X33))))) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X32),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl8 | $spl15 | $spl43)),
  inference(superposition,[],[f10376,f2024])).
fof(f10376,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,'const.bool.IN_2'(s(bool,X18),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl8),
  inference(superposition,[],[f10226,f140])).
fof(f11649,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692,f1511])).
fof(f11648,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611,f692])).
fof(f11647,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f613,f889,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128,f611])).
fof(f11646,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127,f128])).
fof(f11645,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129,f127])).
fof(f11644,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130,f129])).
fof(f11643,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131,f130])).
fof(f11642,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283,f131])).
fof(f11641,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890,f283])).
fof(f11640,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399,f890])).
fof(f11639,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1240,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f890,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399])).
fof(f1240,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),X31)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),'const.pred_set.UNION_2'(s(fun(X28,bool),X30),s(fun(X28,bool),X31)))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1239,f89])).
fof(f1239,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),X31)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),'const.pred_set.UNION_2'(s(fun(X28,bool),X30),s(fun(X28,bool),X31)))))) ) | $spl15),
  inference(forward_demodulation,[],[f965,f468])).
fof(f965,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | ~p(s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),X31)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),'const.pred_set.UNION_2'(s(fun(X28,bool),X30),s(fun(X28,bool),X31)))))) ) | $spl15),
  inference(superposition,[],[f131,f496])).
fof(f11638,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f890,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656,f1399])).
fof(f11637,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f889,f890,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658,f656])).
fof(f11636,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f656,f889,f890,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659,f658])).
fof(f11635,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f656,f658,f889,f890,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693,f659])).
fof(f11634,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f656,f658,f659,f889,f890,f685,f686,f692,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373,f693])).
fof(f11633,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f656,f658,f659,f889,f890,f685,f686,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606,f1373])).
fof(f11632,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f611,f613,f656,f658,f659,f889,f890,f685,f686,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608,f606])).
fof(f11631,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f611,f613,f656,f658,f659,f889,f890,f685,f686,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609,f608])).
fof(f11630,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f611,f613,f656,f658,f659,f889,f890,f685,f686,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612,f609])).
fof(f11629,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f613,f656,f658,f659,f889,f890,f685,f686,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687,f612])).
fof(f11628,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689,f687])).
fof(f11627,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690,f689])).
fof(f11626,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374,f690])).
fof(f11625,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378,f1374])).
fof(f11624,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86,f1378])).
fof(f11623,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010,f86])).
fof(f11622,plain,(
  $false | ($spl17 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1063,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010])).
fof(f11621,plain,(
  $false | ($spl23 | $spl28 | $spl50)),
  inference(global_subsumption,[],[f1174,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010])).
fof(f11620,plain,(
  $false | ($spl23 | $spl30 | $spl50)),
  inference(global_subsumption,[],[f7154,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010])).
fof(f11619,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8417,f8418,f10051,f9810,f9811,f5345,f4344,f8010])).
fof(f11618,plain,(
  $false | ($spl23 | $spl30 | $spl50)),
  inference(global_subsumption,[],[f7823,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f7823,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f7765,f87])).
fof(f7765,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0)))))) ) | $spl30),
  inference(superposition,[],[f283,f1175])).
fof(f11617,plain,(
  $false | ($spl8 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f232,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f11616,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f432,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f11615,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f438,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f11614,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f496,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f11613,plain,(
  $false | ($spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1048,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f1048,plain,(
  ( ! [X23] : (s(bool,X23) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X23)) ) | $spl15),
  inference(subsumption_resolution,[],[f992,f86])).
fof(f992,plain,(
  ( ! [X23] : (s(bool,X23) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X23)) ) | $spl15),
  inference(superposition,[],[f136,f496])).
fof(f11612,plain,(
  $false | ($spl23 | $spl30 | $spl50)),
  inference(global_subsumption,[],[f1175,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f11611,plain,(
  $false | ($spl23 | $spl32 | $spl50)),
  inference(global_subsumption,[],[f1237,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f11610,plain,(
  $false | ($spl8 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1243,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f11609,plain,(
  $false | ($spl8 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f1269,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f1269,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(trivial_inequality_removal,[],[f1264])).
fof(f1264,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f1243,f140])).
fof(f11608,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345,f4344])).
fof(f11607,plain,(
  $false | ($spl23 | $spl50)),
  inference(global_subsumption,[],[f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f4344,f410,f8010,f8417,f8418,f10051,f9810,f9811,f5345])).
fof(f11606,plain,(
  $false | ($spl8 | $spl15 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f10691,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f4344,f5345,f410,f8010,f8417,f8418,f10051,f9810,f9811])).
fof(f10691,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f10673,f10226])).
fof(f10673,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.EMPTY_0')))) ) | ($spl8 | $spl15)),
  inference(superposition,[],[f10376,f496])).
fof(f11605,plain,(
  $false | ($spl8 | $spl23 | $spl50)),
  inference(global_subsumption,[],[f10692,f86,f87,f90,f126,f135,f132,f136,f140,f142,f148,f164,f174,f177,f178,f128,f127,f131,f130,f129,f89,f163,f207,f225,f227,f228,f229,f230,f276,f277,f283,f353,f411,f416,f423,f444,f549,f604,f605,f606,f608,f609,f611,f612,f613,f656,f658,f659,f889,f890,f685,f686,f687,f689,f690,f692,f693,f694,f1056,f1057,f1163,f1164,f1166,f1165,f1069,f226,f1281,f1282,f1283,f1284,f1285,f1372,f1373,f1374,f1378,f1380,f1399,f1498,f1501,f1511,f1647,f1746,f1838,f2194,f2447,f2456,f1058,f409,f1059,f4344,f5345,f410,f8010,f8417,f8418,f10051,f9810,f9811])).
fof(f10692,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl8),
  inference(subsumption_resolution,[],[f10674,f10226])).
fof(f10674,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(superposition,[],[f10376,f1269])).
fof(f11595,plain,(
  $false | ($spl20 | $spl43 | $spl50)),
  inference(subsumption_resolution,[],[f11329,f11309])).
fof(f11309,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) | ($spl20 | $spl50)),
  inference(subsumption_resolution,[],[f11308,f87])).
fof(f11308,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X3)) ) | ($spl20 | $spl50)),
  inference(forward_demodulation,[],[f11307,f5345])).
fof(f11307,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl20),
  inference(subsumption_resolution,[],[f10892,f142])).
fof(f10892,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl20),
  inference(superposition,[],[f131,f1067])).
fof(f11329,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl20 | $spl43 | $spl50)),
  inference(backward_demodulation,[],[f11309,f1293])).
fof(f11594,plain,(
  $false | ($spl8 | $spl20 | $spl50)),
  inference(subsumption_resolution,[],[f11586,f11309])).
fof(f11586,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl8 | $spl20 | $spl50)),
  inference(backward_demodulation,[],[f11570,f1243])).
fof(f11570,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl20 | $spl50)),
  inference(subsumption_resolution,[],[f11318,f11309])).
fof(f11318,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl20 | $spl50)),
  inference(backward_demodulation,[],[f11309,f411])).
fof(f11593,plain,(
  $false | ($spl15 | $spl20 | $spl50)),
  inference(subsumption_resolution,[],[f11581,f11309])).
fof(f11581,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl15 | $spl20 | $spl50)),
  inference(backward_demodulation,[],[f11570,f438])).
fof(f11592,plain,(
  $false | ($spl20 | $spl50)),
  inference(subsumption_resolution,[],[f11571,f11570])).
fof(f11571,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ($spl20 | $spl50)),
  inference(backward_demodulation,[],[f11570,f86])).
fof(f11569,plain,(
  $false | ($spl20 | $spl50)),
  inference(subsumption_resolution,[],[f11310,f87])).
fof(f11310,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl50)),
  inference(backward_demodulation,[],[f11309,f89])).
fof(f175,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f9992,plain,(
  ( ! [X59,X58] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X58,X59),s(fun(X58,bool),'const.pred_set.EMPTY_0')))) ) | $spl52),
  inference(cnf_transformation,[],[f9992_D])).
fof(f9992_D,plain,(
  ( ! [X59,X58] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X58,X59),s(fun(X58,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f10223,plain,(
  ( ! [X412,X413,X411] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X412,X413),s(fun(X412,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X411)) ) | ($spl15 | $spl43 | $spl50)),
  inference(subsumption_resolution,[],[f9769,f6725])).
fof(f9769,plain,(
  ( ! [X412,X413,X411] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X412,X413),s(fun(X412,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X411) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X411)))))))))) ) | ($spl15 | $spl43 | $spl50)),
  inference(superposition,[],[f6725,f225])).
fof(f10213,plain,(
  ( ! [X391,X390,X392] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X391,X392),s(fun(X391,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X390)) ) | ($spl15 | $spl43 | $spl50)),
  inference(subsumption_resolution,[],[f9762,f6727])).
fof(f9762,plain,(
  ( ! [X391,X390,X392] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X391,X392),s(fun(X391,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X390) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X390)))))))) ) | ($spl15 | $spl43 | $spl50)),
  inference(superposition,[],[f6727,f225])).
fof(f10182,plain,(
  ( ! [X346,X345,X347] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X346,X347),s(fun(X346,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X345)) ) | ($spl15 | $spl43 | $spl50)),
  inference(subsumption_resolution,[],[f9747,f6401])).
fof(f9747,plain,(
  ( ! [X346,X345,X347] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X346,X347),s(fun(X346,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X345) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X345)))))) ) | ($spl15 | $spl43 | $spl50)),
  inference(superposition,[],[f6401,f225])).
fof(f10051,plain,(
  ( ! [X149,X147,X148] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X148,X149),s(fun(X148,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X147)) ) | $spl50),
  inference(forward_demodulation,[],[f9881,f5345])).
fof(f9881,plain,(
  ( ! [X149,X147,X148] : (s(bool,'const.bool.IN_2'(s(X148,X149),s(fun(X148,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X147)) ) | $spl50),
  inference(duplicate_literal_removal,[],[f9700])).
fof(f9700,plain,(
  ( ! [X149,X147,X148] : (s(bool,'const.bool.IN_2'(s(X148,X149),s(fun(X148,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,X147)) ) | $spl50),
  inference(superposition,[],[f4344,f225])).
fof(f9991,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X58,X59),s(fun(X58,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X57)) ) | ($spl30 | $spl32)),
  inference(forward_demodulation,[],[f9911,f1175])).
fof(f9911,plain,(
  ( ! [X59,X57,X58] : (s(bool,'const.bool.IN_2'(s(X58,X59),s(fun(X58,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X57)) ) | $spl32),
  inference(duplicate_literal_removal,[],[f9673])).
fof(f9673,plain,(
  ( ! [X59,X57,X58] : (s(bool,'const.bool.IN_2'(s(X58,X59),s(fun(X58,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) ) | $spl32),
  inference(superposition,[],[f1237,f225])).
fof(f7152,plain,(
  $false | ($spl15 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f7151,f7111])).
fof(f7111,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(subsumption_resolution,[],[f7033,f86])).
fof(f7033,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))) | $spl24),
  inference(superposition,[],[f128,f1171])).
fof(f7151,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl24 | $spl43)),
  inference(condensation,[],[f7150])).
fof(f7150,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl15 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f7128,f2130])).
fof(f7128,plain,(
  ( ! [X5] : (p(s(bool,f(s(bool,f(s(bool,X5)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl15 | $spl24 | $spl43)),
  inference(superposition,[],[f7096,f2024])).
fof(f7096,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl24),
  inference(subsumption_resolution,[],[f7095,f89])).
fof(f7095,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl24),
  inference(forward_demodulation,[],[f7094,f163])).
fof(f7094,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl24),
  inference(subsumption_resolution,[],[f7034,f87])).
fof(f7034,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl24),
  inference(superposition,[],[f129,f1171])).
fof(f7149,plain,(
  $false | ($spl15 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f7148,f7111])).
fof(f7148,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl24 | $spl43)),
  inference(condensation,[],[f7147])).
fof(f7147,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f7127,f2071])).
fof(f7127,plain,(
  ( ! [X4] : (p(s(bool,f(s(bool,X4)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl24 | $spl43)),
  inference(superposition,[],[f7096,f1483])).
fof(f7146,plain,(
  $false | ($spl8 | $spl15 | $spl24)),
  inference(subsumption_resolution,[],[f7145,f89])).
fof(f7145,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl15 | $spl24)),
  inference(forward_demodulation,[],[f7144,f1270])).
fof(f7144,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | ($spl8 | $spl24)),
  inference(subsumption_resolution,[],[f7126,f7111])).
fof(f7126,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl8 | $spl24)),
  inference(superposition,[],[f7096,f1269])).
fof(f7143,plain,(
  $false | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f7142,f89])).
fof(f7142,plain,(
  p(s(bool,'F_0')) | ($spl15 | $spl24)),
  inference(forward_demodulation,[],[f7141,f468])).
fof(f7141,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f7125,f7111])).
fof(f7125,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl24)),
  inference(superposition,[],[f7096,f496])).
fof(f7140,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f7139,f7111])).
fof(f7139,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(condensation,[],[f7138])).
fof(f7138,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(subsumption_resolution,[],[f7123,f142])).
fof(f7123,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(superposition,[],[f7096,f140])).
fof(f7137,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f7136,f7111])).
fof(f7136,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(condensation,[],[f7135])).
fof(f7135,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl24),
  inference(subsumption_resolution,[],[f7122,f142])).
fof(f7122,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl24),
  inference(superposition,[],[f7096,f140])).
fof(f7134,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f7133,f7111])).
fof(f7133,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(subsumption_resolution,[],[f7120,f89])).
fof(f7120,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(superposition,[],[f7096,f90])).
fof(f7131,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f7115,f7111])).
fof(f7115,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(unit_resulting_resolution,[],[f87,f7096,f135])).
fof(f7130,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f7117,f7111])).
fof(f7117,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(unit_resulting_resolution,[],[f87,f7096,f135])).
fof(f7129,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f7119,f7111])).
fof(f7119,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(unit_resulting_resolution,[],[f7096,f142])).
fof(f7098,plain,(
  $false | ($spl3 | $spl6 | $spl15 | $spl24 | $spl31 | $spl43)),
  inference(subsumption_resolution,[],[f7097,f89])).
fof(f7097,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl6 | $spl15 | $spl24 | $spl31 | $spl43)),
  inference(forward_demodulation,[],[f7096,f3177])).
fof(f3177,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl6 | $spl15 | $spl31 | $spl43)),
  inference(forward_demodulation,[],[f3176,f1365])).
fof(f3176,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl6 | $spl15 | $spl31 | $spl43)),
  inference(backward_demodulation,[],[f3175,f3160])).
fof(f3160,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))))))) | ($spl3 | $spl6 | $spl15 | $spl31 | $spl43)),
  inference(unit_resulting_resolution,[],[f1176,f2392])).
fof(f3175,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))))) | ($spl3 | $spl6 | $spl15 | $spl31 | $spl43)),
  inference(forward_demodulation,[],[f3159,f3160])).
fof(f3159,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))))))))) | ($spl6 | $spl15 | $spl31 | $spl43)),
  inference(unit_resulting_resolution,[],[f1176,f2375])).
fof(f1064,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | $spl16),
  inference(cnf_transformation,[],[f1064_D])).
fof(f1064_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4344,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl50),
  inference(cnf_transformation,[],[f4344_D])).
fof(f4344_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1059,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f409,f140])).
fof(f1058,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f409,f140])).
fof(f4343,plain,(
  $false | ($spl15 | $spl16 | $spl43)),
  inference(subsumption_resolution,[],[f4342,f4306])).
fof(f4306,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(subsumption_resolution,[],[f4242,f86])).
fof(f4242,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))) | $spl16),
  inference(superposition,[],[f128,f1064])).
fof(f4342,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl16 | $spl43)),
  inference(condensation,[],[f4341])).
fof(f4341,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl16 | $spl43)),
  inference(subsumption_resolution,[],[f4322,f2071])).
fof(f4322,plain,(
  ( ! [X4] : (p(s(bool,f(s(bool,X4)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl16 | $spl43)),
  inference(superposition,[],[f4293,f1483])).
fof(f4293,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl16),
  inference(subsumption_resolution,[],[f4292,f89])).
fof(f4292,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl16),
  inference(forward_demodulation,[],[f4291,f163])).
fof(f4291,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl16),
  inference(subsumption_resolution,[],[f4243,f87])).
fof(f4243,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl16),
  inference(superposition,[],[f129,f1064])).
fof(f4340,plain,(
  $false | ($spl8 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f4339,f89])).
fof(f4339,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl15 | $spl16)),
  inference(forward_demodulation,[],[f4338,f1270])).
fof(f4338,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f4321,f4306])).
fof(f4321,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl8 | $spl16)),
  inference(superposition,[],[f4293,f1269])).
fof(f4337,plain,(
  $false | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f4336,f89])).
fof(f4336,plain,(
  p(s(bool,'F_0')) | ($spl15 | $spl16)),
  inference(forward_demodulation,[],[f4335,f468])).
fof(f4335,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f4320,f4306])).
fof(f4320,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl16)),
  inference(superposition,[],[f4293,f496])).
fof(f4334,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4333,f4306])).
fof(f4333,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(condensation,[],[f4332])).
fof(f4332,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(subsumption_resolution,[],[f4318,f142])).
fof(f4318,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(superposition,[],[f4293,f140])).
fof(f4331,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4330,f4306])).
fof(f4330,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(condensation,[],[f4329])).
fof(f4329,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl16),
  inference(subsumption_resolution,[],[f4317,f142])).
fof(f4317,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl16),
  inference(superposition,[],[f4293,f140])).
fof(f4328,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4327,f4306])).
fof(f4327,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(subsumption_resolution,[],[f4315,f89])).
fof(f4315,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(superposition,[],[f4293,f90])).
fof(f4325,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4310,f4306])).
fof(f4310,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(unit_resulting_resolution,[],[f87,f4293,f135])).
fof(f4324,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4312,f4306])).
fof(f4312,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(unit_resulting_resolution,[],[f87,f4293,f135])).
fof(f4323,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4314,f4306])).
fof(f4314,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(unit_resulting_resolution,[],[f4293,f142])).
fof(f4295,plain,(
  $false | ($spl6 | $spl15 | $spl16 | $spl23 | $spl43)),
  inference(subsumption_resolution,[],[f4294,f89])).
fof(f4294,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl15 | $spl16 | $spl23 | $spl43)),
  inference(forward_demodulation,[],[f4293,f2460])).
fof(f2460,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl6 | $spl15 | $spl23 | $spl43)),
  inference(forward_demodulation,[],[f2437,f2439])).
fof(f2439,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) | ($spl6 | $spl15 | $spl23 | $spl43)),
  inference(unit_resulting_resolution,[],[f1069,f1495])).
fof(f2437,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) | ($spl15 | $spl23 | $spl43)),
  inference(unit_resulting_resolution,[],[f1069,f1069,f1483])).
fof(f2790,plain,(
  $false | ($spl6 | $spl15 | $spl27 | $spl31 | $spl43)),
  inference(subsumption_resolution,[],[f2789,f1176])).
fof(f2789,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl6 | $spl15 | $spl27 | $spl43)),
  inference(condensation,[],[f2788])).
fof(f2788,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl6 | $spl15 | $spl27 | $spl43)),
  inference(subsumption_resolution,[],[f2773,f1495])).
fof(f2773,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl27 | $spl43)),
  inference(superposition,[],[f1173,f1483])).
fof(f2787,plain,(
  $false | ($spl15 | $spl27 | $spl31)),
  inference(subsumption_resolution,[],[f2786,f1176])).
fof(f2786,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl27)),
  inference(subsumption_resolution,[],[f2772,f468])).
fof(f2772,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl27)),
  inference(superposition,[],[f1173,f496])).
fof(f2785,plain,(
  $false | ($spl27 | $spl31)),
  inference(subsumption_resolution,[],[f2784,f1176])).
fof(f2784,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl27),
  inference(condensation,[],[f2783])).
fof(f2783,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl27),
  inference(subsumption_resolution,[],[f2770,f90])).
fof(f2770,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl27),
  inference(superposition,[],[f1173,f140])).
fof(f2782,plain,(
  $false | ($spl27 | $spl31)),
  inference(subsumption_resolution,[],[f2781,f1176])).
fof(f2781,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl27),
  inference(condensation,[],[f2780])).
fof(f2780,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl27),
  inference(subsumption_resolution,[],[f2769,f90])).
fof(f2769,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl27),
  inference(superposition,[],[f1173,f140])).
fof(f2779,plain,(
  $false | ($spl27 | $spl31)),
  inference(subsumption_resolution,[],[f2774,f1176])).
fof(f2774,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl27),
  inference(trivial_inequality_removal,[],[f2767])).
fof(f2767,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl27),
  inference(superposition,[],[f1173,f90])).
fof(f2778,plain,(
  $false | ($spl27 | $spl31)),
  inference(subsumption_resolution,[],[f2760,f1176])).
fof(f2760,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl27),
  inference(unit_resulting_resolution,[],[f1173,f90])).
fof(f2777,plain,(
  $false | ($spl3 | $spl27 | $spl31)),
  inference(subsumption_resolution,[],[f2761,f1176])).
fof(f2761,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl27)),
  inference(unit_resulting_resolution,[],[f151,f1173,f140])).
fof(f2776,plain,(
  $false | ($spl3 | $spl27 | $spl31)),
  inference(subsumption_resolution,[],[f2764,f1176])).
fof(f2764,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl27)),
  inference(unit_resulting_resolution,[],[f151,f1173,f140])).
fof(f2775,plain,(
  $false | ($spl27 | $spl31)),
  inference(subsumption_resolution,[],[f2765,f1176])).
fof(f2765,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl27),
  inference(unit_resulting_resolution,[],[f89,f1173,f177])).
fof(f1513,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) | $spl46),
  inference(cnf_transformation,[],[f1513_D])).
fof(f1513_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1515,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl49),
  inference(cnf_transformation,[],[f1515_D])).
fof(f1515_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f1501,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK3(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f229,f140])).
fof(f1498,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f229,f140])).
fof(f1291,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl40),
  inference(cnf_transformation,[],[f1291_D])).
fof(f1291_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1288,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) | $spl36),
  inference(cnf_transformation,[],[f1288_D])).
fof(f1288_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1293,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl43),
  inference(cnf_transformation,[],[f1293_D])).
fof(f1293_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1285,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f228,f140])).
fof(f1287,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl34),
  inference(cnf_transformation,[],[f1287_D])).
fof(f1287_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1294,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl44),
  inference(cnf_transformation,[],[f1294_D])).
fof(f1294_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1284,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK3(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f228,f140])).
fof(f1283,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f228,f140])).
fof(f1282,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f228,f140])).
fof(f1290,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl39),
  inference(cnf_transformation,[],[f1290_D])).
fof(f1290_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f1281,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f228,f90])).
fof(f1171,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | $spl24),
  inference(cnf_transformation,[],[f1171_D])).
fof(f1171_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1237,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(cnf_transformation,[],[f1237_D])).
fof(f1237_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1165,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f410,f140])).
fof(f1166,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f410,f140])).
fof(f1235,plain,(
  $false | ($spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1182,f1173])).
fof(f1182,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f1179,f90])).
fof(f1179,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(backward_demodulation,[],[f1178,f410])).
fof(f1178,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | $spl11),
  inference(condensation,[],[f1177])).
fof(f1177,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))))) ) | $spl11),
  inference(subsumption_resolution,[],[f1165,f891])).
fof(f891,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),'const.pred_set.UNION_2'(s(fun(X53,bool),X56),s(fun(X53,bool),X55))))) | s(bool,X57) = s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),X55))) | s(bool,'T_0') = s(bool,X57)) ) | $spl11),
  inference(subsumption_resolution,[],[f684,f87])).
fof(f684,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),'const.pred_set.UNION_2'(s(fun(X53,bool),X56),s(fun(X53,bool),X55))))) | s(bool,X57) = s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),X55))) | s(bool,'T_0') = s(bool,X57)) ) | $spl11),
  inference(superposition,[],[f366,f140])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) ) | $spl11),
  inference(subsumption_resolution,[],[f365,f89])).
fof(f365,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) ) | $spl11),
  inference(forward_demodulation,[],[f358,f253])).
fof(f253,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f245,f90])).
fof(f245,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl11),
  inference(backward_demodulation,[],[f239,f226])).
fof(f239,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f89,f234,f177])).
fof(f358,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) ) | $spl11),
  inference(superposition,[],[f131,f273])).
fof(f273,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl11),
  inference(subsumption_resolution,[],[f266,f89])).
fof(f266,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl11),
  inference(superposition,[],[f177,f253])).
fof(f1234,plain,(
  $false | ($spl3 | $spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1183,f1173])).
fof(f1183,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl11)),
  inference(unit_resulting_resolution,[],[f151,f1179,f140])).
fof(f1233,plain,(
  $false | ($spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1232,f1173])).
fof(f1232,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(forward_demodulation,[],[f1184,f253])).
fof(f1184,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f245,f1179,f140])).
fof(f1231,plain,(
  $false | ($spl11 | $spl15 | $spl27)),
  inference(subsumption_resolution,[],[f1230,f1173])).
fof(f1230,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(forward_demodulation,[],[f1185,f468])).
fof(f1185,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(unit_resulting_resolution,[],[f438,f1179,f140])).
fof(f1229,plain,(
  $false | ($spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1228,f1173])).
fof(f1228,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(forward_demodulation,[],[f1186,f1075])).
fof(f1075,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f1072,f90])).
fof(f1072,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(backward_demodulation,[],[f1071,f409])).
fof(f1071,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | $spl11),
  inference(condensation,[],[f1070])).
fof(f1070,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))))) ) | $spl11),
  inference(subsumption_resolution,[],[f1058,f891])).
fof(f1186,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f1072,f1179,f140])).
fof(f1227,plain,(
  $false | ($spl3 | $spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1189,f1173])).
fof(f1189,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl11)),
  inference(unit_resulting_resolution,[],[f151,f1179,f140])).
fof(f1226,plain,(
  $false | ($spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1225,f1173])).
fof(f1225,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(forward_demodulation,[],[f1190,f253])).
fof(f1190,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f245,f1179,f140])).
fof(f1224,plain,(
  $false | ($spl11 | $spl15 | $spl27)),
  inference(subsumption_resolution,[],[f1223,f1173])).
fof(f1223,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(forward_demodulation,[],[f1191,f468])).
fof(f1191,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(unit_resulting_resolution,[],[f438,f1179,f140])).
fof(f1222,plain,(
  $false | ($spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1221,f1173])).
fof(f1221,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(forward_demodulation,[],[f1192,f1075])).
fof(f1192,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f1072,f1179,f140])).
fof(f1220,plain,(
  $false | ($spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1196,f1173])).
fof(f1196,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f89,f1179,f177])).
fof(f1219,plain,(
  $false | ($spl11 | $spl27)),
  inference(subsumption_resolution,[],[f1218,f1173])).
fof(f1218,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(forward_demodulation,[],[f1198,f253])).
fof(f1198,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f1179,f273])).
fof(f1216,plain,(
  $false | ($spl11 | $spl15 | $spl27)),
  inference(subsumption_resolution,[],[f1215,f1173])).
fof(f1215,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(forward_demodulation,[],[f1200,f468])).
fof(f1200,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(unit_resulting_resolution,[],[f1179,f496])).
fof(f1174,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f1174_D])).
fof(f1174_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1176,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl31),
  inference(cnf_transformation,[],[f1176_D])).
fof(f1176_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1164,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f410,f140])).
fof(f1173,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl27),
  inference(cnf_transformation,[],[f1173_D])).
fof(f1173_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1163,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f410,f90])).
fof(f1122,plain,(
  $false | ($spl11 | $spl19)),
  inference(subsumption_resolution,[],[f1075,f1066])).
fof(f1121,plain,(
  $false | ($spl3 | $spl11 | $spl19)),
  inference(subsumption_resolution,[],[f1076,f1066])).
fof(f1076,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl11)),
  inference(unit_resulting_resolution,[],[f151,f1072,f140])).
fof(f1120,plain,(
  $false | ($spl11 | $spl19)),
  inference(subsumption_resolution,[],[f1119,f1066])).
fof(f1119,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(forward_demodulation,[],[f1077,f253])).
fof(f1077,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f245,f1072,f140])).
fof(f1118,plain,(
  $false | ($spl11 | $spl15 | $spl19)),
  inference(subsumption_resolution,[],[f1117,f1066])).
fof(f1117,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(forward_demodulation,[],[f1078,f468])).
fof(f1078,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(unit_resulting_resolution,[],[f438,f1072,f140])).
fof(f1116,plain,(
  $false | ($spl3 | $spl11 | $spl19)),
  inference(subsumption_resolution,[],[f1081,f1066])).
fof(f1081,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl11)),
  inference(unit_resulting_resolution,[],[f151,f1072,f140])).
fof(f1115,plain,(
  $false | ($spl11 | $spl19)),
  inference(subsumption_resolution,[],[f1114,f1066])).
fof(f1114,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(forward_demodulation,[],[f1082,f253])).
fof(f1082,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f245,f1072,f140])).
fof(f1113,plain,(
  $false | ($spl11 | $spl15 | $spl19)),
  inference(subsumption_resolution,[],[f1112,f1066])).
fof(f1112,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(forward_demodulation,[],[f1083,f468])).
fof(f1083,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(unit_resulting_resolution,[],[f438,f1072,f140])).
fof(f1111,plain,(
  $false | ($spl11 | $spl19)),
  inference(subsumption_resolution,[],[f1087,f1066])).
fof(f1087,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f89,f1072,f177])).
fof(f1110,plain,(
  $false | ($spl11 | $spl19)),
  inference(subsumption_resolution,[],[f1109,f1066])).
fof(f1109,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(forward_demodulation,[],[f1089,f253])).
fof(f1089,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(unit_resulting_resolution,[],[f1072,f273])).
fof(f1107,plain,(
  $false | ($spl11 | $spl15 | $spl19)),
  inference(subsumption_resolution,[],[f1106,f1066])).
fof(f1106,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(forward_demodulation,[],[f1091,f468])).
fof(f1091,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl15)),
  inference(unit_resulting_resolution,[],[f1072,f496])).
fof(f1067,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f1067_D])).
fof(f1067_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1069,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl23),
  inference(cnf_transformation,[],[f1069_D])).
fof(f1069_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1057,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f409,f140])).
fof(f1066,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl19),
  inference(cnf_transformation,[],[f1066_D])).
fof(f1066_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1056,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f409,f90])).
fof(f425,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f427,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl15),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f423,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),bool,sK0))))),
  inference(superposition,[],[f227,f90])).
fof(f232,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl8),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f234,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f230,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(superposition,[],[f226,f90])).
fof(f187,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f180,f87])).
fof(f180,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f175,f89])).
fof(f176,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl6),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f178,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f170,f142])).
fof(f170,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f132,f90])).
fof(f174,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f168,f142])).
fof(f168,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f132,f90])).
fof(f149,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f151,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f151_D])).
fof(f151_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f148,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f90])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 101
% Passive clauses: 363
% Generated clauses: 11911
% Final active clauses: 16
% Final passive clauses: 17
% Input formulas: 23
% Initial clauses: 51
% 
% Duplicate literals: 760
% Trivial inequalities: 43
% Fw subsumption resolutions: 779
% Fw demodulations: 2766
% Bw demodulations: 277
% Condensations: 655
% Global subsumptions: 225
% 
% Simple tautologies: 129
% Equational tautologies: 126
% Forward subsumptions: 5712
% Fw demodulations to eq. taut.: 47
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 797
% Factoring: 1
% Forward superposition: 2757
% Backward superposition: 2353
% Self superposition: 88
% Equality factoring: 40
% Equality resolution: 5
% 
% Split clauses: 25
% Split components: 55
% SAT solver clauses: 921
% SAT solver unit clauses: 32
% SAT solver binary clauses: 188
% 
% Sat splits: 25
% Sat splitting refutations: 281
% 
% Memory used [KB]: 6396
% Time elapsed: 0.350 s
% ------------------------------
% ------------------------------
% Success in time 1.666 s
