Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (85581)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 747
% Passive clauses: 5605
% Generated clauses: 31742
% Final active clauses: 727
% Final passive clauses: 4523
% Input formulas: 69
% Initial clauses: 184
% 
% Duplicate literals: 593
% Trivial inequalities: 69
% Fw subsumption resolutions: 3003
% Fw demodulations: 7382
% Bw demodulations: 20
% 
% Simple tautologies: 730
% Equational tautologies: 27
% Forward subsumptions: 14195
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 9265
% Factoring: 10
% Forward superposition: 4586
% Backward superposition: 6476
% Self superposition: 57
% Equality factoring: 26
% Equality resolution: 3
% 
% Split clauses: 247
% Split components: 495
% SAT solver clauses: 450
% SAT solver unit clauses: 16
% SAT solver binary clauses: 246
% 
% Sat splits: 247
% 
% TWLsolver clauses: 450
% TWLsolver calls for satisfiability: 127
% 
% Memory used [KB]: 17910
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (85596)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 259
% Passive clauses: 2135
% Generated clauses: 11504
% Final active clauses: 256
% Final passive clauses: 1875
% Input formulas: 69
% Initial clauses: 184
% 
% Duplicate literals: 241
% Trivial inequalities: 90
% Fw subsumption resolutions: 569
% Fw demodulations: 1532
% Bw demodulations: 1
% Condensations: 285
% Global subsumptions: 18
% 
% Simple tautologies: 225
% Equational tautologies: 25
% Forward subsumptions: 3894
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 710
% Factoring: 4
% Forward superposition: 2327
% Backward superposition: 3361
% Self superposition: 65
% Equality factoring: 31
% Equality resolution: 12
% 
% Split clauses: 106
% Split components: 212
% SAT solver clauses: 2688
% SAT solver unit clauses: 192
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 106
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2616
% TWLsolver calls for satisfiability: 1224
% 
% Memory used [KB]: 22515
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f8685,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f310,f504,f503,f309,f506,f308,f508,f502,f1319,f1327,f1326,f1320,f1330,f1329,f2101,f2103,f2104,f2105,f2102,f2106,f2108,f2110,f2114,f2116,f2120,f2122,f2140,f2194,f2200,f2205,f2207,f2215,f2224,f2233,f2235,f2241,f2246,f2248,f2253,f2330,f2336,f2341,f2343,f2351,f2360,f2369,f2371,f2377,f2382,f2384,f2389,f2430,f2432,f2434,f2442,f2444,f2446,f2462,f2464,f2472,f2474,f2490,f2492,f2500,f2502,f2518,f2520,f2528,f2530,f2546,f2548,f2716,f2717,f2718,f2562,f2564,f2566,f2571,f2573,f2575,f2583,f2585,f2590,f2592,f2597,f2610,f2617,f2619,f2621,f2630,f2632,f2638,f2643,f2648,f2654,f2664,f2666,f2676,f2678,f2729,f2730,f2731,f2732,f2733,f2734,f2735,f2736,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2749,f2750,f2751,f2752,f2753,f2754,f2755,f2756,f2757,f2758,f2760,f2761,f2762,f2763,f2764,f2765,f2766,f2767,f2768,f2769,f2770,f2771,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f2781,f2782,f2783,f2784,f2785,f2786,f2787,f2788,f2789,f2790,f2791,f2792,f2793,f2794,f2795,f2796,f2797,f2798,f2799,f2800,f2801,f2802,f2803,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f2813,f2814,f2815,f2816,f2817,f2818,f2819,f2820,f2821,f2822,f2823,f2824,f2825,f2826,f2827,f2828,f2829,f2830,f2831,f2832,f2833,f2834,f2835,f2836,f2837,f2838,f2839,f2840,f2841,f2842,f2843,f2844,f2845,f2846,f2847,f2848,f2849,f2850,f2851,f2852,f2853,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f4509,f4552,f4553,f4554,f4797,f5072,f5347,f5622,f5897,f6172,f6444,f6467,f6490,f6513,f6536,f6559,f6571,f6979,f6990,f6991,f6998,f8663])).
fof(f8663,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f8661,f7977])).
fof(f7977,plain,(
  s('type.list.list'(sK8),sK10) != s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK23(sK11,sK9,sK10,'const.list.NIL_0',sK8,sK7)),s('type.list.list'(sK8),'const.list.NIL_0'))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f7008,f504])).
fof(f7008,plain,(
  s(sK7,sK9) = s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK23(sK11,sK9,sK10,'const.list.NIL_0',sK8,sK7)))) | $spl0),
  inference(unit_resulting_resolution,[],[f502,f432])).
fof(f432,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,sK23(X5,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK23(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK24(X5,X4,X3,X2,X1,X0)))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,sK23(X5,X4,X3,X2,X1,X0)))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),sK24(X5,X4,X3,X2,X1,X0)))))) & (! [X6,X7] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | ? [X8,X9] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X8),s('type.list.list'(X1),X9))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X9))))) & (! [X6,X7] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3)))))),
  inference(rectify,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7))))) & (! [X6,X7] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3)))))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) <=> ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) <=> ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : ! [X13,X9,X12,X10] : (s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X10),s('type.list.list'(X7),X9))) <=> ? [X11,X14] : (s('type.list.list'(X7),X9) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X11),s('type.list.list'(X7),X14))) & s(X6,X12) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X11))) & s('type.list.list'(X6),X13) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X10),s('type.list.list'(X7),X14)))))),
  file('thm.list.MAP_EQ_SING','thm.list.MAP_EQ_CONS')).
fof(f8661,plain,(
  s('type.list.list'(sK8),sK10) = s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK23(sK11,sK9,sK10,'const.list.NIL_0',sK8,sK7)),s('type.list.list'(sK8),'const.list.NIL_0'))) | $spl0),
  inference(backward_demodulation,[],[f8514,f7006])).
fof(f7006,plain,(
  s('type.list.list'(sK8),sK10) = s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK23(sK11,sK9,sK10,'const.list.NIL_0',sK8,sK7)),s('type.list.list'(sK8),sK24(sK11,sK9,sK10,'const.list.NIL_0',sK8,sK7)))) | $spl0),
  inference(unit_resulting_resolution,[],[f502,f431])).
fof(f431,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK23(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK24(X5,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f287])).
fof(f8514,plain,(
  s('type.list.list'(sK8),'const.list.NIL_0') = s('type.list.list'(sK8),sK24(sK11,sK9,sK10,'const.list.NIL_0',sK8,sK7)) | $spl0),
  inference(unit_resulting_resolution,[],[f7007,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3] : ((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2)) & (s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X2) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X2))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : ! [X9,X10] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X10),s('type.list.list'(X7),X9))) <=> s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X9))),
  file('thm.list.MAP_EQ_SING','thm.list.MAP_EQ_NIL|split|1')).
fof(f7007,plain,(
  s('type.list.list'(sK7),'const.list.NIL_0') = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK24(sK11,sK9,sK10,'const.list.NIL_0',sK8,sK7)))) | $spl0),
  inference(unit_resulting_resolution,[],[f502,f433])).
fof(f433,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),sK24(X5,X4,X3,X2,X1,X0)))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3)))) )),
  inference(cnf_transformation,[],[f287])).
fof(f6998,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6989,f506])).
fof(f6989,plain,(
  s(sK7,sK9) != s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK12))) | ($spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f6988])).
fof(f6988,plain,(
  s('type.list.list'(sK8),sK10) != s('type.list.list'(sK8),sK10) | s(sK7,sK9) != s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK12))) | ($spl2 | $spl6)),
  inference(superposition,[],[f504,f508])).
fof(f6991,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6977,f508])).
fof(f6977,plain,(
  s('type.list.list'(sK8),sK10) != s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))) | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f506,f504])).
fof(f6990,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6978,f506])).
fof(f6978,plain,(
  s(sK7,sK9) != s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK12))) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f508,f504])).
fof(f6979,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f504])).
fof(f6571,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f3689,f503])).
fof(f3689,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f421,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7)))) )),
  inference(cnf_transformation,[],[f287])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : ! [X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X6] : ! [X10] : s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X10),s('type.list.list'(X7),'const.list.NIL_0')))),
  file('thm.list.MAP_EQ_SING','thm.list.MAP|split|0')).
fof(f6559,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6558,f503])).
fof(f6558,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f6557,f508])).
fof(f6557,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))) | $spl4),
  inference(forward_demodulation,[],[f3701,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : ! [X9] : s('type.list.list'(X7),X9) = s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X7,X7),'const.combin.I_0'),s('type.list.list'(X7),X9)))),
  file('thm.list.MAP_EQ_SING','thm.list.MAP_ID|split|1')).
fof(f3701,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.MAP_2'(s(fun(sK8,sK8),'const.combin.I_0'),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | $spl4),
  inference(unit_resulting_resolution,[],[f506,f372,f421,f434])).
fof(f6536,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6535,f503])).
fof(f6535,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f6534,f508])).
fof(f6534,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))) | $spl4),
  inference(forward_demodulation,[],[f3715,f373])).
fof(f373,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'_i_0_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'_i_0_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'_i_0_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : ! [X9] : s('type.list.list'(X7),X9) = s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X7,X7),'_i_0_0'),s('type.list.list'(X7),X9)))),
  file('thm.list.MAP_EQ_SING','thm.list.MAP_ID|split|0')).
fof(f3715,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.MAP_2'(s(fun(sK8,sK8),'_i_0_0'),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | $spl4),
  inference(unit_resulting_resolution,[],[f506,f373,f421,f434])).
fof(f6513,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6512,f503])).
fof(f6512,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f6511,f508])).
fof(f6511,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))) | $spl4),
  inference(forward_demodulation,[],[f3729,f369])).
fof(f369,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X7] : ! [X8] : s(X7,X8) = s(X7,'const.combin.I_1'(s(X7,X8)))),
  file('thm.list.MAP_EQ_SING','thm.combin.I_THM')).
fof(f3729,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.combin.I_1'(s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | $spl4),
  inference(unit_resulting_resolution,[],[f506,f369,f421,f434])).
fof(f6490,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6489,f503])).
fof(f6489,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f6488,f508])).
fof(f6488,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))) | $spl4),
  inference(forward_demodulation,[],[f3743,f370])).
fof(f370,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X19] : ! [X21] : s(X19,X21) = s(X19,happ(s(fun(X19,X19),'_i_0_0'),s(X19,X21)))),
  file('thm.list.MAP_EQ_SING','_i0_ax')).
fof(f3743,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),happ(s(fun('type.list.list'(sK8),'type.list.list'(sK8)),'_i_0_0'),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | $spl4),
  inference(unit_resulting_resolution,[],[f506,f370,f421,f434])).
fof(f6467,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6466,f503])).
fof(f6466,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f6465,f508])).
fof(f6465,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))) | $spl4),
  inference(forward_demodulation,[],[f3757,f510])).
fof(f510,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f371,f369])).
fof(f371,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X7] : ! [X23] : s(X7,'const.combin.I_1'(s(X7,X23))) = s(X7,happ(s(fun(X7,X7),'const.combin.I_0'),s(X7,X23)))),
  file('thm.list.MAP_EQ_SING','ARITY_0_1_const.combin.I')).
fof(f3757,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),happ(s(fun('type.list.list'(sK8),'type.list.list'(sK8)),'const.combin.I_0'),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | $spl4),
  inference(unit_resulting_resolution,[],[f506,f510,f421,f434])).
fof(f6444,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6443,f503])).
fof(f6443,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f6442,f508])).
fof(f6442,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))))) | $spl4),
  inference(forward_demodulation,[],[f3771,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7] : ! [X35,X10] : s(X7,X35) = s(X7,'const.ind_type.FCONS_3'(s(X7,X35),s(fun('type.num.num',X7),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.list.MAP_EQ_SING','thm.ind_type.FCONS|split|0')).
fof(f3771,plain,(
  ( ! [X0] : (s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.ind_type.FCONS_3'(s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))),s(fun('type.num.num','type.list.list'(sK8)),X0),s('type.num.num','const.num.|48|_0')))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f506,f383,f421,f434])).
fof(f6172,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f6171,f503])).
fof(f6171,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f6170,f421])).
fof(f6170,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f3880,f372])).
fof(f3880,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK7,sK7),'const.combin.I_0'),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f372,f434])).
fof(f5897,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f5896,f503])).
fof(f5896,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f5895,f421])).
fof(f5895,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f3987,f373])).
fof(f3987,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK7,sK7),'_i_0_0'),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f373,f434])).
fof(f5622,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f5621,f503])).
fof(f5621,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f5620,f421])).
fof(f5620,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f4094,f369])).
fof(f4094,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.combin.I_1'(s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f369,f434])).
fof(f5347,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f5346,f503])).
fof(f5346,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f5345,f421])).
fof(f5345,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f4201,f370])).
fof(f4201,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),happ(s(fun('type.list.list'(sK7),'type.list.list'(sK7)),'_i_0_0'),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f370,f434])).
fof(f5072,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f5071,f503])).
fof(f5071,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f5070,f421])).
fof(f5070,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f4308,f510])).
fof(f4308,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),happ(s(fun('type.list.list'(sK7),'type.list.list'(sK7)),'const.combin.I_0'),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f510,f434])).
fof(f4797,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f4796,f503])).
fof(f4796,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f4795,f421])).
fof(f4795,plain,(
  s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f4415,f383])).
fof(f4415,plain,(
  ( ! [X0] : (s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) = s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.ind_type.FCONS_3'(s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))),s(fun('type.num.num','type.list.list'(sK7)),X0),s('type.num.num','const.num.|48|_0')))))) ) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f383,f434])).
fof(f4554,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f4506,f421])).
fof(f4506,plain,(
  s('type.list.list'(sK7),'const.list.NIL_0') != s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),'const.list.NIL_0'))) | ($spl1 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f508,f506,f503,f434])).
fof(f4553,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f4507,f508])).
fof(f4507,plain,(
  s('type.list.list'(sK8),sK10) != s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f506,f421,f503,f434])).
fof(f4552,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f4508,f506])).
fof(f4508,plain,(
  s(sK7,sK9) != s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK12))) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f508,f421,f503,f434])).
fof(f4509,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f506,f508,f421,f503,f434])).
fof(f2860,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2107,f2118])).
fof(f2118,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f2102,f2102,f398])).
fof(f398,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X15,X16] : ((p(s(bool,X15)) => p(s(bool,X16))) => ((p(s(bool,X16)) => p(s(bool,X15))) => s(bool,X15) = s(bool,X16)))),
  file('thm.list.MAP_EQ_SING','thm.bool.IMP_ANTISYM_AX')).
fof(f2107,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f2102,f328])).
fof(f328,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.list.MAP_EQ_SING','thm.bool.EQ_CLAUSES|split|0')).
fof(f2859,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2115,f509])).
fof(f509,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f314,f311])).
fof(f311,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.MAP_EQ_SING',aHL_TRUTH)).
fof(f314,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f46])).
fof(f46,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.list.MAP_EQ_SING','thm.bool.NOT_CLAUSES|split|2')).
fof(f2115,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f531,f2102,f398])).
fof(f531,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f509,f328])).
fof(f2858,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2121,f509])).
fof(f2121,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f531,f2102,f398])).
fof(f2857,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2139,f2326])).
fof(f2326,plain,(
  ( ! [X0] : (sP4('T_0','T_0',X0)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f520,f514,f517,f2102,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~sP2(X1,X2,X0) | p(s(bool,X0)) | ~sP3(X0,X2,X1) | sP4(X0,X1,X2) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (sP1(X2,X1,X0) & sP2(X1,X2,X0) & sP3(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP1(X2,X1,X0) | ~sP2(X1,X2,X0) | ~sP3(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f289])).
fof(f289,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP4(X2,X1,X0)))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP4(X2,X1,X0)))),
  inference(nnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X2,X1,X0] : (sP4(X2,X1,X0) <=> (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f517,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f311,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f295])).
fof(f295,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(nnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X1,X0,X2] : (sP2(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f514,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f311,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f292])).
fof(f292,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f520,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f311,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(nnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : (sP1(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2139,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',X0)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f311,f311,f2102,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f290])).
fof(f2856,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2141,f2445])).
fof(f2445,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f2102,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f296])).
fof(f2141,plain,(
  ~sP2('T_0','F_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f514,f2102,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (~sP2(X1,X2,X0) | p(s(bool,X1)) | ~sP3(X0,X2,X1) | sP4(X0,X1,X2) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f290])).
fof(f706,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f311,f531,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP4(X2,X1,X0)) & (~sP4(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP4(X2,X1,X0))),
  inference(definition_folding,[],[f175,f225,f224,f223,f222])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X30,X39,X40] : ((p(s(bool,X40)) <=> s(bool,X30) = s(bool,X39)) <=> ((p(s(bool,X40)) | p(s(bool,X39)) | p(s(bool,X30))) & (p(s(bool,X40)) | ~p(s(bool,X30)) | ~p(s(bool,X39))) & (p(s(bool,X39)) | ~p(s(bool,X30)) | ~p(s(bool,X40))) & (p(s(bool,X30)) | ~p(s(bool,X39)) | ~p(s(bool,X40)))))),
  file('thm.list.MAP_EQ_SING','thm.sat.sat_2|49|_prove_ltb')).
fof(f521,plain,(
  ( ! [X0,X1] : (sP1(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f311,f452])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f298])).
fof(f2855,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2146,f2445])).
fof(f2146,plain,(
  ~sP2('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f514,f2102,f440])).
fof(f1382,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'T_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f531,f1326,f456])).
fof(f2854,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2147,f2445])).
fof(f2147,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f1361,f2102,f440])).
fof(f1361,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1326,f443])).
fof(f1341,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f311,f509,f1326,f438])).
fof(f1372,plain,(
  ( ! [X0,X1] : (sP1(X0,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1326,f452])).
fof(f2853,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2148,f2445])).
fof(f2148,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f716,f1361,f2102,f440])).
fof(f716,plain,(
  ( ! [X0] : (~sP4('F_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f509,f369,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f2852,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2149,f2445])).
fof(f2149,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f717,f1361,f2102,f440])).
fof(f717,plain,(
  ( ! [X0] : (~sP4('F_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f509,f370,f457])).
fof(f2851,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2150,f2445])).
fof(f2150,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f737,f1361,f2102,f440])).
fof(f737,plain,(
  ( ! [X0] : (~sP4('F_0',X0,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f509,f510,f457])).
fof(f2850,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2151,f2445])).
fof(f2151,plain,(
  ~sP2('T_0','F_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f515,f2102,f440])).
fof(f515,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f509,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f293])).
fof(f2849,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2152,f2445])).
fof(f2152,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f515,f2102,f440])).
fof(f2848,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2153,f2445])).
fof(f2153,plain,(
  ~sP2('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f515,f2102,f440])).
fof(f2847,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2155,f2445])).
fof(f2155,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f716,f516,f2102,f440])).
fof(f516,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f509,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f293])).
fof(f2846,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2157,f2445])).
fof(f2157,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f717,f516,f2102,f440])).
fof(f2845,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2159,f2445])).
fof(f2159,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f737,f516,f2102,f440])).
fof(f2844,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2162,f2445])).
fof(f2162,plain,(
  ~sP2(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f715,f516,f2102,f440])).
fof(f715,plain,(
  ( ! [X0] : (~sP4('F_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f509,f370,f457])).
fof(f1373,plain,(
  ( ! [X0,X1] : (sP1('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1326,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f2843,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2164,f2445])).
fof(f2164,plain,(
  ~sP2(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f735,f516,f2102,f440])).
fof(f735,plain,(
  ( ! [X0] : (~sP4('F_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f509,f510,f457])).
fof(f2842,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2166,f2445])).
fof(f2166,plain,(
  ~sP2('const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f713,f516,f2102,f440])).
fof(f713,plain,(
  ( ! [X0] : (~sP4('F_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f509,f369,f457])).
fof(f2841,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2168,f2433])).
fof(f2433,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f2102,f445])).
fof(f2168,plain,(
  ~sP3('T_0','F_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f517,f2102,f440])).
fof(f2840,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2169,f2433])).
fof(f2169,plain,(
  ~sP3('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f517,f2102,f440])).
fof(f2839,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2208,f2433])).
fof(f2208,plain,(
  ~sP3('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f1370,f2102,f440])).
fof(f1370,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1326,f447])).
fof(f2838,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2225,f2433])).
fof(f2225,plain,(
  ~sP3('T_0','F_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f518,f2102,f440])).
fof(f518,plain,(
  ( ! [X0,X1] : (sP2(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f509,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f296])).
fof(f2837,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2226,f2433])).
fof(f2226,plain,(
  ~sP3('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f518,f2102,f440])).
fof(f2836,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2227,f2433])).
fof(f2227,plain,(
  ~sP3('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f518,f2102,f440])).
fof(f2835,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2277,f2445])).
fof(f2277,plain,(
  ~sP2('T_0','F_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f514,f2102,f441])).
fof(f2834,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2282,f2445])).
fof(f2282,plain,(
  ~sP2('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f514,f2102,f441])).
fof(f2833,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2283,f2445])).
fof(f2283,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f1361,f2102,f441])).
fof(f2832,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2284,f2445])).
fof(f2284,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f716,f1361,f2102,f441])).
fof(f2831,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2285,f2445])).
fof(f2285,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f717,f1361,f2102,f441])).
fof(f2830,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2286,f2445])).
fof(f2286,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f737,f1361,f2102,f441])).
fof(f2829,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2287,f2445])).
fof(f2287,plain,(
  ~sP2('T_0','F_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f515,f2102,f441])).
fof(f2828,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2288,f2445])).
fof(f2288,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f515,f2102,f441])).
fof(f2827,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2289,f2445])).
fof(f2289,plain,(
  ~sP2('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f515,f2102,f441])).
fof(f2826,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2291,f2445])).
fof(f2291,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f716,f516,f2102,f441])).
fof(f2825,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2293,f2445])).
fof(f2293,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f717,f516,f2102,f441])).
fof(f2824,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2295,f2445])).
fof(f2295,plain,(
  ~sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f737,f516,f2102,f441])).
fof(f2823,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2298,f2445])).
fof(f2298,plain,(
  ~sP2(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f715,f516,f2102,f441])).
fof(f2822,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2300,f2445])).
fof(f2300,plain,(
  ~sP2(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f735,f516,f2102,f441])).
fof(f2821,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2302,f2445])).
fof(f2302,plain,(
  ~sP2('const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f713,f516,f2102,f441])).
fof(f2820,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2304,f2433])).
fof(f2304,plain,(
  ~sP3('T_0','F_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f517,f2102,f441])).
fof(f2819,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2305,f2433])).
fof(f2305,plain,(
  ~sP3('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f517,f2102,f441])).
fof(f2818,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2344,f2433])).
fof(f2344,plain,(
  ~sP3('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f1370,f2102,f441])).
fof(f2817,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2361,f2433])).
fof(f2361,plain,(
  ~sP3('T_0','F_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f518,f2102,f441])).
fof(f2816,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2362,f2433])).
fof(f2362,plain,(
  ~sP3('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f518,f2102,f441])).
fof(f2815,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2363,f2433])).
fof(f2363,plain,(
  ~sP3('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f518,f2102,f441])).
fof(f2814,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2386,f1341])).
fof(f2386,plain,(
  sP4('T_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f520,f1361,f518,f2102,f441])).
fof(f2813,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2387,f1342])).
fof(f1342,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1326,f509,f1326,f438])).
fof(f2387,plain,(
  sP4('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f1361,f518,f2102,f441])).
fof(f1371,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1326,f451])).
fof(f2812,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2429,f2433])).
fof(f2429,plain,(
  ( ! [X0] : (~sP3('T_0','T_0',X0)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f311,f311,f2102,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f293])).
fof(f2811,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2441,f2445])).
fof(f2441,plain,(
  ( ! [X0] : (~sP2('T_0','T_0',X0)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f311,f311,f2102,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f296])).
fof(f2810,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2448,f2102])).
fof(f2448,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f509,f2102,f450])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f2809,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2449,f2102])).
fof(f2449,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f522,f509,f2102,f450])).
fof(f522,plain,(
  ( ! [X0,X1] : (sP1('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f311,f453])).
fof(f2808,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2450,f2102])).
fof(f2450,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f509,f2102,f450])).
fof(f2807,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2451,f2102])).
fof(f2451,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f2102,f450])).
fof(f2806,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2452,f2102])).
fof(f2452,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f2102,f2102,f450])).
fof(f2805,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2454,f2102])).
fof(f2454,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f2102,f2102,f450])).
fof(f2804,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2455,f2102])).
fof(f2455,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f522,f2102,f2102,f450])).
fof(f2803,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2456,f2102])).
fof(f2456,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f2102,f2102,f450])).
fof(f2802,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2458,f2102])).
fof(f2458,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f509,f2102,f450])).
fof(f2801,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2459,f2102])).
fof(f2459,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f522,f509,f2102,f450])).
fof(f2800,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2460,f2102])).
fof(f2460,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f509,f2102,f450])).
fof(f2799,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2466,f2102])).
fof(f2466,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f2102,f2102,f450])).
fof(f2798,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2467,f2102])).
fof(f2467,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f521,f2102,f2102,f450])).
fof(f2797,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2468,f2102])).
fof(f2468,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f2102,f2102,f450])).
fof(f2796,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2469,f2102])).
fof(f2469,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f522,f2102,f2102,f450])).
fof(f2795,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2470,f2102])).
fof(f2470,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f2102,f2102,f450])).
fof(f2794,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2476,f2102])).
fof(f2476,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f509,f2102,f450])).
fof(f2793,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2477,f2102])).
fof(f2477,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f521,f509,f2102,f450])).
fof(f2792,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2478,f2102])).
fof(f2478,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f509,f2102,f450])).
fof(f2791,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2480,f2102])).
fof(f2480,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f2102,f2102,f450])).
fof(f2790,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2481,f2102])).
fof(f2481,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f521,f2102,f2102,f450])).
fof(f2789,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2482,f2102])).
fof(f2482,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f2102,f2102,f450])).
fof(f2788,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2483,f2102])).
fof(f2483,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f522,f2102,f2102,f450])).
fof(f2787,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2484,f2102])).
fof(f2484,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f2102,f2102,f450])).
fof(f2786,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2485,f2102])).
fof(f2485,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f521,f509,f2102,f450])).
fof(f2785,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2486,f2102])).
fof(f2486,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f509,f2102,f450])).
fof(f2784,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2488,f2102])).
fof(f2488,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f509,f2102,f450])).
fof(f2783,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2493,f2102])).
fof(f2493,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f2102,f450])).
fof(f2782,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2494,f2102])).
fof(f2494,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f2102,f2102,f450])).
fof(f2781,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2495,f2102])).
fof(f2495,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f521,f2102,f2102,f450])).
fof(f2780,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2496,f2102])).
fof(f2496,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f2102,f2102,f450])).
fof(f2779,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2498,f2102])).
fof(f2498,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f2102,f2102,f450])).
fof(f2778,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2503,f2102])).
fof(f2503,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f520,f509,f2102,f450])).
fof(f2777,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2504,f2102])).
fof(f2504,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f509,f2102,f450])).
fof(f2776,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2506,f2102])).
fof(f2506,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f509,f2102,f450])).
fof(f2775,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2507,f2102])).
fof(f2507,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f2102,f450])).
fof(f2774,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2508,f2102])).
fof(f2508,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f2102,f2102,f450])).
fof(f2773,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2510,f2102])).
fof(f2510,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f2102,f2102,f450])).
fof(f2772,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2511,f2102])).
fof(f2511,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f522,f2102,f2102,f450])).
fof(f2771,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2512,f2102])).
fof(f2512,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f2102,f2102,f450])).
fof(f2770,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2513,f2102])).
fof(f2513,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f520,f509,f2102,f450])).
fof(f2769,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2514,f2102])).
fof(f2514,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f509,f2102,f450])).
fof(f2768,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2516,f2102])).
fof(f2516,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f509,f2102,f450])).
fof(f2767,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2521,f2102])).
fof(f2521,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f2102,f450])).
fof(f2766,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2522,f2102])).
fof(f2522,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1371,f2102,f2102,f450])).
fof(f2765,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2523,f2102])).
fof(f2523,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f521,f2102,f2102,f450])).
fof(f2764,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2524,f2102])).
fof(f2524,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f2102,f2102,f450])).
fof(f2763,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2526,f2102])).
fof(f2526,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1373,f2102,f2102,f450])).
fof(f2762,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2531,f531])).
fof(f2531,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl16),
  inference(unit_resulting_resolution,[],[f704,f2102,f454])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f704,plain,(
  ~sP4('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f311,f531,f456])).
fof(f2761,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2532,f531])).
fof(f2532,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl16),
  inference(unit_resulting_resolution,[],[f706,f2102,f454])).
fof(f2760,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2759,f531])).
fof(f2759,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl16)),
  inference(forward_demodulation,[],[f2533,f1335])).
fof(f1335,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl8),
  inference(unit_resulting_resolution,[],[f311,f1326,f401])).
fof(f401,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f201])).
fof(f2533,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1341,f2102,f454])).
fof(f2758,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2543,f531])).
fof(f2543,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1382,f2102,f454])).
fof(f2757,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2544,f531])).
fof(f2544,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1383,f2102,f454])).
fof(f1383,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f531,f1326,f456])).
fof(f2756,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2551,f369])).
fof(f2551,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f690,f2102,f457])).
fof(f690,plain,(
  ( ! [X0] : (sP4('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f311,f369,f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f2755,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2552,f370])).
fof(f2552,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f691,f2102,f457])).
fof(f691,plain,(
  ( ! [X0] : (sP4('T_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f311,f370,f455])).
fof(f2754,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2553,f510])).
fof(f2553,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f738,f2102,f457])).
fof(f738,plain,(
  ( ! [X0] : (sP4('T_0',X0,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f311,f510,f455])).
fof(f2753,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2555,f1335])).
fof(f2555,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1349,f2102,f457])).
fof(f1349,plain,(
  sP4('T_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl8),
  inference(unit_resulting_resolution,[],[f521,f514,f517,f1326,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (~sP2(X1,X2,X0) | ~p(s(bool,X2)) | ~sP3(X0,X2,X1) | sP4(X0,X1,X2) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f290])).
fof(f2752,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2557,f370])).
fof(f2557,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f689,f2102,f457])).
fof(f689,plain,(
  ( ! [X0] : (sP4('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f311,f370,f455])).
fof(f2751,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2558,f510])).
fof(f2558,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f736,f2102,f457])).
fof(f736,plain,(
  ( ! [X0] : (sP4('T_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f311,f510,f455])).
fof(f2750,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2559,f369])).
fof(f2559,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f688,f2102,f457])).
fof(f688,plain,(
  ( ! [X0] : (sP4('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f311,f369,f455])).
fof(f2749,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2579,f509])).
fof(f2579,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f567,f2102,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP5(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP5(X0,X1,X2) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f301])).
fof(f301,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP5(X2,X0,X1) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP5(X2,X0,X1)) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f567,plain,(
  ( ! [X0] : (sP6('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f509,f311,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP6(X2,X0,X1))),
  inference(definition_folding,[],[f215,f228,f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X30,X39,X40] : ((p(s(bool,X40)) <=> (p(s(bool,X39)) => p(s(bool,X30)))) <=> ((p(s(bool,X40)) | p(s(bool,X39))) & (p(s(bool,X40)) | ~p(s(bool,X30))) & (~p(s(bool,X39)) | p(s(bool,X30)) | ~p(s(bool,X40)))))),
  file('thm.list.MAP_EQ_SING','thm.sat.sat_27_prove_ltb')).
fof(f2748,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2580,f2102])).
fof(f2580,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f574,f2102,f458])).
fof(f574,plain,(
  ( ! [X0] : (sP6('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f311,f311,f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f307])).
fof(f2747,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2588,f509])).
fof(f2588,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f907,f2102,f458])).
fof(f907,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f525,f311,f509,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f525,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f509,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f304])).
fof(f304,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f227])).
fof(f2746,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2609,f2620])).
fof(f2620,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f2102,f468])).
fof(f2609,plain,(
  ~sP5('F_0','T_0','T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f564,f311,f2102,f464])).
fof(f564,plain,(
  ( ! [X0] : (~sP6('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f311,f509,f459])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f302])).
fof(f2745,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2616,f2620])).
fof(f2616,plain,(
  ( ! [X0] : (~sP5('T_0',X0,'T_0')) ) | $spl16),
  inference(unit_resulting_resolution,[],[f311,f311,f2102,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f305])).
fof(f2744,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2627,f509])).
fof(f2627,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f564,f311,f2102,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f307])).
fof(f2743,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2634,f509])).
fof(f2634,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f925,f311,f2102,f469])).
fof(f925,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f920,f460])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f302])).
fof(f920,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f311,f311,f509,f465])).
fof(f2742,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2656,f2102])).
fof(f2656,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl16),
  inference(unit_resulting_resolution,[],[f543,f2102,f511])).
fof(f511,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f406,f509])).
fof(f406,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f275])).
fof(f275,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f543,plain,(
  ( ! [X0] : (sP0(X0,'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f311,f530])).
fof(f530,plain,(
  ( ! [X2,X3] : (sP0(X3,'const.combin.I_1'(s(bool,X2))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f407,f369])).
fof(f407,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f276])).
fof(f2741,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2657,f2102])).
fof(f2657,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1337,f2102,f511])).
fof(f1337,plain,(
  ( ! [X0] : (sP0(X0,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1326,f407])).
fof(f2740,plain,(
  $false | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f2658,f2102])).
fof(f2658,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | ($spl12 | $spl16)),
  inference(unit_resulting_resolution,[],[f1458,f2102,f511])).
fof(f1458,plain,(
  ( ! [X0] : (sP0(X0,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f1329,f407])).
fof(f2739,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2659,f2102])).
fof(f2659,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f513,f2102,f511])).
fof(f513,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f311,f408])).
fof(f408,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f276])).
fof(f2738,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2660,f2102])).
fof(f2660,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1338,f2102,f511])).
fof(f1338,plain,(
  ( ! [X0] : (sP0('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1326,f408])).
fof(f2737,plain,(
  $false | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f2661,f2102])).
fof(f2661,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl12 | $spl16)),
  inference(unit_resulting_resolution,[],[f1459,f2102,f511])).
fof(f1459,plain,(
  ( ! [X0] : (sP0('|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')),X0)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f1329,f408])).
fof(f2736,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2662,f2102])).
fof(f2662,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f537,f2102,f511])).
fof(f537,plain,(
  ( ! [X0] : (sP0('const.combin.I_1'(s(bool,'T_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f311,f529])).
fof(f529,plain,(
  ( ! [X0,X1] : (sP0('const.combin.I_1'(s(bool,X0)),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f408,f369])).
fof(f2735,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2667,f2102])).
fof(f2667,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f512,f2102,f511])).
fof(f512,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f311,f407])).
fof(f2734,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2668,f2102])).
fof(f2668,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f543,f2102,f511])).
fof(f2733,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2669,f2102])).
fof(f2669,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1337,f2102,f511])).
fof(f2732,plain,(
  $false | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f2670,f2102])).
fof(f2670,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl12 | $spl16)),
  inference(unit_resulting_resolution,[],[f1458,f2102,f511])).
fof(f2731,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f2672,f2102])).
fof(f2672,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1338,f2102,f511])).
fof(f2730,plain,(
  $false | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f2673,f2102])).
fof(f2673,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | ($spl12 | $spl16)),
  inference(unit_resulting_resolution,[],[f1459,f2102,f511])).
fof(f2729,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2674,f2102])).
fof(f2674,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl16),
  inference(unit_resulting_resolution,[],[f537,f2102,f511])).
fof(f2678,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f512,f2102,f2102,f511])).
fof(f2676,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f512,f509,f2102,f511])).
fof(f2666,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f512,f2102,f2102,f511])).
fof(f2664,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f513,f509,f2102,f511])).
fof(f2654,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f574,f311,f311,f2102,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f307])).
fof(f2648,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f925,f311,f2102,f470])).
fof(f2643,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f564,f311,f2102,f2102,f469])).
fof(f2638,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f925,f311,f509,f2102,f469])).
fof(f2632,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f564,f2102,f311,f2102,f469])).
fof(f2630,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f564,f509,f311,f2102,f469])).
fof(f2621,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f920,f2102,f468])).
fof(f2619,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f920,f2102,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f305])).
fof(f2617,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f524,f311,f311,f2102,f465])).
fof(f524,plain,(
  ( ! [X0,X1] : (sP5(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f311,f467])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f305])).
fof(f2610,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f525,f564,f311,f2102,f464])).
fof(f2597,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f574,f311,f2102,f459])).
fof(f2592,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f567,f2102,f2102,f458])).
fof(f2590,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f907,f509,f2102,f458])).
fof(f2585,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f567,f2102,f2102,f458])).
fof(f2583,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f567,f509,f2102,f458])).
fof(f2575,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f738,f510,f2102,f457])).
fof(f2573,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f691,f370,f2102,f457])).
fof(f2571,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f690,f369,f2102,f457])).
fof(f2566,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f736,f510,f2102,f457])).
fof(f2564,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f689,f370,f2102,f457])).
fof(f2562,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f688,f369,f2102,f457])).
fof(f2718,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f2560])).
fof(f2560,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,X0))) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl16),
  inference(unit_resulting_resolution,[],[f700,f2102,f457])).
fof(f700,plain,(
  ( ! [X2] : (sP4('T_0','const.combin.I_1'(s(bool,X2)),'const.combin.I_1'(s(bool,X2)))) )),
  inference(superposition,[],[f688,f369])).
fof(f2717,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f2556])).
fof(f2556,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl16),
  inference(unit_resulting_resolution,[],[f1217,f2102,f457])).
fof(f1217,plain,(
  sP4('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f520,f515,f518,f509,f440])).
fof(f2716,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f2554])).
fof(f2554,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl16),
  inference(unit_resulting_resolution,[],[f1184,f2102,f457])).
fof(f1184,plain,(
  sP4('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f522,f514,f517,f311,f439])).
fof(f2548,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f706,f531,f2102,f454])).
fof(f2546,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f704,f531,f2102,f454])).
fof(f2530,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f2102,f2102,f450])).
fof(f2528,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f520,f509,f2102,f2102,f450])).
fof(f2520,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f509,f2102,f450])).
fof(f2518,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f520,f509,f509,f2102,f450])).
fof(f2502,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f2102,f2102,f450])).
fof(f2500,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f520,f509,f2102,f2102,f450])).
fof(f2492,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f2102,f509,f2102,f450])).
fof(f2490,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f509,f509,f2102,f450])).
fof(f2474,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f2102,f2102,f450])).
fof(f2472,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f509,f2102,f2102,f450])).
fof(f2464,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f520,f2102,f509,f2102,f450])).
fof(f2462,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f522,f509,f509,f2102,f450])).
fof(f2446,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f847,f2102,f449])).
fof(f847,plain,(
  ~sP2('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f311,f509,f311,f446])).
fof(f2444,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f847,f2102,f448])).
fof(f2442,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f517,f311,f311,f2102,f446])).
fof(f2434,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f832,f2102,f445])).
fof(f832,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f311,f509,f311,f442])).
fof(f2432,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f832,f2102,f444])).
fof(f2430,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f695,f311,f311,f2102,f442])).
fof(f695,plain,(
  ( ! [X0] : (sP3('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f688,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP3(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f2389,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f1361,f518,f2102,f441])).
fof(f2384,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f514,f518,f2102,f441])).
fof(f2382,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f514,f518,f2102,f441])).
fof(f2377,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f515,f518,f2102,f441])).
fof(f2371,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f515,f518,f2102,f441])).
fof(f2369,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f515,f518,f2102,f441])).
fof(f2360,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f515,f1370,f2102,f441])).
fof(f2351,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f514,f1370,f2102,f441])).
fof(f2343,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f515,f517,f2102,f441])).
fof(f2341,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f515,f517,f2102,f441])).
fof(f2336,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f1361,f517,f2102,f441])).
fof(f2330,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f514,f517,f2102,f441])).
fof(f2253,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f1361,f518,f2102,f440])).
fof(f2248,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f514,f518,f2102,f440])).
fof(f2246,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f514,f518,f2102,f440])).
fof(f2241,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f515,f518,f2102,f440])).
fof(f2235,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f515,f518,f2102,f440])).
fof(f2233,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f515,f518,f2102,f440])).
fof(f2224,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f515,f1370,f2102,f440])).
fof(f2215,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f521,f1382,f514,f1370,f2102,f440])).
fof(f2207,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f515,f517,f2102,f440])).
fof(f2205,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f515,f517,f2102,f440])).
fof(f2200,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f1372,f1341,f1361,f517,f2102,f440])).
fof(f2194,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f521,f706,f514,f517,f2102,f440])).
fof(f2140,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f1184,f311,f311,f2102,f438])).
fof(f2122,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f509,f531,f2102,f398])).
fof(f2120,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f2102,f531,f2102,f398])).
fof(f2116,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f509,f531,f2102,f398])).
fof(f2114,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f2102,f531,f2102,f398])).
fof(f2110,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f531,f2102,f333])).
fof(f333,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.list.MAP_EQ_SING','thm.bool.EQ_CLAUSES|split|3')).
fof(f2108,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f369,f2102,f328])).
fof(f2106,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f311,f2102])).
fof(f2102,plain,(
  ( ! [X4] : (~p(s(bool,X4))) ) | $spl16),
  inference(cnf_transformation,[],[f2102_D])).
fof(f2102_D,plain,(
  ( ! [X4] : (~p(s(bool,X4))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2105,plain,(
  ( ! [X7] : (p(s(bool,sK18('const.combin.I_0',X7,bool))) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,sK19('const.combin.I_0',X7,bool)))))) ) | $spl20),
  inference(cnf_transformation,[],[f2105_D])).
fof(f2105_D,plain,(
  ( ! [X7] : (p(s(bool,sK18('const.combin.I_0',X7,bool))) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,sK19('const.combin.I_0',X7,bool)))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2104,plain,(
  ( ! [X6,X7] : (p(s(bool,sK18('const.combin.I_0',X7,bool))) | ~p(s(bool,X6)) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,sK19('const.combin.I_0',X7,bool)))))) )),
  inference(forward_demodulation,[],[f2093,f510])).
fof(f2093,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,sK19('const.combin.I_0',X7,bool))))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK18('const.combin.I_0',X7,bool)))))) )),
  inference(superposition,[],[f386,f510])).
fof(f386,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7] : ! [X31,X26] : (! [X8] : (p(s(bool,happ(s(fun(X7,bool),X26),s(X7,X8)))) => p(s(bool,happ(s(fun(X7,bool),X31),s(X7,X8))))) => (? [X8] : p(s(bool,happ(s(fun(X7,bool),X26),s(X7,X8)))) => ? [X8] : p(s(bool,happ(s(fun(X7,bool),X31),s(X7,X8))))))),
  file('thm.list.MAP_EQ_SING','thm.bool.MONO_EXISTS')).
fof(f2103,plain,(
  ( ! [X5] : (p(s(bool,sK18('_i_0_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK19('_i_0_0',X5,bool)))))) ) | $spl18),
  inference(cnf_transformation,[],[f2103_D])).
fof(f2103_D,plain,(
  ( ! [X5] : (p(s(bool,sK18('_i_0_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK19('_i_0_0',X5,bool)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2101,plain,(
  ( ! [X4,X5] : (p(s(bool,sK18('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK19('_i_0_0',X5,bool)))))) )),
  inference(forward_demodulation,[],[f2092,f370])).
fof(f2092,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK19('_i_0_0',X5,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK18('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f386,f370])).
fof(f1329,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl12),
  inference(cnf_transformation,[],[f1329_D])).
fof(f1329_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1330,plain,(
  ( ! [X4] : (~p(s(bool,sK17(X4,'const.combin.I_0',bool)))) ) | $spl14),
  inference(cnf_transformation,[],[f1330_D])).
fof(f1330_D,plain,(
  ( ! [X4] : (~p(s(bool,sK17(X4,'const.combin.I_0',bool)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1320,plain,(
  ( ! [X4] : (~p(s(bool,sK17(X4,'const.combin.I_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f385,f510])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK17(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7] : ! [X26,X13] : (! [X8] : (s(X7,X8) = s(X7,X13) => p(s(bool,happ(s(fun(X7,bool),X26),s(X7,X8))))) => p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X26)))))),
  file('thm.list.MAP_EQ_SING','thm.bool.JRH_INDUCT_UTIL')).
fof(f1326,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl8),
  inference(cnf_transformation,[],[f1326_D])).
fof(f1326_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1327,plain,(
  ( ! [X3] : (~p(s(bool,sK17(X3,'_i_0_0',bool)))) ) | $spl10),
  inference(cnf_transformation,[],[f1327_D])).
fof(f1327_D,plain,(
  ( ! [X3] : (~p(s(bool,sK17(X3,'_i_0_0',bool)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1319,plain,(
  ( ! [X3] : (~p(s(bool,sK17(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f385,f370])).
fof(f502,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | $spl0),
  inference(cnf_transformation,[],[f502_D])).
fof(f502_D,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f508,plain,(
  s('type.list.list'(sK8),sK10) = s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))) | $spl6),
  inference(cnf_transformation,[],[f508_D])).
fof(f508_D,plain,(
  s('type.list.list'(sK8),sK10) = s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f308,plain,(
  s('type.list.list'(sK8),sK10) = s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))) | s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10)))),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  (s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | (s('type.list.list'(sK8),sK10) = s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,sK12),s('type.list.list'(sK8),'const.list.NIL_0'))) & s(sK7,sK9) = s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK12))))) & (s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) != s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | ! [X5] : (s('type.list.list'(sK8),sK10) != s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,X5),s('type.list.list'(sK8),'const.list.NIL_0'))) | s(sK7,sK9) != s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,X5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9,sK10,sK11,sK12])],[f231])).
fof(f231,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X3))) | ? [X6] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),'const.list.NIL_0'))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))) & (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.NIL_0'))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X3))) | ! [X5] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),'const.list.NIL_0'))) | s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(rectify,[],[f230])).
fof(f230,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X3))) | ? [X5] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),'const.list.NIL_0'))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))) & (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.NIL_0'))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X3))) | ! [X5] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),'const.list.NIL_0'))) | s(X0,X2) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ? [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X3))) <~> ? [X5] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),'const.list.NIL_0'))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X3))) <=> ? [X5] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),'const.list.NIL_0'))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X3))) <=> ? [X5] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),'const.list.NIL_0'))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),'const.list.NIL_0'))) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X10),s('type.list.list'(X7),X9))) <=> ? [X11] : (s('type.list.list'(X7),X9) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X11),s('type.list.list'(X7),'const.list.NIL_0'))) & s(X6,X8) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X11)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),'const.list.NIL_0'))) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X10),s('type.list.list'(X7),X9))) <=> ? [X11] : (s('type.list.list'(X7),X9) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X11),s('type.list.list'(X7),'const.list.NIL_0'))) & s(X6,X8) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X11)))))),
  file('thm.list.MAP_EQ_SING','thm.list.MAP_EQ_SING')).
fof(f506,plain,(
  s(sK7,sK9) = s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK12))) | $spl4),
  inference(cnf_transformation,[],[f506_D])).
fof(f506_D,plain,(
  s(sK7,sK9) = s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK12))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f309,plain,(
  s(sK7,sK9) = s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,sK12))) | s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) = s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10)))),
  inference(cnf_transformation,[],[f232])).
fof(f503,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) != s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) | $spl1),
  inference(cnf_transformation,[],[f503_D])).
fof(f503_D,plain,(
  s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) != s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f504,plain,(
  ( ! [X5] : (s('type.list.list'(sK8),sK10) != s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,X5),s('type.list.list'(sK8),'const.list.NIL_0'))) | s(sK7,sK9) != s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,X5)))) ) | $spl2),
  inference(cnf_transformation,[],[f504_D])).
fof(f504_D,plain,(
  ( ! [X5] : (s('type.list.list'(sK8),sK10) != s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,X5),s('type.list.list'(sK8),'const.list.NIL_0'))) | s(sK7,sK9) != s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,X5)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f310,plain,(
  ( ! [X5] : (s(sK7,sK9) != s(sK7,happ(s(fun(sK8,sK7),sK11),s(sK8,X5))) | s('type.list.list'(sK8),sK10) != s('type.list.list'(sK8),'const.list.CONS_2'(s(sK8,X5),s('type.list.list'(sK8),'const.list.NIL_0'))) | s('type.list.list'(sK7),'const.list.CONS_2'(s(sK7,sK9),s('type.list.list'(sK7),'const.list.NIL_0'))) != s('type.list.list'(sK7),'const.list.MAP_2'(s(fun(sK8,sK7),sK11),s('type.list.list'(sK8),sK10)))) )),
  inference(cnf_transformation,[],[f232])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 221
% Passive clauses: 1116
% Generated clauses: 8403
% Final active clauses: 211
% Final passive clauses: 828
% Input formulas: 69
% Initial clauses: 184
% 
% Duplicate literals: 8
% Trivial inequalities: 9
% Fw subsumption resolutions: 161
% Fw demodulations: 3570
% Bw demodulations: 25
% Condensations: 7
% 
% Simple tautologies: 69
% Equational tautologies: 33
% Forward subsumptions: 2860
% Fw demodulations to eq. taut.: 330
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 2931
% Factoring: 1
% Forward superposition: 758
% Backward superposition: 668
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 3
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 1365
% SAT solver unit clauses: 504
% SAT solver binary clauses: 626
% 
% Sat splits: 7
% Sat splitting refutations: 233
% 
% Memory used [KB]: 11257
% Time elapsed: 0.325 s
% ------------------------------
% ------------------------------
% Success in time 1.65 s
